<!DOCTYPE html>
<html>
<head>
    <meta charset="UTF-8">
    <title>E-Prover Proof Graphs Grid</title>
    <script type="text/javascript" src="https://unpkg.com/vis-network/standalone/umd/vis-network.min.js"></script>
    <style>
        * {
            margin: 0;
            padding: 0;
            box-sizing: border-box;
        }
        body {
            font-family: 'Courier New', 'Liberation Mono', 'DejaVu Sans Mono', monospace;
            background-color: #FFFFFF;
            padding: 0;
            margin: 0;
            overflow: hidden;
        }
        .nav-bar {
            position: fixed;
            top: 0;
            left: 0;
            right: 0;
            height: 50px;
            background-color: #000000;
            color: #FFFFFF;
            display: flex;
            align-items: center;
            justify-content: space-between;
            padding: 0 20px;
            z-index: 1000;
            border-bottom: 4px solid #000000;
        }
        .legend-box {
            position: fixed;
            top: 50px;
            left: 50%;
            transform: translateX(-50%);
            background-color: #FFFFFF;
            border: 3px solid #000000;
            border-top: none;
            padding: 6px 12px;
            font-size: 8px;
            line-height: 1.2;
            z-index: 999;
            display: flex;
            align-items: center;
            gap: 12px;
            white-space: nowrap;
        }
        .legend-title {
            font-weight: bold;
            margin-right: 4px;
        }
        .nav-button {
            background-color: #FFFFFF;
            color: #000000;
            border: 2px solid #000000;
            padding: 8px 20px;
            font-family: inherit;
            font-size: 12px;
            cursor: pointer;
            font-weight: bold;
        }
        .nav-button:hover {
            background-color: #000000;
            color: #FFFFFF;
        }
        .nav-button:disabled {
            background-color: #CCCCCC;
            color: #666666;
            cursor: not-allowed;
        }
        .page-info {
            font-size: 11px;
            font-weight: bold;
        }
        .grid-container {
            display: grid;
            grid-template-columns: repeat(5, 1fr);
            grid-template-rows: repeat(5, 1fr);
            width: 100vw;
            height: calc(100vh - 80px);
            margin-top: 80px;
            gap: 0;
        }
        .grid-item {
            border: 4px solid #000000;
            background-color: #FFFFFF;
            overflow: hidden;
            position: relative;
        }
        .grid-label {
            position: absolute;
            top: 0;
            left: 0;
            background-color: #FFFFFF;
            color: #000000;
            padding: 3px 5px;
            font-size: 7px;
            font-weight: bold;
            z-index: 10;
            border-right: 2px solid #000000;
            border-bottom: 2px solid #000000;
            max-width: calc(100% - 10px);
            overflow: hidden;
            text-overflow: ellipsis;
            white-space: nowrap;
            line-height: 1.1;
        }
        .grid-network {
            width: 100%;
            height: 100%;
        }
        .node-detail-box {
            position: fixed;
            bottom: 10px;
            left: 50%;
            transform: translateX(-50%);
            background-color: #FFFFFF;
            border: 3px solid #000000;
            padding: 8px 12px;
            font-size: 9px;
            font-family: inherit;
            z-index: 1001;
            max-width: 90%;
            display: none;
            white-space: pre-wrap;
        }
        .node-detail-box.active {
            display: block;
        }
    </style>
</head>
<body>
    <div class="nav-bar">
        <button class="nav-button" id="prev-btn" onclick="changePage(-1)">◀ PREV</button>
        <div class="page-info" id="page-info">Page 1 / 1</div>
        <button class="nav-button" id="next-btn" onclick="changePage(1)">NEXT ▶</button>
    </div>
    <div class="legend-box">
        <span class="legend-title">LEGEND:</span>
        <span>White thick-border: initial</span>
        <span>Black: $false</span>
        <span>Edge label: child inference rule</span>
        <span>Click node: show (name, role, rule, formula)</span>
    </div>
    <div class="grid-container">
        
        <div class="grid-item" id="grid-item-0">
            <div class="grid-label" id="label-0"></div>
            <div class="grid-network" id="network-0"></div>
        </div>
        <div class="grid-item" id="grid-item-1">
            <div class="grid-label" id="label-1"></div>
            <div class="grid-network" id="network-1"></div>
        </div>
        <div class="grid-item" id="grid-item-2">
            <div class="grid-label" id="label-2"></div>
            <div class="grid-network" id="network-2"></div>
        </div>
        <div class="grid-item" id="grid-item-3">
            <div class="grid-label" id="label-3"></div>
            <div class="grid-network" id="network-3"></div>
        </div>
        <div class="grid-item" id="grid-item-4">
            <div class="grid-label" id="label-4"></div>
            <div class="grid-network" id="network-4"></div>
        </div>
        <div class="grid-item" id="grid-item-5">
            <div class="grid-label" id="label-5"></div>
            <div class="grid-network" id="network-5"></div>
        </div>
        <div class="grid-item" id="grid-item-6">
            <div class="grid-label" id="label-6"></div>
            <div class="grid-network" id="network-6"></div>
        </div>
        <div class="grid-item" id="grid-item-7">
            <div class="grid-label" id="label-7"></div>
            <div class="grid-network" id="network-7"></div>
        </div>
        <div class="grid-item" id="grid-item-8">
            <div class="grid-label" id="label-8"></div>
            <div class="grid-network" id="network-8"></div>
        </div>
        <div class="grid-item" id="grid-item-9">
            <div class="grid-label" id="label-9"></div>
            <div class="grid-network" id="network-9"></div>
        </div>
        <div class="grid-item" id="grid-item-10">
            <div class="grid-label" id="label-10"></div>
            <div class="grid-network" id="network-10"></div>
        </div>
        <div class="grid-item" id="grid-item-11">
            <div class="grid-label" id="label-11"></div>
            <div class="grid-network" id="network-11"></div>
        </div>
        <div class="grid-item" id="grid-item-12">
            <div class="grid-label" id="label-12"></div>
            <div class="grid-network" id="network-12"></div>
        </div>
        <div class="grid-item" id="grid-item-13">
            <div class="grid-label" id="label-13"></div>
            <div class="grid-network" id="network-13"></div>
        </div>
        <div class="grid-item" id="grid-item-14">
            <div class="grid-label" id="label-14"></div>
            <div class="grid-network" id="network-14"></div>
        </div>
        <div class="grid-item" id="grid-item-15">
            <div class="grid-label" id="label-15"></div>
            <div class="grid-network" id="network-15"></div>
        </div>
        <div class="grid-item" id="grid-item-16">
            <div class="grid-label" id="label-16"></div>
            <div class="grid-network" id="network-16"></div>
        </div>
        <div class="grid-item" id="grid-item-17">
            <div class="grid-label" id="label-17"></div>
            <div class="grid-network" id="network-17"></div>
        </div>
        <div class="grid-item" id="grid-item-18">
            <div class="grid-label" id="label-18"></div>
            <div class="grid-network" id="network-18"></div>
        </div>
        <div class="grid-item" id="grid-item-19">
            <div class="grid-label" id="label-19"></div>
            <div class="grid-network" id="network-19"></div>
        </div>
        <div class="grid-item" id="grid-item-20">
            <div class="grid-label" id="label-20"></div>
            <div class="grid-network" id="network-20"></div>
        </div>
        <div class="grid-item" id="grid-item-21">
            <div class="grid-label" id="label-21"></div>
            <div class="grid-network" id="network-21"></div>
        </div>
        <div class="grid-item" id="grid-item-22">
            <div class="grid-label" id="label-22"></div>
            <div class="grid-network" id="network-22"></div>
        </div>
        <div class="grid-item" id="grid-item-23">
            <div class="grid-label" id="label-23"></div>
            <div class="grid-network" id="network-23"></div>
        </div>
        <div class="grid-item" id="grid-item-24">
            <div class="grid-label" id="label-24"></div>
            <div class="grid-network" id="network-24"></div>
        </div>
    </div>
    <div class="node-detail-box" id="node-detail-box"></div>

    <script type="text/javascript">
        const allGraphsData = [{"problem_file": "PUZ012_1.p", "status": "# SZS status Theorem", "nodes": [{"id": 0, "label": "0", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": null, "raw_label": "tff(label_is_wrong, axiom, ![X3:box, X4:fruit]:(~((label(X3,X4)&contains(X3,X4)))),\\nfile('Problems/PUZ/PUZ012_1.p', label_is_wrong)).\n", "fillcolor": "forestgreen"}, {"id": 1, "label": "1", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": null, "raw_label": "tff(prove_boxa_contains_bananas_and_boxc_oranges, conjecture, (contains(boxa,bananas)&contains(boxc,oranges)),\\nfile('Problems/PUZ/PUZ012_1.p', prove_boxa_contains_bananas_and_boxc_oranges)).\n", "fillcolor": "firebrick1"}, {"id": 2, "label": "2", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": null, "raw_label": "tff(boxc_labelled_bananas, hypothesis, label(boxc,bananas),\\nfile('Problems/PUZ/PUZ012_1.p', boxc_labelled_bananas)).\n", "fillcolor": "forestgreen"}, {"id": 3, "label": "3", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": null, "raw_label": "tff(each_box_contains_something, axiom, ![X5:box]:((contains(X5,apples)|contains(X5,bananas)|contains(X5,oranges))),\\nfile('Problems/PUZ/PUZ012_1.p', each_box_contains_something)).\n", "fillcolor": "forestgreen"}, {"id": 4, "label": "4", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": null, "raw_label": "tff(each_thing_is_in_a_box, axiom, ![X1:fruit]:((contains(boxa,X1)|contains(boxb,X1)|contains(boxc,X1))),\\nfile('Problems/PUZ/PUZ012_1.p', each_thing_is_in_a_box)).\n", "fillcolor": "forestgreen"}, {"id": 5, "label": "5", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": null, "raw_label": "tff(boxb_labelled_oranges, hypothesis, label(boxb,oranges),\\nfile('Problems/PUZ/PUZ012_1.p', boxb_labelled_oranges)).\n", "fillcolor": "forestgreen"}, {"id": 6, "label": "6", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": null, "raw_label": "tff(boxa_not_boxc, axiom, ~(equal_boxes(boxa,boxc)),\\nfile('Problems/PUZ/PUZ012_1.p', boxa_not_boxc)).\n", "fillcolor": "forestgreen"}, {"id": 7, "label": "7", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": null, "raw_label": "tff(contains_is_well_defined2, axiom, ![X8:box, X4:fruit, X9:box]:(((contains(X8,X4)&contains(X9,X4))=>equal_boxes(X8,X9))),\\nfile('Problems/PUZ/PUZ012_1.p', contains_is_well_defined2)).\n", "fillcolor": "forestgreen"}, {"id": 8, "label": "8", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": null, "raw_label": "tff(boxa_labelled_apples, hypothesis, label(boxa,apples),\\nfile('Problems/PUZ/PUZ012_1.p', boxa_labelled_apples)).\n", "fillcolor": "forestgreen"}, {"id": 9, "label": "9", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": null, "raw_label": "tff(boxb_not_boxc, axiom, ~(equal_boxes(boxb,boxc)),\\nfile('Problems/PUZ/PUZ012_1.p', boxb_not_boxc)).\n", "fillcolor": "forestgreen"}, {"id": 10, "label": "10", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": null, "raw_label": "tff(boxb_contains_apples, hypothesis, contains(boxb,apples),\\nfile('Problems/PUZ/PUZ012_1.p', boxb_contains_apples)).\n", "fillcolor": "forestgreen"}, {"id": 11, "label": "11", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "fof_nnf", "raw_label": "tff(c_0_11, plain, ![X12:box, X13:fruit]:((~label(X12,X13)|~contains(X12,X13))),\\ninference(fof_nnf,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[label_is_wrong])])])).\n", "fillcolor": "palegreen"}, {"id": 12, "label": "12", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "assume_negation", "raw_label": "tff(c_0_12, negated_conjecture, ~((contains(boxa,bananas)&contains(boxc,oranges))),\\ninference(assume_negation,[status(cth)],[prove_boxa_contains_bananas_and_boxc_oranges])).\n", "fillcolor": "lightskyblue1"}, {"id": 13, "label": "13", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "split_conjunct", "raw_label": "tcf(c_0_13, plain, ![X2:box, X1:fruit]:((~label(X2,X1)|~contains(X2,X1))),\\ninference(split_conjunct,[status(thm)],[c_0_11])).\n", "fillcolor": "palegreen"}, {"id": 14, "label": "14", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "split_conjunct", "raw_label": "tcf(c_0_14, hypothesis, label(boxc,bananas),\\ninference(split_conjunct,[status(thm)],[boxc_labelled_bananas])).\n", "fillcolor": "palegreen"}, {"id": 15, "label": "15", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "variable_rename", "raw_label": "tff(c_0_15, plain, ![X15:box]:((contains(X15,apples)|contains(X15,bananas)|contains(X15,oranges))),\\ninference(variable_rename,[status(thm)],[each_box_contains_something])).\n", "fillcolor": "palegreen"}, {"id": 16, "label": "16", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "fof_nnf", "raw_label": "tff(c_0_16, negated_conjecture, (~contains(boxa,bananas)|~contains(boxc,oranges)),\\ninference(fof_nnf,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])).\n", "fillcolor": "lightskyblue1"}, {"id": 17, "label": "17", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "variable_rename", "raw_label": "tff(c_0_17, plain, ![X14:fruit]:((contains(boxa,X14)|contains(boxb,X14)|contains(boxc,X14))),\\ninference(variable_rename,[status(thm)],[each_thing_is_in_a_box])).\n", "fillcolor": "palegreen"}, {"id": 18, "label": "18", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "split_conjunct", "raw_label": "tcf(c_0_18, hypothesis, label(boxb,oranges),\\ninference(split_conjunct,[status(thm)],[boxb_labelled_oranges])).\n", "fillcolor": "palegreen"}, {"id": 19, "label": "19", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "fof_simplification", "raw_label": "tff(c_0_19, plain, ~equal_boxes(boxa,boxc),\\ninference(fof_simplification,[status(thm)],[boxa_not_boxc])).\n", "fillcolor": "palegreen"}, {"id": 20, "label": "20", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "fof_nnf", "raw_label": "tff(c_0_20, plain, ![X19:box, X20:fruit, X21:box]:((~contains(X19,X20)|~contains(X21,X20)|equal_boxes(X19,X21))),\\ninference(fof_nnf,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[contains_is_well_defined2])])])).\n", "fillcolor": "palegreen"}, {"id": 21, "label": "21", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "spm", "raw_label": "tcf(c_0_21, hypothesis, ~contains(boxc,bananas),\\ninference(spm,[status(thm)],[c_0_13, c_0_14])).\n", "fillcolor": "palegreen"}, {"id": 22, "label": "22", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "split_conjunct", "raw_label": "tcf(c_0_22, plain, ![X2:box]:((contains(X2,apples)|contains(X2,bananas)|contains(X2,oranges))),\\ninference(split_conjunct,[status(thm)],[c_0_15])).\n", "fillcolor": "palegreen"}, {"id": 23, "label": "23", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "split_conjunct", "raw_label": "tcf(c_0_23, negated_conjecture, (~contains(boxa,bananas)|~contains(boxc,oranges)),\\ninference(split_conjunct,[status(thm)],[c_0_16])).\n", "fillcolor": "lightskyblue1"}, {"id": 24, "label": "24", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "split_conjunct", "raw_label": "tcf(c_0_24, plain, ![X1:fruit]:((contains(boxa,X1)|contains(boxb,X1)|contains(boxc,X1))),\\ninference(split_conjunct,[status(thm)],[c_0_17])).\n", "fillcolor": "palegreen"}, {"id": 25, "label": "25", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "spm", "raw_label": "tcf(c_0_25, hypothesis, ~contains(boxb,oranges),\\ninference(spm,[status(thm)],[c_0_13, c_0_18])).\n", "fillcolor": "palegreen"}, {"id": 26, "label": "26", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "split_conjunct", "raw_label": "tcf(c_0_26, hypothesis, label(boxa,apples),\\ninference(split_conjunct,[status(thm)],[boxa_labelled_apples])).\n", "fillcolor": "palegreen"}, {"id": 27, "label": "27", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "fof_nnf", "raw_label": "tff(c_0_27, plain, ~equal_boxes(boxa,boxc),\\ninference(fof_nnf,[status(thm)],[c_0_19])).\n", "fillcolor": "palegreen"}, {"id": 28, "label": "28", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "split_conjunct", "raw_label": "tcf(c_0_28, plain, ![X3:box, X2:box, X1:fruit]:((equal_boxes(X2,X3)|~contains(X2,X1)|~contains(X3,X1))),\\ninference(split_conjunct,[status(thm)],[c_0_20])).\n", "fillcolor": "palegreen"}, {"id": 29, "label": "29", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "spm", "raw_label": "tcf(c_0_29, hypothesis, (contains(boxc,oranges)|contains(boxc,apples)),\\ninference(spm,[status(thm)],[c_0_21, c_0_22])).\n", "fillcolor": "palegreen"}, {"id": 30, "label": "30", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "sr", "raw_label": "tcf(c_0_30, negated_conjecture, (contains(boxa,oranges)|~contains(boxa,bananas)),\\ninference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_25])).\n", "fillcolor": "lightskyblue1"}, {"id": 31, "label": "31", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "spm", "raw_label": "tcf(c_0_31, hypothesis, ~contains(boxa,apples),\\ninference(spm,[status(thm)],[c_0_13, c_0_26])).\n", "fillcolor": "palegreen"}, {"id": 32, "label": "32", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "fof_simplification", "raw_label": "tff(c_0_32, plain, ~equal_boxes(boxb,boxc),\\ninference(fof_simplification,[status(thm)],[boxb_not_boxc])).\n", "fillcolor": "palegreen"}, {"id": 33, "label": "33", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "split_conjunct", "raw_label": "tcf(c_0_33, plain, ~equal_boxes(boxa,boxc),\\ninference(split_conjunct,[status(thm)],[c_0_27])).\n", "fillcolor": "palegreen"}, {"id": 34, "label": "34", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "spm", "raw_label": "tcf(c_0_34, hypothesis, ![X2:box]:((contains(boxc,apples)|equal_boxes(X2,boxc)|~contains(X2,oranges))),\\ninference(spm,[status(thm)],[c_0_28, c_0_29])).\n", "fillcolor": "palegreen"}, {"id": 35, "label": "35", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "sr", "raw_label": "tcf(c_0_35, negated_conjecture, contains(boxa,oranges),\\ninference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_22]), c_0_31])).\n", "fillcolor": "lightskyblue1"}, {"id": 36, "label": "36", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "fof_nnf", "raw_label": "tff(c_0_36, plain, ~equal_boxes(boxb,boxc),\\ninference(fof_nnf,[status(thm)],[c_0_32])).\n", "fillcolor": "palegreen"}, {"id": 37, "label": "37", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "cn", "raw_label": "tcf(c_0_37, hypothesis, contains(boxc,apples),\\ninference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_35])])).\n", "fillcolor": "palegreen"}, {"id": 38, "label": "38", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "split_conjunct", "raw_label": "tcf(c_0_38, plain, ~equal_boxes(boxb,boxc),\\ninference(split_conjunct,[status(thm)],[c_0_36])).\n", "fillcolor": "palegreen"}, {"id": 39, "label": "39", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "spm", "raw_label": "tcf(c_0_39, hypothesis, ![X2:box]:((equal_boxes(X2,boxc)|~contains(X2,apples))),\\ninference(spm,[status(thm)],[c_0_28, c_0_37])).\n", "fillcolor": "palegreen"}, {"id": 40, "label": "40", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "split_conjunct", "raw_label": "tcf(c_0_40, hypothesis, contains(boxb,apples),\\ninference(split_conjunct,[status(thm)],[boxb_contains_apples])).\n", "fillcolor": "palegreen"}, {"id": 41, "label": "c_0_41", "color": {"background": "#000000", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#FFFFFF"}, "cnf_name": "c_0_41", "role": "hypothesis", "formula": "$false", "inference_rule": "cn", "raw_label": "cnf(c_0_41, hypothesis, ($false),\\ninference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_40])])).\n", "fillcolor": "darkorchid1"}], "edges": [{"from": 0, "to": 11, "label": "fof_nnf", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 1, "to": 12, "label": "assume_negation", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 11, "to": 13, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 2, "to": 14, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 3, "to": 15, "label": "variable_rename", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 12, "to": 16, "label": "fof_nnf", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 4, "to": 17, "label": "variable_rename", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 5, "to": 18, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 6, "to": 19, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 7, "to": 20, "label": "fof_nnf", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 14, "to": 21, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 13, "to": 21, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 15, "to": 22, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 16, "to": 23, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 17, "to": 24, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 18, "to": 25, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 13, "to": 25, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 8, "to": 26, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 19, "to": 27, "label": "fof_nnf", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 20, "to": 28, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 22, "to": 29, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 21, "to": 29, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 25, "to": 30, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 24, "to": 30, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 23, "to": 30, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 26, "to": 31, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 13, "to": 31, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 9, "to": 32, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 27, "to": 33, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 29, "to": 34, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 28, "to": 34, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 31, "to": 35, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 22, "to": 35, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 30, "to": 35, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 32, "to": 36, "label": "fof_nnf", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 35, "to": 37, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 34, "to": 37, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 33, "to": 37, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 36, "to": 38, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 37, "to": 39, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 28, "to": 39, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 10, "to": 40, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 40, "to": 41, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 39, "to": 41, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 38, "to": 41, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}]}, {"problem_file": "PUZ031+2.p", "status": "# SZS status Theorem", "nodes": [{"id": 0, "label": "0", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": null, "raw_label": "fof(pel47_13, axiom, ![X2, X3]:(((bird(X2)&snail(X3))=>~(eats(X2,X3)))),\\nfile('Problems/PUZ/PUZ031+2.p', pel47_13)).\n", "fillcolor": "forestgreen"}, {"id": 1, "label": "1", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": null, "raw_label": "fof(pel47_7, axiom, ![X2]:((animal(X2)=>(![X3]:((plant(X3)=>eats(X2,X3)))|![X4]:((((animal(X4)&much_smaller(X4,X2))&?[X5]:((plant(X5)&eats(X4,X5))))=>eats(X2,X4)))))),\\nfile('Problems/PUZ/PUZ031+2.p', pel47_7)).\n", "fillcolor": "forestgreen"}, {"id": 2, "label": "2", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": null, "raw_label": "fof(pel47_8, axiom, ![X3, X2]:(((bird(X3)&snail(X2))=>much_smaller(X2,X3))),\\nfile('Problems/PUZ/PUZ031+2.p', pel47_8)).\n", "fillcolor": "forestgreen"}, {"id": 3, "label": "3", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": null, "raw_label": "fof(pel47_3_1, axiom, ![X2]:((bird(X2)=>animal(X2))),\\nfile('Problems/PUZ/PUZ031+2.p', pel47_3_1)).\n", "fillcolor": "forestgreen"}, {"id": 4, "label": "4", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": null, "raw_label": "fof(pel47_4_2, axiom, ![X2]:((snail(X2)=>animal(X2))),\\nfile('Problems/PUZ/PUZ031+2.p', pel47_4_2)).\n", "fillcolor": "forestgreen"}, {"id": 5, "label": "5", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": null, "raw_label": "fof(pel47_14a, axiom, ![X2]:((snail(X2)=>?[X3]:((plant(X3)&eats(X2,X3))))),\\nfile('Problems/PUZ/PUZ031+2.p', pel47_14a)).\n", "fillcolor": "forestgreen"}, {"id": 6, "label": "6", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": null, "raw_label": "fof(pel47_6_2, axiom, ![X2]:((grain(X2)=>plant(X2))),\\nfile('Problems/PUZ/PUZ031+2.p', pel47_6_2)).\n", "fillcolor": "forestgreen"}, {"id": 7, "label": "7", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": null, "raw_label": "fof(grain_type, axiom, ?[X1]:(grain(X1)),\\nfile('Problems/PUZ/PUZ031+2.p', grain_type)).\n", "fillcolor": "forestgreen"}, {"id": 8, "label": "8", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": null, "raw_label": "fof(snail_type, axiom, ?[X1]:(snail(X1)),\\nfile('Problems/PUZ/PUZ031+2.p', snail_type)).\n", "fillcolor": "forestgreen"}, {"id": 9, "label": "9", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": null, "raw_label": "fof(pel47_9, axiom, ![X2, X3]:(((bird(X2)&fox(X3))=>much_smaller(X2,X3))),\\nfile('Problems/PUZ/PUZ031+2.p', pel47_9)).\n", "fillcolor": "forestgreen"}, {"id": 10, "label": "10", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": null, "raw_label": "fof(pel47_2_1, axiom, ![X2]:((fox(X2)=>animal(X2))),\\nfile('Problems/PUZ/PUZ031+2.p', pel47_2_1)).\n", "fillcolor": "forestgreen"}, {"id": 11, "label": "11", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": null, "raw_label": "fof(bird_type, axiom, ?[X1]:(bird(X1)),\\nfile('Problems/PUZ/PUZ031+2.p', bird_type)).\n", "fillcolor": "forestgreen"}, {"id": 12, "label": "12", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": null, "raw_label": "fof(pel47, conjecture, ?[X2, X3]:(((animal(X2)&animal(X3))&?[X5]:(((grain(X5)&eats(X3,X5))&eats(X2,X3))))),\\nfile('Problems/PUZ/PUZ031+2.p', pel47)).\n", "fillcolor": "firebrick1"}, {"id": 13, "label": "13", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": null, "raw_label": "fof(fox_type, axiom, ?[X1]:(fox(X1)),\\nfile('Problems/PUZ/PUZ031+2.p', fox_type)).\n", "fillcolor": "forestgreen"}, {"id": 14, "label": "14", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": null, "raw_label": "fof(pel47_11, axiom, ![X2, X3]:(((wolf(X2)&fox(X3))=>~(eats(X2,X3)))),\\nfile('Problems/PUZ/PUZ031+2.p', pel47_11)).\n", "fillcolor": "forestgreen"}, {"id": 15, "label": "15", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": null, "raw_label": "fof(pel47_10, axiom, ![X2, X3]:(((fox(X2)&wolf(X3))=>much_smaller(X2,X3))),\\nfile('Problems/PUZ/PUZ031+2.p', pel47_10)).\n", "fillcolor": "forestgreen"}, {"id": 16, "label": "16", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": null, "raw_label": "fof(pel47_1_1, axiom, ![X2]:((wolf(X2)=>animal(X2))),\\nfile('Problems/PUZ/PUZ031+2.p', pel47_1_1)).\n", "fillcolor": "forestgreen"}, {"id": 17, "label": "17", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": null, "raw_label": "fof(pel47_11a, axiom, ![X2, X3]:(((wolf(X2)&grain(X3))=>~(eats(X2,X3)))),\\nfile('Problems/PUZ/PUZ031+2.p', pel47_11a)).\n", "fillcolor": "forestgreen"}, {"id": 18, "label": "18", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": null, "raw_label": "fof(wolf_type, axiom, ?[X1]:(wolf(X1)),\\nfile('Problems/PUZ/PUZ031+2.p', wolf_type)).\n", "fillcolor": "forestgreen"}, {"id": 19, "label": "19", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "fof_simplification", "raw_label": "fof(c_0_19, plain, ![X2, X3]:(((bird(X2)&snail(X3))=>~eats(X2,X3))),\\ninference(fof_simplification,[status(thm)],[pel47_13])).\n", "fillcolor": "palegreen"}, {"id": 20, "label": "20", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "fof_nnf", "raw_label": "fof(c_0_20, plain, ![X18, X19, X20, X21]:((~animal(X18)|(~plant(X19)|eats(X18,X19)|(~animal(X20)|~much_smaller(X20,X18)|(~plant(X21)|~eats(X20,X21))|eats(X18,X20))))),\\ninference(fof_nnf,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[pel47_7])])])])).\n", "fillcolor": "palegreen"}, {"id": 21, "label": "21", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "fof_nnf", "raw_label": "fof(c_0_21, plain, ![X22, X23]:((~bird(X22)|~snail(X23)|much_smaller(X23,X22))),\\ninference(fof_nnf,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[pel47_8])])])).\n", "fillcolor": "palegreen"}, {"id": 22, "label": "22", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "fof_nnf", "raw_label": "fof(c_0_22, plain, ![X13]:((~bird(X13)|animal(X13))),\\ninference(fof_nnf,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[pel47_3_1])])])).\n", "fillcolor": "palegreen"}, {"id": 23, "label": "23", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "fof_nnf", "raw_label": "fof(c_0_23, plain, ![X15]:((~snail(X15)|animal(X15))),\\ninference(fof_nnf,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[pel47_4_2])])])).\n", "fillcolor": "palegreen"}, {"id": 24, "label": "24", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "fof_nnf", "raw_label": "fof(c_0_24, plain, ![X36, X37]:((~bird(X36)|~snail(X37)|~eats(X36,X37))),\\ninference(fof_nnf,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])])).\n", "fillcolor": "palegreen"}, {"id": 25, "label": "c_0_25", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_25", "role": "plain", "formula": "eats(X1,X2)|eats(X1,X3)|~animal(X1)|~plant(X2)|~animal(X3)|~much_smaller(X3,X1)|~plant(X4)|~eats(X3,X4)", "inference_rule": "split_conjunct", "raw_label": "cnf(c_0_25, plain, (eats(X1,X2)|eats(X1,X3)|~animal(X1)|~plant(X2)|~animal(X3)|~much_smaller(X3,X1)|~plant(X4)|~eats(X3,X4)),\\ninference(split_conjunct,[status(thm)],[c_0_20])).\n", "fillcolor": "palegreen"}, {"id": 26, "label": "c_0_26", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_26", "role": "plain", "formula": "much_smaller(X2,X1)|~bird(X1)|~snail(X2)", "inference_rule": "split_conjunct", "raw_label": "cnf(c_0_26, plain, (much_smaller(X2,X1)|~bird(X1)|~snail(X2)),\\ninference(split_conjunct,[status(thm)],[c_0_21])).\n", "fillcolor": "palegreen"}, {"id": 27, "label": "c_0_27", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_27", "role": "plain", "formula": "animal(X1)|~bird(X1)", "inference_rule": "split_conjunct", "raw_label": "cnf(c_0_27, plain, (animal(X1)|~bird(X1)),\\ninference(split_conjunct,[status(thm)],[c_0_22])).\n", "fillcolor": "palegreen"}, {"id": 28, "label": "c_0_28", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_28", "role": "plain", "formula": "animal(X1)|~snail(X1)", "inference_rule": "split_conjunct", "raw_label": "cnf(c_0_28, plain, (animal(X1)|~snail(X1)),\\ninference(split_conjunct,[status(thm)],[c_0_23])).\n", "fillcolor": "palegreen"}, {"id": 29, "label": "c_0_29", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_29", "role": "plain", "formula": "~bird(X1)|~snail(X2)|~eats(X1,X2)", "inference_rule": "split_conjunct", "raw_label": "cnf(c_0_29, plain, (~bird(X1)|~snail(X2)|~eats(X1,X2)),\\ninference(split_conjunct,[status(thm)],[c_0_24])).\n", "fillcolor": "palegreen"}, {"id": 30, "label": "30", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "distribute", "raw_label": "fof(c_0_30, plain, ![X40]:(((plant(esk8_1(X40))|~snail(X40))&(eats(X40,esk8_1(X40))|~snail(X40)))),\\ninference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[pel47_14a])])])])])).\n", "fillcolor": "palegreen"}, {"id": 31, "label": "31", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "fof_nnf", "raw_label": "fof(c_0_31, plain, ![X17]:((~grain(X17)|plant(X17))),\\ninference(fof_nnf,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[pel47_6_2])])])).\n", "fillcolor": "palegreen"}, {"id": 32, "label": "32", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "skolemize", "raw_label": "fof(c_0_32, plain, grain(esk6_0),\\ninference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[grain_type])])).\n", "fillcolor": "palegreen"}, {"id": 33, "label": "c_0_33", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_33", "role": "plain", "formula": "eats(X1,X2)|~eats(X3,X4)|~plant(X4)|~plant(X2)|~snail(X3)|~bird(X1)", "inference_rule": "csr", "raw_label": "cnf(c_0_33, plain, (eats(X1,X2)|~eats(X3,X4)|~plant(X4)|~plant(X2)|~snail(X3)|~bird(X1)),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_27]), c_0_28]), c_0_29])).\n", "fillcolor": "palegreen"}, {"id": 34, "label": "c_0_34", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_34", "role": "plain", "formula": "eats(X1,esk8_1(X1))|~snail(X1)", "inference_rule": "split_conjunct", "raw_label": "cnf(c_0_34, plain, (eats(X1,esk8_1(X1))|~snail(X1)),\\ninference(split_conjunct,[status(thm)],[c_0_30])).\n", "fillcolor": "palegreen"}, {"id": 35, "label": "c_0_35", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_35", "role": "plain", "formula": "plant(esk8_1(X1))|~snail(X1)", "inference_rule": "split_conjunct", "raw_label": "cnf(c_0_35, plain, (plant(esk8_1(X1))|~snail(X1)),\\ninference(split_conjunct,[status(thm)],[c_0_30])).\n", "fillcolor": "palegreen"}, {"id": 36, "label": "c_0_36", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_36", "role": "plain", "formula": "plant(X1)|~grain(X1)", "inference_rule": "split_conjunct", "raw_label": "cnf(c_0_36, plain, (plant(X1)|~grain(X1)),\\ninference(split_conjunct,[status(thm)],[c_0_31])).\n", "fillcolor": "palegreen"}, {"id": 37, "label": "c_0_37", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_37", "role": "plain", "formula": "grain(esk6_0)", "inference_rule": "split_conjunct", "raw_label": "cnf(c_0_37, plain, (grain(esk6_0)),\\ninference(split_conjunct,[status(thm)],[c_0_32])).\n", "fillcolor": "palegreen"}, {"id": 38, "label": "c_0_38", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_38", "role": "plain", "formula": "eats(X1,X2)|~plant(X2)|~snail(X3)|~bird(X1)", "inference_rule": "csr", "raw_label": "cnf(c_0_38, plain, (eats(X1,X2)|~plant(X2)|~snail(X3)|~bird(X1)),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_35])).\n", "fillcolor": "palegreen"}, {"id": 39, "label": "c_0_39", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_39", "role": "plain", "formula": "plant(esk6_0)", "inference_rule": "spm", "raw_label": "cnf(c_0_39, plain, (plant(esk6_0)),\\ninference(spm,[status(thm)],[c_0_36, c_0_37])).\n", "fillcolor": "palegreen"}, {"id": 40, "label": "40", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "skolemize", "raw_label": "fof(c_0_40, plain, snail(esk5_0),\\ninference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[snail_type])])).\n", "fillcolor": "palegreen"}, {"id": 41, "label": "41", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "fof_nnf", "raw_label": "fof(c_0_41, plain, ![X26, X27]:((~bird(X26)|~fox(X27)|much_smaller(X26,X27))),\\ninference(fof_nnf,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[pel47_9])])])).\n", "fillcolor": "palegreen"}, {"id": 42, "label": "42", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "fof_nnf", "raw_label": "fof(c_0_42, plain, ![X12]:((~fox(X12)|animal(X12))),\\ninference(fof_nnf,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[pel47_2_1])])])).\n", "fillcolor": "palegreen"}, {"id": 43, "label": "c_0_43", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_43", "role": "plain", "formula": "eats(X1,esk6_0)|~snail(X2)|~bird(X1)", "inference_rule": "spm", "raw_label": "cnf(c_0_43, plain, (eats(X1,esk6_0)|~snail(X2)|~bird(X1)),\\ninference(spm,[status(thm)],[c_0_38, c_0_39])).\n", "fillcolor": "palegreen"}, {"id": 44, "label": "c_0_44", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_44", "role": "plain", "formula": "snail(esk5_0)", "inference_rule": "split_conjunct", "raw_label": "cnf(c_0_44, plain, (snail(esk5_0)),\\ninference(split_conjunct,[status(thm)],[c_0_40])).\n", "fillcolor": "palegreen"}, {"id": 45, "label": "45", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "skolemize", "raw_label": "fof(c_0_45, plain, bird(esk3_0),\\ninference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[bird_type])])).\n", "fillcolor": "palegreen"}, {"id": 46, "label": "46", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "assume_negation", "raw_label": "fof(c_0_46, negated_conjecture, ~(?[X2, X3]:(((animal(X2)&animal(X3))&?[X5]:(((grain(X5)&eats(X3,X5))&eats(X2,X3)))))),\\ninference(assume_negation,[status(cth)],[pel47])).\n", "fillcolor": "lightskyblue1"}, {"id": 47, "label": "c_0_47", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_47", "role": "plain", "formula": "much_smaller(X1,X2)|~bird(X1)|~fox(X2)", "inference_rule": "split_conjunct", "raw_label": "cnf(c_0_47, plain, (much_smaller(X1,X2)|~bird(X1)|~fox(X2)),\\ninference(split_conjunct,[status(thm)],[c_0_41])).\n", "fillcolor": "palegreen"}, {"id": 48, "label": "c_0_48", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_48", "role": "plain", "formula": "animal(X1)|~fox(X1)", "inference_rule": "split_conjunct", "raw_label": "cnf(c_0_48, plain, (animal(X1)|~fox(X1)),\\ninference(split_conjunct,[status(thm)],[c_0_42])).\n", "fillcolor": "palegreen"}, {"id": 49, "label": "c_0_49", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_49", "role": "plain", "formula": "eats(X1,esk6_0)|~bird(X1)", "inference_rule": "spm", "raw_label": "cnf(c_0_49, plain, (eats(X1,esk6_0)|~bird(X1)),\\ninference(spm,[status(thm)],[c_0_43, c_0_44])).\n", "fillcolor": "palegreen"}, {"id": 50, "label": "c_0_50", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_50", "role": "plain", "formula": "bird(esk3_0)", "inference_rule": "split_conjunct", "raw_label": "cnf(c_0_50, plain, (bird(esk3_0)),\\ninference(split_conjunct,[status(thm)],[c_0_45])).\n", "fillcolor": "palegreen"}, {"id": 51, "label": "51", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "fof_nnf", "raw_label": "fof(c_0_51, negated_conjecture, ![X42, X43, X44]:((~animal(X42)|~animal(X43)|(~grain(X44)|~eats(X43,X44)|~eats(X42,X43)))),\\ninference(fof_nnf,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_46])])])])).\n", "fillcolor": "lightskyblue1"}, {"id": 52, "label": "52", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "skolemize", "raw_label": "fof(c_0_52, plain, fox(esk2_0),\\ninference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[fox_type])])).\n", "fillcolor": "palegreen"}, {"id": 53, "label": "53", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "fof_simplification", "raw_label": "fof(c_0_53, plain, ![X2, X3]:(((wolf(X2)&fox(X3))=>~eats(X2,X3))),\\ninference(fof_simplification,[status(thm)],[pel47_11])).\n", "fillcolor": "palegreen"}, {"id": 54, "label": "c_0_54", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_54", "role": "plain", "formula": "eats(X1,X2)|eats(X1,X3)|~eats(X3,X4)|~plant(X4)|~plant(X2)|~bird(X3)|~fox(X1)", "inference_rule": "csr", "raw_label": "cnf(c_0_54, plain, (eats(X1,X2)|eats(X1,X3)|~eats(X3,X4)|~plant(X4)|~plant(X2)|~bird(X3)|~fox(X1)),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_47]), c_0_48]), c_0_27])).\n", "fillcolor": "palegreen"}, {"id": 55, "label": "c_0_55", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_55", "role": "plain", "formula": "eats(esk3_0,esk6_0)", "inference_rule": "spm", "raw_label": "cnf(c_0_55, plain, (eats(esk3_0,esk6_0)),\\ninference(spm,[status(thm)],[c_0_49, c_0_50])).\n", "fillcolor": "palegreen"}, {"id": 56, "label": "c_0_56", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_56", "role": "negated_conjecture", "formula": "~animal(X1)|~animal(X2)|~grain(X3)|~eats(X2,X3)|~eats(X1,X2)", "inference_rule": "split_conjunct", "raw_label": "cnf(c_0_56, negated_conjecture, (~animal(X1)|~animal(X2)|~grain(X3)|~eats(X2,X3)|~eats(X1,X2)),\\ninference(split_conjunct,[status(thm)],[c_0_51])).\n", "fillcolor": "lightskyblue1"}, {"id": 57, "label": "c_0_57", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_57", "role": "plain", "formula": "animal(esk3_0)", "inference_rule": "spm", "raw_label": "cnf(c_0_57, plain, (animal(esk3_0)),\\ninference(spm,[status(thm)],[c_0_27, c_0_50])).\n", "fillcolor": "palegreen"}, {"id": 58, "label": "c_0_58", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_58", "role": "plain", "formula": "fox(esk2_0)", "inference_rule": "split_conjunct", "raw_label": "cnf(c_0_58, plain, (fox(esk2_0)),\\ninference(split_conjunct,[status(thm)],[c_0_52])).\n", "fillcolor": "palegreen"}, {"id": 59, "label": "59", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "fof_nnf", "raw_label": "fof(c_0_59, plain, ![X28, X29]:((~fox(X28)|~wolf(X29)|much_smaller(X28,X29))),\\ninference(fof_nnf,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[pel47_10])])])).\n", "fillcolor": "palegreen"}, {"id": 60, "label": "60", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "fof_nnf", "raw_label": "fof(c_0_60, plain, ![X11]:((~wolf(X11)|animal(X11))),\\ninference(fof_nnf,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[pel47_1_1])])])).\n", "fillcolor": "palegreen"}, {"id": 61, "label": "61", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "fof_nnf", "raw_label": "fof(c_0_61, plain, ![X30, X31]:((~wolf(X30)|~fox(X31)|~eats(X30,X31))),\\ninference(fof_nnf,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_53])])])).\n", "fillcolor": "palegreen"}, {"id": 62, "label": "c_0_62", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_62", "role": "plain", "formula": "eats(X1,esk3_0)|eats(X1,X2)|~plant(X2)|~fox(X1)", "inference_rule": "cn", "raw_label": "cnf(c_0_62, plain, (eats(X1,esk3_0)|eats(X1,X2)|~plant(X2)|~fox(X1)),\\ninference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_39]), c_0_50])])).\n", "fillcolor": "palegreen"}, {"id": 63, "label": "c_0_63", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_63", "role": "negated_conjecture", "formula": "~eats(X1,esk3_0)|~animal(X1)", "inference_rule": "cn", "raw_label": "cnf(c_0_63, negated_conjecture, (~eats(X1,esk3_0)|~animal(X1)),\\ninference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_55]), c_0_37]), c_0_57])])).\n", "fillcolor": "lightskyblue1"}, {"id": 64, "label": "c_0_64", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_64", "role": "plain", "formula": "animal(esk2_0)", "inference_rule": "spm", "raw_label": "cnf(c_0_64, plain, (animal(esk2_0)),\\ninference(spm,[status(thm)],[c_0_48, c_0_58])).\n", "fillcolor": "palegreen"}, {"id": 65, "label": "c_0_65", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_65", "role": "plain", "formula": "much_smaller(X1,X2)|~fox(X1)|~wolf(X2)", "inference_rule": "split_conjunct", "raw_label": "cnf(c_0_65, plain, (much_smaller(X1,X2)|~fox(X1)|~wolf(X2)),\\ninference(split_conjunct,[status(thm)],[c_0_59])).\n", "fillcolor": "palegreen"}, {"id": 66, "label": "c_0_66", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_66", "role": "plain", "formula": "animal(X1)|~wolf(X1)", "inference_rule": "split_conjunct", "raw_label": "cnf(c_0_66, plain, (animal(X1)|~wolf(X1)),\\ninference(split_conjunct,[status(thm)],[c_0_60])).\n", "fillcolor": "palegreen"}, {"id": 67, "label": "c_0_67", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_67", "role": "plain", "formula": "~wolf(X1)|~fox(X2)|~eats(X1,X2)", "inference_rule": "split_conjunct", "raw_label": "cnf(c_0_67, plain, (~wolf(X1)|~fox(X2)|~eats(X1,X2)),\\ninference(split_conjunct,[status(thm)],[c_0_61])).\n", "fillcolor": "palegreen"}, {"id": 68, "label": "c_0_68", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_68", "role": "plain", "formula": "eats(X1,esk6_0)|eats(X1,esk3_0)|~fox(X1)", "inference_rule": "spm", "raw_label": "cnf(c_0_68, plain, (eats(X1,esk6_0)|eats(X1,esk3_0)|~fox(X1)),\\ninference(spm,[status(thm)],[c_0_62, c_0_39])).\n", "fillcolor": "palegreen"}, {"id": 69, "label": "c_0_69", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_69", "role": "negated_conjecture", "formula": "~eats(esk2_0,esk3_0)", "inference_rule": "spm", "raw_label": "cnf(c_0_69, negated_conjecture, (~eats(esk2_0,esk3_0)),\\ninference(spm,[status(thm)],[c_0_63, c_0_64])).\n", "fillcolor": "lightskyblue1"}, {"id": 70, "label": "c_0_70", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_70", "role": "plain", "formula": "eats(X1,X2)|~eats(X3,X4)|~plant(X4)|~plant(X2)|~fox(X3)|~wolf(X1)", "inference_rule": "csr", "raw_label": "cnf(c_0_70, plain, (eats(X1,X2)|~eats(X3,X4)|~plant(X4)|~plant(X2)|~fox(X3)|~wolf(X1)),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_65]), c_0_66]), c_0_48]), c_0_67])).\n", "fillcolor": "palegreen"}, {"id": 71, "label": "c_0_71", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_71", "role": "plain", "formula": "eats(esk2_0,esk6_0)", "inference_rule": "sr", "raw_label": "cnf(c_0_71, plain, (eats(esk2_0,esk6_0)),\\ninference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_58]), c_0_69])).\n", "fillcolor": "palegreen"}, {"id": 72, "label": "72", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "fof_simplification", "raw_label": "fof(c_0_72, plain, ![X2, X3]:(((wolf(X2)&grain(X3))=>~eats(X2,X3))),\\ninference(fof_simplification,[status(thm)],[pel47_11a])).\n", "fillcolor": "palegreen"}, {"id": 73, "label": "c_0_73", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_73", "role": "plain", "formula": "eats(X1,X2)|~plant(X2)|~wolf(X1)", "inference_rule": "cn", "raw_label": "cnf(c_0_73, plain, (eats(X1,X2)|~plant(X2)|~wolf(X1)),\\ninference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_39]), c_0_58])])).\n", "fillcolor": "palegreen"}, {"id": 74, "label": "74", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "skolemize", "raw_label": "fof(c_0_74, plain, wolf(esk1_0),\\ninference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[wolf_type])])).\n", "fillcolor": "palegreen"}, {"id": 75, "label": "75", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "fof_nnf", "raw_label": "fof(c_0_75, plain, ![X32, X33]:((~wolf(X32)|~grain(X33)|~eats(X32,X33))),\\ninference(fof_nnf,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_72])])])).\n", "fillcolor": "palegreen"}, {"id": 76, "label": "c_0_76", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_76", "role": "plain", "formula": "eats(X1,esk6_0)|~wolf(X1)", "inference_rule": "spm", "raw_label": "cnf(c_0_76, plain, (eats(X1,esk6_0)|~wolf(X1)),\\ninference(spm,[status(thm)],[c_0_73, c_0_39])).\n", "fillcolor": "palegreen"}, {"id": 77, "label": "c_0_77", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_77", "role": "plain", "formula": "wolf(esk1_0)", "inference_rule": "split_conjunct", "raw_label": "cnf(c_0_77, plain, (wolf(esk1_0)),\\ninference(split_conjunct,[status(thm)],[c_0_74])).\n", "fillcolor": "palegreen"}, {"id": 78, "label": "c_0_78", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_78", "role": "plain", "formula": "~wolf(X1)|~grain(X2)|~eats(X1,X2)", "inference_rule": "split_conjunct", "raw_label": "cnf(c_0_78, plain, (~wolf(X1)|~grain(X2)|~eats(X1,X2)),\\ninference(split_conjunct,[status(thm)],[c_0_75])).\n", "fillcolor": "palegreen"}, {"id": 79, "label": "c_0_79", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_79", "role": "plain", "formula": "eats(esk1_0,esk6_0)", "inference_rule": "spm", "raw_label": "cnf(c_0_79, plain, (eats(esk1_0,esk6_0)),\\ninference(spm,[status(thm)],[c_0_76, c_0_77])).\n", "fillcolor": "palegreen"}, {"id": 80, "label": "c_0_80", "color": {"background": "#000000", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#FFFFFF"}, "cnf_name": "c_0_80", "role": "plain", "formula": "$false", "inference_rule": "cn", "raw_label": "cnf(c_0_80, plain, ($false),\\ninference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_79]), c_0_37]), c_0_77])])).\n", "fillcolor": "darkorchid1"}], "edges": [{"from": 0, "to": 19, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 1, "to": 20, "label": "fof_nnf", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 2, "to": 21, "label": "fof_nnf", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 3, "to": 22, "label": "fof_nnf", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 4, "to": 23, "label": "fof_nnf", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 19, "to": 24, "label": "fof_nnf", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 20, "to": 25, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 21, "to": 26, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 22, "to": 27, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 23, "to": 28, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 24, "to": 29, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 5, "to": 30, "label": "distribute", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 6, "to": 31, "label": "fof_nnf", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 7, "to": 32, "label": "skolemize", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 29, "to": 33, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 28, "to": 33, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 27, "to": 33, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 26, "to": 33, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 25, "to": 33, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 30, "to": 34, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 30, "to": 35, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 31, "to": 36, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 32, "to": 37, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 35, "to": 38, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 34, "to": 38, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 33, "to": 38, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 37, "to": 39, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 36, "to": 39, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 8, "to": 40, "label": "skolemize", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 9, "to": 41, "label": "fof_nnf", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 10, "to": 42, "label": "fof_nnf", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 39, "to": 43, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 38, "to": 43, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 40, "to": 44, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 11, "to": 45, "label": "skolemize", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 12, "to": 46, "label": "assume_negation", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 41, "to": 47, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 42, "to": 48, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 44, "to": 49, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 43, "to": 49, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 45, "to": 50, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 46, "to": 51, "label": "fof_nnf", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 13, "to": 52, "label": "skolemize", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 14, "to": 53, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 27, "to": 54, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 48, "to": 54, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 47, "to": 54, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 25, "to": 54, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 50, "to": 55, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 49, "to": 55, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 51, "to": 56, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 50, "to": 57, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 27, "to": 57, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 52, "to": 58, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 15, "to": 59, "label": "fof_nnf", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 16, "to": 60, "label": "fof_nnf", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 53, "to": 61, "label": "fof_nnf", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 50, "to": 62, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 39, "to": 62, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 55, "to": 62, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 54, "to": 62, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 57, "to": 63, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 37, "to": 63, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 55, "to": 63, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 56, "to": 63, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 58, "to": 64, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 48, "to": 64, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 59, "to": 65, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 60, "to": 66, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 61, "to": 67, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 39, "to": 68, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 62, "to": 68, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 64, "to": 69, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 63, "to": 69, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 67, "to": 70, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 48, "to": 70, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 66, "to": 70, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 65, "to": 70, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 25, "to": 70, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 69, "to": 71, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 58, "to": 71, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 68, "to": 71, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 17, "to": 72, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 58, "to": 73, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 39, "to": 73, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 71, "to": 73, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 70, "to": 73, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 18, "to": 74, "label": "skolemize", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 72, "to": 75, "label": "fof_nnf", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 39, "to": 76, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 73, "to": 76, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 74, "to": 77, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 75, "to": 78, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 77, "to": 79, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 76, "to": 79, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 77, "to": 80, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 37, "to": 80, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 79, "to": 80, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 78, "to": 80, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}]}, {"problem_file": "PUZ021-1.p", "status": "# SZS status Unsatisfiable", "nodes": [{"id": 0, "label": "conjunction1", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "conjunction1", "role": "axiom", "formula": "a_truth(X1,X3)|~a_truth(and(X1,X2", "inference_rule": null, "raw_label": "cnf(conjunction1, axiom, (a_truth(X1,X3)|~a_truth(and(X1,X2),X3)),\\nfile('Problems/PUZ/PUZ021-1.p', conjunction1)).\n", "fillcolor": "forestgreen"}, {"id": 1, "label": "knaves_lie2", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "knaves_lie2", "role": "axiom", "formula": "says(X1,X3)|a_truth(X3,X2)|~a_truth(knave(X1", "inference_rule": null, "raw_label": "cnf(knaves_lie2, axiom, (says(X1,X3)|a_truth(X3,X2)|~a_truth(knave(X1),X2)),\\nfile('Problems/PUZ/PUZ021-1.p', knaves_lie2)).\n", "fillcolor": "forestgreen"}, {"id": 2, "label": "prove_statement_ex", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "prove_statement_exists2", "role": "negated_conjecture", "formula": "says(me,X1)|a_truth(and(knave(me", "inference_rule": null, "raw_label": "cnf(prove_statement_exists2, negated_conjecture, (says(me,X1)|a_truth(and(knave(me),rich(me)),X1)),\\nfile('Problems/PUZ/PUZ021-1.p', prove_statement_exists2)).\n", "fillcolor": "dodgerblue"}, {"id": 3, "label": "knights_tell_truth", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "knights_tell_truth1", "role": "axiom", "formula": "a_truth(X3,X2)|~a_truth(knight(X1", "inference_rule": null, "raw_label": "cnf(knights_tell_truth1, axiom, (a_truth(X3,X2)|~a_truth(knight(X1),X2)|~says(X1,X3)),\\nfile('Problems/PUZ/PUZ021-1.p', knights_tell_truth1)).\n", "fillcolor": "forestgreen"}, {"id": 4, "label": "knight_or_knave", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "knight_or_knave", "role": "axiom", "formula": "a_truth(knight(X1", "inference_rule": null, "raw_label": "cnf(knight_or_knave, axiom, (a_truth(knight(X1),X2)|a_truth(knave(X1),X2)),\\nfile('Problems/PUZ/PUZ021-1.p', knight_or_knave)).\n", "fillcolor": "forestgreen"}, {"id": 5, "label": "knaves_lie1", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "knaves_lie1", "role": "axiom", "formula": "~a_truth(knave(X1", "inference_rule": null, "raw_label": "cnf(knaves_lie1, axiom, (~a_truth(knave(X1),X2)|~says(X1,X3)|~a_truth(X3,X2)),\\nfile('Problems/PUZ/PUZ021-1.p', knaves_lie1)).\n", "fillcolor": "forestgreen"}, {"id": 6, "label": "conjunction3", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "conjunction3", "role": "axiom", "formula": "a_truth(and(X1,X2", "inference_rule": null, "raw_label": "cnf(conjunction3, axiom, (a_truth(and(X1,X2),X3)|~a_truth(X1,X3)|~a_truth(X2,X3)),\\nfile('Problems/PUZ/PUZ021-1.p', conjunction3)).\n", "fillcolor": "forestgreen"}, {"id": 7, "label": "not_rich_and_poor", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "not_rich_and_poor", "role": "axiom", "formula": "~a_truth(rich(X1", "inference_rule": null, "raw_label": "cnf(not_rich_and_poor, axiom, (~a_truth(rich(X1),X2)|~a_truth(poor(X1),X2)),\\nfile('Problems/PUZ/PUZ021-1.p', not_rich_and_poor)).\n", "fillcolor": "forestgreen"}, {"id": 8, "label": "conjunction2", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "conjunction2", "role": "axiom", "formula": "a_truth(X2,X3)|~a_truth(and(X1,X2", "inference_rule": null, "raw_label": "cnf(conjunction2, axiom, (a_truth(X2,X3)|~a_truth(and(X1,X2),X3)),\\nfile('Problems/PUZ/PUZ021-1.p', conjunction2)).\n", "fillcolor": "forestgreen"}, {"id": 9, "label": "not_knight_and_kna", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "not_knight_and_knave", "role": "axiom", "formula": "~a_truth(knight(X1", "inference_rule": null, "raw_label": "cnf(not_knight_and_knave, axiom, (~a_truth(knight(X1),X2)|~a_truth(knave(X1),X2)),\\nfile('Problems/PUZ/PUZ021-1.p', not_knight_and_knave)).\n", "fillcolor": "forestgreen"}, {"id": 10, "label": "prove_statement_ex", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "prove_statement_exists1", "role": "negated_conjecture", "formula": "~says(me,X1)|~a_truth(and(knave(me", "inference_rule": null, "raw_label": "cnf(prove_statement_exists1, negated_conjecture, (~says(me,X1)|~a_truth(and(knave(me),rich(me)),X1)),\\nfile('Problems/PUZ/PUZ021-1.p', prove_statement_exists1)).\n", "fillcolor": "dodgerblue"}, {"id": 11, "label": "rich_or_poor", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "rich_or_poor", "role": "axiom", "formula": "a_truth(rich(X1", "inference_rule": null, "raw_label": "cnf(rich_or_poor, axiom, (a_truth(rich(X1),X2)|a_truth(poor(X1),X2)),\\nfile('Problems/PUZ/PUZ021-1.p', rich_or_poor)).\n", "fillcolor": "forestgreen"}, {"id": 12, "label": "c_0_12", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_12", "role": "plain", "formula": "a_truth(X1,X3)|~a_truth(and(X1,X2", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_12, plain, (a_truth(X1,X3)|~a_truth(and(X1,X2),X3)),\\ninference(fof_simplification,[status(thm)],[conjunction1])).\n", "fillcolor": "palegreen"}, {"id": 13, "label": "c_0_13", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_13", "role": "plain", "formula": "says(X1,X3)|a_truth(X3,X2)|~a_truth(knave(X1", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_13, plain, (says(X1,X3)|a_truth(X3,X2)|~a_truth(knave(X1),X2)),\\ninference(fof_simplification,[status(thm)],[knaves_lie2])).\n", "fillcolor": "palegreen"}, {"id": 14, "label": "c_0_14", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_14", "role": "plain", "formula": "a_truth(X1,X3)|~a_truth(and(X1,X2", "inference_rule": null, "raw_label": "cnf(c_0_14, plain, (a_truth(X1,X3)|~a_truth(and(X1,X2),X3)),\\nc_0_12).\n", "fillcolor": "palegreen"}, {"id": 15, "label": "c_0_15", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_15", "role": "negated_conjecture", "formula": "says(me,X1)|a_truth(and(knave(me", "inference_rule": null, "raw_label": "cnf(c_0_15, negated_conjecture, (says(me,X1)|a_truth(and(knave(me),rich(me)),X1)),\\nprove_statement_exists2).\n", "fillcolor": "lightskyblue1"}, {"id": 16, "label": "c_0_16", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_16", "role": "plain", "formula": "a_truth(X3,X2)|~a_truth(knight(X1", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_16, plain, (a_truth(X3,X2)|~a_truth(knight(X1),X2)|~says(X1,X3)),\\ninference(fof_simplification,[status(thm)],[knights_tell_truth1])).\n", "fillcolor": "palegreen"}, {"id": 17, "label": "c_0_17", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_17", "role": "plain", "formula": "says(X1,X3)|a_truth(X3,X2)|~a_truth(knave(X1", "inference_rule": null, "raw_label": "cnf(c_0_17, plain, (says(X1,X3)|a_truth(X3,X2)|~a_truth(knave(X1),X2)),\\nc_0_13).\n", "fillcolor": "palegreen"}, {"id": 18, "label": "c_0_18", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_18", "role": "negated_conjecture", "formula": "says(me,X1)|a_truth(knave(me", "inference_rule": "spm", "raw_label": "cnf(c_0_18, negated_conjecture, (says(me,X1)|a_truth(knave(me),X1)),\\ninference(spm,[status(thm)],[c_0_14, c_0_15])).\n", "fillcolor": "lightskyblue1"}, {"id": 19, "label": "c_0_19", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_19", "role": "plain", "formula": "a_truth(X3,X2)|~a_truth(knight(X1", "inference_rule": null, "raw_label": "cnf(c_0_19, plain, (a_truth(X3,X2)|~a_truth(knight(X1),X2)|~says(X1,X3)),\\nc_0_16).\n", "fillcolor": "palegreen"}, {"id": 20, "label": "c_0_20", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_20", "role": "axiom", "formula": "a_truth(knight(X1", "inference_rule": null, "raw_label": "cnf(c_0_20, axiom, (a_truth(knight(X1),X2)|a_truth(knave(X1),X2)),\\nknight_or_knave).\n", "fillcolor": "palegreen"}, {"id": 21, "label": "c_0_21", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_21", "role": "negated_conjecture", "formula": "says(me,X1)|says(me,X2)|a_truth(X2,X1)", "inference_rule": "spm", "raw_label": "cnf(c_0_21, negated_conjecture, (says(me,X1)|says(me,X2)|a_truth(X2,X1)),\\ninference(spm,[status(thm)],[c_0_17, c_0_18])).\n", "fillcolor": "lightskyblue1"}, {"id": 22, "label": "c_0_22", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_22", "role": "plain", "formula": "~a_truth(knave(X1", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_22, plain, (~a_truth(knave(X1),X2)|~says(X1,X3)|~a_truth(X3,X2)),\\ninference(fof_simplification,[status(thm)],[knaves_lie1])).\n", "fillcolor": "palegreen"}, {"id": 23, "label": "c_0_23", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_23", "role": "plain", "formula": "a_truth(knave(X1", "inference_rule": "spm", "raw_label": "cnf(c_0_23, plain, (a_truth(knave(X1),X2)|a_truth(X3,X2)|~says(X1,X3)),\\ninference(spm,[status(thm)],[c_0_19, c_0_20])).\n", "fillcolor": "palegreen"}, {"id": 24, "label": "c_0_24", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_24", "role": "negated_conjecture", "formula": "says(me,X1)|a_truth(X1,X1)", "inference_rule": "ef", "raw_label": "cnf(c_0_24, negated_conjecture, (says(me,X1)|a_truth(X1,X1)),\\ninference(ef,[status(thm)],[c_0_21])).\n", "fillcolor": "lightskyblue1"}, {"id": 25, "label": "c_0_25", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_25", "role": "plain", "formula": "~a_truth(knave(X1", "inference_rule": null, "raw_label": "cnf(c_0_25, plain, (~a_truth(knave(X1),X2)|~says(X1,X3)|~a_truth(X3,X2)),\\nc_0_22).\n", "fillcolor": "palegreen"}, {"id": 26, "label": "c_0_26", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_26", "role": "negated_conjecture", "formula": "a_truth(knave(me", "inference_rule": "spm", "raw_label": "cnf(c_0_26, negated_conjecture, (a_truth(knave(me),X1)|a_truth(X2,X2)|a_truth(X2,X1)),\\ninference(spm,[status(thm)],[c_0_23, c_0_24])).\n", "fillcolor": "lightskyblue1"}, {"id": 27, "label": "c_0_27", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_27", "role": "negated_conjecture", "formula": "says(me,X1)|~says(me,X2)|~a_truth(X2,X1)", "inference_rule": "spm", "raw_label": "cnf(c_0_27, negated_conjecture, (says(me,X1)|~says(me,X2)|~a_truth(X2,X1)),\\ninference(spm,[status(thm)],[c_0_25, c_0_18])).\n", "fillcolor": "lightskyblue1"}, {"id": 28, "label": "c_0_28", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_28", "role": "negated_conjecture", "formula": "a_truth(knave(me", "inference_rule": "ef", "raw_label": "cnf(c_0_28, negated_conjecture, (a_truth(knave(me),knave(me))),\\ninference(ef,[status(thm)],[c_0_26])).\n", "fillcolor": "lightskyblue1"}, {"id": 29, "label": "c_0_29", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_29", "role": "plain", "formula": "a_truth(and(X1,X2", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_29, plain, (a_truth(and(X1,X2),X3)|~a_truth(X1,X3)|~a_truth(X2,X3)),\\ninference(fof_simplification,[status(thm)],[conjunction3])).\n", "fillcolor": "palegreen"}, {"id": 30, "label": "c_0_30", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_30", "role": "negated_conjecture", "formula": "says(me,X1)|a_truth(X2,X2)|~a_truth(X2,X1)", "inference_rule": "spm", "raw_label": "cnf(c_0_30, negated_conjecture, (says(me,X1)|a_truth(X2,X2)|~a_truth(X2,X1)),\\ninference(spm,[status(thm)],[c_0_27, c_0_24])).\n", "fillcolor": "lightskyblue1"}, {"id": 31, "label": "c_0_31", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_31", "role": "negated_conjecture", "formula": "~says(me,X1)|~a_truth(X1,knave(me))", "inference_rule": "spm", "raw_label": "cnf(c_0_31, negated_conjecture, (~says(me,X1)|~a_truth(X1,knave(me))),\\ninference(spm,[status(thm)],[c_0_25, c_0_28])).\n", "fillcolor": "lightskyblue1"}, {"id": 32, "label": "c_0_32", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_32", "role": "plain", "formula": "a_truth(and(X1,X2", "inference_rule": null, "raw_label": "cnf(c_0_32, plain, (a_truth(and(X1,X2),X3)|~a_truth(X1,X3)|~a_truth(X2,X3)),\\nc_0_29).\n", "fillcolor": "palegreen"}, {"id": 33, "label": "c_0_33", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_33", "role": "negated_conjecture", "formula": "says(me,X1)|a_truth(and(knave(me", "inference_rule": "spm", "raw_label": "cnf(c_0_33, negated_conjecture, (says(me,X1)|a_truth(and(knave(me),rich(me)),and(knave(me),rich(me)))),\\ninference(spm,[status(thm)],[c_0_30, c_0_15])).\n", "fillcolor": "lightskyblue1"}, {"id": 34, "label": "c_0_34", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_34", "role": "negated_conjecture", "formula": "~says(me,and(X1,X2))|~a_truth(X2,knave(me))|~a_truth(X1,knave(me))", "inference_rule": "spm", "raw_label": "cnf(c_0_34, negated_conjecture, (~says(me,and(X1,X2))|~a_truth(X2,knave(me))|~a_truth(X1,knave(me))),\\ninference(spm,[status(thm)],[c_0_31, c_0_32])).\n", "fillcolor": "lightskyblue1"}, {"id": 35, "label": "c_0_35", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_35", "role": "negated_conjecture", "formula": "says(me,X1)|a_truth(knave(me", "inference_rule": "spm", "raw_label": "cnf(c_0_35, negated_conjecture, (says(me,X1)|a_truth(knave(me),and(knave(me),rich(me)))),\\ninference(spm,[status(thm)],[c_0_14, c_0_33])).\n", "fillcolor": "lightskyblue1"}, {"id": 36, "label": "c_0_36", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_36", "role": "negated_conjecture", "formula": "a_truth(knave(me", "inference_rule": "spm", "raw_label": "cnf(c_0_36, negated_conjecture, (a_truth(knave(me),and(knave(me),rich(me)))|~a_truth(X1,knave(me))|~a_truth(X2,knave(me))),\\ninference(spm,[status(thm)],[c_0_34, c_0_35])).\n", "fillcolor": "lightskyblue1"}, {"id": 37, "label": "c_0_37", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_37", "role": "negated_conjecture", "formula": "~says(me,knave(me))", "inference_rule": "spm", "raw_label": "cnf(c_0_37, negated_conjecture, (~says(me,knave(me))),\\ninference(spm,[status(thm)],[c_0_31, c_0_28])).\n", "fillcolor": "lightskyblue1"}, {"id": 38, "label": "c_0_38", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_38", "role": "plain", "formula": "~a_truth(rich(X1", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_38, plain, (~a_truth(rich(X1),X2)|~a_truth(poor(X1),X2)),\\ninference(fof_simplification,[status(thm)],[not_rich_and_poor])).\n", "fillcolor": "palegreen"}, {"id": 39, "label": "c_0_39", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_39", "role": "plain", "formula": "a_truth(X2,X3)|~a_truth(and(X1,X2", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_39, plain, (a_truth(X2,X3)|~a_truth(and(X1,X2),X3)),\\ninference(fof_simplification,[status(thm)],[conjunction2])).\n", "fillcolor": "palegreen"}, {"id": 40, "label": "c_0_40", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_40", "role": "negated_conjecture", "formula": "a_truth(knave(me", "inference_rule": "sr", "raw_label": "cnf(c_0_40, negated_conjecture, (a_truth(knave(me),and(knave(me),rich(me)))|~a_truth(X1,knave(me))),\\ninference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_15]), c_0_37])).\n", "fillcolor": "lightskyblue1"}, {"id": 41, "label": "c_0_41", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_41", "role": "plain", "formula": "~a_truth(rich(X1", "inference_rule": null, "raw_label": "cnf(c_0_41, plain, (~a_truth(rich(X1),X2)|~a_truth(poor(X1),X2)),\\nc_0_38).\n", "fillcolor": "palegreen"}, {"id": 42, "label": "c_0_42", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_42", "role": "plain", "formula": "a_truth(X2,X3)|~a_truth(and(X1,X2", "inference_rule": null, "raw_label": "cnf(c_0_42, plain, (a_truth(X2,X3)|~a_truth(and(X1,X2),X3)),\\nc_0_39).\n", "fillcolor": "palegreen"}, {"id": 43, "label": "c_0_43", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_43", "role": "negated_conjecture", "formula": "a_truth(knave(me", "inference_rule": "sr", "raw_label": "cnf(c_0_43, negated_conjecture, (a_truth(knave(me),and(knave(me),rich(me)))),\\ninference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_15]), c_0_37])).\n", "fillcolor": "lightskyblue1"}, {"id": 44, "label": "c_0_44", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_44", "role": "negated_conjecture", "formula": "says(me,poor(X1))|~a_truth(rich(X1", "inference_rule": "spm", "raw_label": "cnf(c_0_44, negated_conjecture, (says(me,poor(X1))|~a_truth(rich(X1),poor(X1))),\\ninference(spm,[status(thm)],[c_0_41, c_0_24])).\n", "fillcolor": "lightskyblue1"}, {"id": 45, "label": "c_0_45", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_45", "role": "negated_conjecture", "formula": "says(me,X1)|a_truth(rich(me", "inference_rule": "spm", "raw_label": "cnf(c_0_45, negated_conjecture, (says(me,X1)|a_truth(rich(me),X1)),\\ninference(spm,[status(thm)],[c_0_42, c_0_15])).\n", "fillcolor": "lightskyblue1"}, {"id": 46, "label": "c_0_46", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_46", "role": "plain", "formula": "~a_truth(knight(X1", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_46, plain, (~a_truth(knight(X1),X2)|~a_truth(knave(X1),X2)),\\ninference(fof_simplification,[status(thm)],[not_knight_and_knave])).\n", "fillcolor": "palegreen"}, {"id": 47, "label": "c_0_47", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_47", "role": "negated_conjecture", "formula": "~says(me,X1)|~a_truth(X1,and(knave(me", "inference_rule": "spm", "raw_label": "cnf(c_0_47, negated_conjecture, (~says(me,X1)|~a_truth(X1,and(knave(me),rich(me)))),\\ninference(spm,[status(thm)],[c_0_25, c_0_43])).\n", "fillcolor": "lightskyblue1"}, {"id": 48, "label": "c_0_48", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_48", "role": "negated_conjecture", "formula": "says(me,X1)|a_truth(knight(X2", "inference_rule": "spm", "raw_label": "cnf(c_0_48, negated_conjecture, (says(me,X1)|a_truth(knight(X2),knight(X2))|a_truth(knave(X2),X1)),\\ninference(spm,[status(thm)],[c_0_30, c_0_20])).\n", "fillcolor": "lightskyblue1"}, {"id": 49, "label": "c_0_49", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_49", "role": "negated_conjecture", "formula": "~says(me,and(knave(me", "inference_rule": "sr", "raw_label": "cnf(c_0_49, negated_conjecture, (~says(me,and(knave(me),rich(me)))),\\ninference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_15]), c_0_37])).\n", "fillcolor": "lightskyblue1"}, {"id": 50, "label": "c_0_50", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_50", "role": "negated_conjecture", "formula": "says(me,poor(me))", "inference_rule": "spm", "raw_label": "cnf(c_0_50, negated_conjecture, (says(me,poor(me))),\\ninference(spm,[status(thm)],[c_0_44, c_0_45])).\n", "fillcolor": "lightskyblue1"}, {"id": 51, "label": "c_0_51", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_51", "role": "plain", "formula": "~a_truth(knight(X1", "inference_rule": null, "raw_label": "cnf(c_0_51, plain, (~a_truth(knight(X1),X2)|~a_truth(knave(X1),X2)),\\nc_0_46).\n", "fillcolor": "palegreen"}, {"id": 52, "label": "c_0_52", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_52", "role": "negated_conjecture", "formula": "a_truth(knight(X1", "inference_rule": "sr", "raw_label": "cnf(c_0_52, negated_conjecture, (a_truth(knight(X1),knight(X1))|~says(me,knave(X1))),\\ninference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_49])).\n", "fillcolor": "lightskyblue1"}, {"id": 53, "label": "c_0_53", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_53", "role": "negated_conjecture", "formula": "says(me,knave(X1))|~says(X1,X2)|~a_truth(X2,knave(X1))", "inference_rule": "spm", "raw_label": "cnf(c_0_53, negated_conjecture, (says(me,knave(X1))|~says(X1,X2)|~a_truth(X2,knave(X1))),\\ninference(spm,[status(thm)],[c_0_25, c_0_24])).\n", "fillcolor": "lightskyblue1"}, {"id": 54, "label": "c_0_54", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_54", "role": "negated_conjecture", "formula": "a_truth(knave(me", "inference_rule": "ef", "raw_label": "cnf(c_0_54, negated_conjecture, (a_truth(knave(me),X1)|a_truth(X1,X1)),\\ninference(ef,[status(thm)],[c_0_26])).\n", "fillcolor": "lightskyblue1"}, {"id": 55, "label": "c_0_55", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_55", "role": "negated_conjecture", "formula": "says(me,X1)|~a_truth(poor(me", "inference_rule": "spm", "raw_label": "cnf(c_0_55, negated_conjecture, (says(me,X1)|~a_truth(poor(me),X1)),\\ninference(spm,[status(thm)],[c_0_27, c_0_50])).\n", "fillcolor": "lightskyblue1"}, {"id": 56, "label": "c_0_56", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_56", "role": "negated_conjecture", "formula": "says(me,knave(X1))|says(X1,X2)|a_truth(X2,knave(X1))", "inference_rule": "spm", "raw_label": "cnf(c_0_56, negated_conjecture, (says(me,knave(X1))|says(X1,X2)|a_truth(X2,knave(X1))),\\ninference(spm,[status(thm)],[c_0_17, c_0_24])).\n", "fillcolor": "lightskyblue1"}, {"id": 57, "label": "c_0_57", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_57", "role": "negated_conjecture", "formula": "~says(me,knave(X1))|~a_truth(knave(X1", "inference_rule": "spm", "raw_label": "cnf(c_0_57, negated_conjecture, (~says(me,knave(X1))|~a_truth(knave(X1),knight(X1))),\\ninference(spm,[status(thm)],[c_0_51, c_0_52])).\n", "fillcolor": "lightskyblue1"}, {"id": 58, "label": "c_0_58", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_58", "role": "negated_conjecture", "formula": "a_truth(X1,knight(X2))|~says(me,knave(X2))|~says(X2,X1)", "inference_rule": "spm", "raw_label": "cnf(c_0_58, negated_conjecture, (a_truth(X1,knight(X2))|~says(me,knave(X2))|~says(X2,X1)),\\ninference(spm,[status(thm)],[c_0_19, c_0_52])).\n", "fillcolor": "lightskyblue1"}, {"id": 59, "label": "c_0_59", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_59", "role": "negated_conjecture", "formula": "says(me,knave(X1))|~says(X1,knave(X1))", "inference_rule": "spm", "raw_label": "cnf(c_0_59, negated_conjecture, (says(me,knave(X1))|~says(X1,knave(X1))),\\ninference(spm,[status(thm)],[c_0_53, c_0_24])).\n", "fillcolor": "lightskyblue1"}, {"id": 60, "label": "c_0_60", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_60", "role": "negated_conjecture", "formula": "~says(me,X1)|~a_truth(and(knave(me", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_60, negated_conjecture, (~says(me,X1)|~a_truth(and(knave(me),rich(me)),X1)),\\ninference(fof_simplification,[status(thm)],[prove_statement_exists1])).\n", "fillcolor": "lightskyblue1"}, {"id": 61, "label": "c_0_61", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_61", "role": "negated_conjecture", "formula": "a_truth(X1,X1)|~says(me,X2)|~a_truth(X2,X1)", "inference_rule": "spm", "raw_label": "cnf(c_0_61, negated_conjecture, (a_truth(X1,X1)|~says(me,X2)|~a_truth(X2,X1)),\\ninference(spm,[status(thm)],[c_0_25, c_0_54])).\n", "fillcolor": "lightskyblue1"}, {"id": 62, "label": "c_0_62", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_62", "role": "negated_conjecture", "formula": "says(X1,poor(me))|says(me,knave(X1))", "inference_rule": "spm", "raw_label": "cnf(c_0_62, negated_conjecture, (says(X1,poor(me))|says(me,knave(X1))),\\ninference(spm,[status(thm)],[c_0_55, c_0_56])).\n", "fillcolor": "lightskyblue1"}, {"id": 63, "label": "c_0_63", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_63", "role": "negated_conjecture", "formula": "~says(X1,knave(X1))", "inference_rule": "csr", "raw_label": "cnf(c_0_63, negated_conjecture, (~says(X1,knave(X1))),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_59])).\n", "fillcolor": "lightskyblue1"}, {"id": 64, "label": "c_0_64", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_64", "role": "negated_conjecture", "formula": "~says(me,X1)|~a_truth(and(knave(me", "inference_rule": null, "raw_label": "cnf(c_0_64, negated_conjecture, (~says(me,X1)|~a_truth(and(knave(me),rich(me)),X1)),\\nc_0_60).\n", "fillcolor": "lightskyblue1"}, {"id": 65, "label": "c_0_65", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_65", "role": "negated_conjecture", "formula": "a_truth(X1,X1)|~a_truth(poor(me", "inference_rule": "sr", "raw_label": "cnf(c_0_65, negated_conjecture, (a_truth(X1,X1)|~a_truth(poor(me),X1)),\\ninference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_63])).\n", "fillcolor": "lightskyblue1"}, {"id": 66, "label": "c_0_66", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_66", "role": "axiom", "formula": "a_truth(rich(X1", "inference_rule": null, "raw_label": "cnf(c_0_66, axiom, (a_truth(rich(X1),X2)|a_truth(poor(X1),X2)),\\nrich_or_poor).\n", "fillcolor": "palegreen"}, {"id": 67, "label": "c_0_67", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_67", "role": "negated_conjecture", "formula": "~says(me,X1)|~a_truth(rich(me", "inference_rule": "spm", "raw_label": "cnf(c_0_67, negated_conjecture, (~says(me,X1)|~a_truth(rich(me),X1)|~a_truth(knave(me),X1)),\\ninference(spm,[status(thm)],[c_0_64, c_0_32])).\n", "fillcolor": "lightskyblue1"}, {"id": 68, "label": "c_0_68", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_68", "role": "negated_conjecture", "formula": "a_truth(rich(me", "inference_rule": "spm", "raw_label": "cnf(c_0_68, negated_conjecture, (a_truth(rich(me),X1)|a_truth(X1,X1)),\\ninference(spm,[status(thm)],[c_0_65, c_0_66])).\n", "fillcolor": "lightskyblue1"}, {"id": 69, "label": "c_0_69", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_69", "role": "negated_conjecture", "formula": "says(me,and(X1,X2))|a_truth(X1,and(X1,X2))", "inference_rule": "spm", "raw_label": "cnf(c_0_69, negated_conjecture, (says(me,and(X1,X2))|a_truth(X1,and(X1,X2))),\\ninference(spm,[status(thm)],[c_0_14, c_0_24])).\n", "fillcolor": "lightskyblue1"}, {"id": 70, "label": "c_0_70", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_70", "role": "negated_conjecture", "formula": "a_truth(X1,X1)", "inference_rule": "csr", "raw_label": "cnf(c_0_70, negated_conjecture, (a_truth(X1,X1)),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68]), c_0_54]), c_0_24])).\n", "fillcolor": "lightskyblue1"}, {"id": 71, "label": "c_0_71", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_71", "role": "negated_conjecture", "formula": "says(me,and(poor(me", "inference_rule": "spm", "raw_label": "cnf(c_0_71, negated_conjecture, (says(me,and(poor(me),X1))),\\ninference(spm,[status(thm)],[c_0_55, c_0_69])).\n", "fillcolor": "lightskyblue1"}, {"id": 72, "label": "c_0_72", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_72", "role": "negated_conjecture", "formula": "a_truth(X1,and(X1,X2))", "inference_rule": "spm", "raw_label": "cnf(c_0_72, negated_conjecture, (a_truth(X1,and(X1,X2))),\\ninference(spm,[status(thm)],[c_0_14, c_0_70])).\n", "fillcolor": "lightskyblue1"}, {"id": 73, "label": "c_0_73", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_73", "role": "negated_conjecture", "formula": "says(me,X1)|~a_truth(and(poor(me", "inference_rule": "spm", "raw_label": "cnf(c_0_73, negated_conjecture, (says(me,X1)|~a_truth(and(poor(me),X2),X1)),\\ninference(spm,[status(thm)],[c_0_27, c_0_71])).\n", "fillcolor": "lightskyblue1"}, {"id": 74, "label": "c_0_74", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_74", "role": "negated_conjecture", "formula": "says(me,and(X1,X2))|a_truth(X2,and(X1,X2))", "inference_rule": "spm", "raw_label": "cnf(c_0_74, negated_conjecture, (says(me,and(X1,X2))|a_truth(X2,and(X1,X2))),\\ninference(spm,[status(thm)],[c_0_42, c_0_24])).\n", "fillcolor": "lightskyblue1"}, {"id": 75, "label": "c_0_75", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_75", "role": "negated_conjecture", "formula": "~says(me,and(and(knave(me", "inference_rule": "spm", "raw_label": "cnf(c_0_75, negated_conjecture, (~says(me,and(and(knave(me),rich(me)),X1))),\\ninference(spm,[status(thm)],[c_0_64, c_0_72])).\n", "fillcolor": "lightskyblue1"}, {"id": 76, "label": "c_0_76", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_76", "role": "negated_conjecture", "formula": "says(me,and(X1,and(poor(me", "inference_rule": "spm", "raw_label": "cnf(c_0_76, negated_conjecture, (says(me,and(X1,and(poor(me),X2)))),\\ninference(spm,[status(thm)],[c_0_73, c_0_74])).\n", "fillcolor": "lightskyblue1"}, {"id": 77, "label": "c_0_77", "color": {"background": "#000000", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#FFFFFF"}, "cnf_name": "c_0_77", "role": "negated_conjecture", "formula": "$false", "inference_rule": "spm", "raw_label": "cnf(c_0_77, negated_conjecture, ($false),\\ninference(spm,[status(thm)],[c_0_75, c_0_76])).\n", "fillcolor": "darkorchid1"}], "edges": [{"from": 0, "to": 12, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 1, "to": 13, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 12, "to": 14, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 2, "to": 15, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 3, "to": 16, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 13, "to": 17, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 15, "to": 18, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 14, "to": 18, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 16, "to": 19, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 4, "to": 20, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 18, "to": 21, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 17, "to": 21, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 5, "to": 22, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 20, "to": 23, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 19, "to": 23, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 21, "to": 24, "label": "ef", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 22, "to": 25, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 24, "to": 26, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 23, "to": 26, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 18, "to": 27, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 25, "to": 27, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 26, "to": 28, "label": "ef", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 6, "to": 29, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 24, "to": 30, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 27, "to": 30, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 28, "to": 31, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 25, "to": 31, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 29, "to": 32, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 15, "to": 33, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 30, "to": 33, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 32, "to": 34, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 31, "to": 34, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 33, "to": 35, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 14, "to": 35, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 35, "to": 36, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 34, "to": 36, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 28, "to": 37, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 31, "to": 37, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 7, "to": 38, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 8, "to": 39, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 37, "to": 40, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 15, "to": 40, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 36, "to": 40, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 38, "to": 41, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 39, "to": 42, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 37, "to": 43, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 15, "to": 43, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 40, "to": 43, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 24, "to": 44, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 41, "to": 44, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 15, "to": 45, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 42, "to": 45, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 9, "to": 46, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 43, "to": 47, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 25, "to": 47, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 20, "to": 48, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 30, "to": 48, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 37, "to": 49, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 15, "to": 49, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 31, "to": 49, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 45, "to": 50, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 44, "to": 50, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 46, "to": 51, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 49, "to": 52, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 48, "to": 52, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 47, "to": 52, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 24, "to": 53, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 25, "to": 53, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 26, "to": 54, "label": "ef", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 50, "to": 55, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 27, "to": 55, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 24, "to": 56, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 17, "to": 56, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 52, "to": 57, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 51, "to": 57, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 52, "to": 58, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 19, "to": 58, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 24, "to": 59, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 53, "to": 59, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 10, "to": 60, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 54, "to": 61, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 25, "to": 61, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 56, "to": 62, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 55, "to": 62, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 59, "to": 63, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 58, "to": 63, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 57, "to": 63, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 60, "to": 64, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 63, "to": 65, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 62, "to": 65, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 61, "to": 65, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 11, "to": 66, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 32, "to": 67, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 64, "to": 67, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 66, "to": 68, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 65, "to": 68, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 24, "to": 69, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 14, "to": 69, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 24, "to": 70, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 54, "to": 70, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 68, "to": 70, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 67, "to": 70, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 69, "to": 71, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 55, "to": 71, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 70, "to": 72, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 14, "to": 72, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 71, "to": 73, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 27, "to": 73, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 24, "to": 74, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 42, "to": 74, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 72, "to": 75, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 64, "to": 75, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 74, "to": 76, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 73, "to": 76, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 76, "to": 77, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 75, "to": 77, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}]}, {"problem_file": "PUZ110^5.p", "status": "# SZS status Theorem", "nodes": [{"id": 0, "label": "c1", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": null, "raw_label": "thf(cCKB6_BLACK_def, axiom, ((cCKB6_BLACK)=(^[X1:$i, X2:$i]:(![X3:$i > $i > $o]:((((X3 @ c1 @ c1)&![X4:$i, X5:$i]:(((X3 @ X4 @ X5)=>((X3 @ (s @ (s @ X4)) @ X5)&(X3 @ (s @ X4) @ (s @ X5))))))=>(X3 @ X1 @ X2)))))),\\nfile('Problems/PUZ/PUZ110^5.p', cCKB6_BLACK_def)).\n", "fillcolor": "forestgreen"}, {"id": 1, "label": "1", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": null, "raw_label": "thf(cCKB6_L8000, conjecture, ![X4:$i, X5:$i]:(((cCKB6_BLACK @ X4 @ X5)=>(cCKB6_BLACK @ (s @ (s @ X4)) @ X5))),\\nfile('Problems/PUZ/PUZ110^5.p', cCKB6_L8000)).\n", "fillcolor": "firebrick1"}, {"id": 2, "label": "c1", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "fof_simplification", "raw_label": "thf(c_0_2, plain, ((cCKB6_BLACK)=(^[Z0/* 19 */:$i, Z1:$i]:(![X3:$i > $i > $o]:((((X3 @ c1 @ c1)&![X4:$i, X5:$i]:(((X3 @ X4 @ X5)=>((X3 @ (s @ (s @ X4)) @ X5)&(X3 @ (s @ X4) @ (s @ X5))))))=>(X3 @ Z0 @ Z1)))))),\\ninference(fof_simplification,[status(thm)],[cCKB6_BLACK_def])).\n", "fillcolor": "palegreen"}, {"id": 3, "label": "c1", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "apply_def", "raw_label": "thf(c_0_3, negated_conjecture, ~(![X4:$i, X5:$i]:((![X8:$i > $i > $o]:((((X8 @ c1 @ c1)&![X9:$i, X10:$i]:(((X8 @ X9 @ X10)=>((X8 @ (s @ (s @ X9)) @ X10)&(X8 @ (s @ X9) @ (s @ X10))))))=>(X8 @ X4 @ X5)))=>![X11:$i > $i > $o]:((((X11 @ c1 @ c1)&![X12:$i, X13:$i]:(((X11 @ X12 @ X13)=>((X11 @ (s @ (s @ X12)) @ X13)&(X11 @ (s @ X12) @ (s @ X13))))))=>(X11 @ (s @ (s @ X4)) @ X5)))))),\\ninference(apply_def,[status(thm)],[inference(assume_negation,[status(cth)],[cCKB6_L8000]), c_0_2])).\n", "fillcolor": "lightskyblue1"}, {"id": 4, "label": "c1", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "distribute", "raw_label": "thf(c_0_4, negated_conjecture, ![X16:$i > $i > $o, X20:$i, X21:$i]:(((((X16 @ (esk3_1 @ X16) @ (esk4_1 @ X16))|~(X16 @ c1 @ c1)|(X16 @ esk1_0 @ esk2_0))&(~(X16 @ (s @ (s @ (esk3_1 @ X16))) @ (esk4_1 @ X16))|~(X16 @ (s @ (esk3_1 @ X16)) @ (s @ (esk4_1 @ X16)))|~(X16 @ c1 @ c1)|(X16 @ esk1_0 @ esk2_0)))&(((epred1_0 @ c1 @ c1)&(((epred1_0 @ (s @ (s @ X20)) @ X21)|~(epred1_0 @ X20 @ X21))&((epred1_0 @ (s @ X20) @ (s @ X21))|~(epred1_0 @ X20 @ X21))))&~(epred1_0 @ (s @ (s @ esk1_0)) @ esk2_0)))),\\ninference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])])])).\n", "fillcolor": "lightskyblue1"}, {"id": 5, "label": "5", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "split_conjunct", "raw_label": "thf(c_0_5, negated_conjecture, ~((epred1_0 @ (s @ (s @ esk1_0)) @ esk2_0)),\\ninference(split_conjunct,[status(thm)],[c_0_4])).\n", "fillcolor": "lightskyblue1"}, {"id": 6, "label": "6", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "split_conjunct", "raw_label": "thf(c_0_6, negated_conjecture, ![X1:$i, X2:$i]:(((epred1_0 @ (s @ (s @ X1)) @ X2)|~((epred1_0 @ X1 @ X2)))),\\ninference(split_conjunct,[status(thm)],[c_0_4])).\n", "fillcolor": "lightskyblue1"}, {"id": 7, "label": "c1", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "split_conjunct", "raw_label": "thf(c_0_7, negated_conjecture, ![X3:$i > $i > $o]:(((X3 @ (esk3_1 @ X3) @ (esk4_1 @ X3))|(X3 @ esk1_0 @ esk2_0)|~((X3 @ c1 @ c1)))),\\ninference(split_conjunct,[status(thm)],[c_0_4])).\n", "fillcolor": "lightskyblue1"}, {"id": 8, "label": "c1", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "split_conjunct", "raw_label": "thf(c_0_8, negated_conjecture, (epred1_0 @ c1 @ c1),\\ninference(split_conjunct,[status(thm)],[c_0_4])).\n", "fillcolor": "lightskyblue1"}, {"id": 9, "label": "9", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "spm", "raw_label": "thf(c_0_9, negated_conjecture, ~((epred1_0 @ esk1_0 @ esk2_0)),\\ninference(spm,[status(thm)],[c_0_5, c_0_6])).\n", "fillcolor": "lightskyblue1"}, {"id": 10, "label": "c1", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "split_conjunct", "raw_label": "thf(c_0_10, negated_conjecture, ![X3:$i > $i > $o]:(((X3 @ esk1_0 @ esk2_0)|~((X3 @ (s @ (s @ (esk3_1 @ X3))) @ (esk4_1 @ X3)))|~((X3 @ (s @ (esk3_1 @ X3)) @ (s @ (esk4_1 @ X3))))|~((X3 @ c1 @ c1)))),\\ninference(split_conjunct,[status(thm)],[c_0_4])).\n", "fillcolor": "lightskyblue1"}, {"id": 11, "label": "11", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "split_conjunct", "raw_label": "thf(c_0_11, negated_conjecture, ![X1:$i, X2:$i]:(((epred1_0 @ (s @ X1) @ (s @ X2))|~((epred1_0 @ X1 @ X2)))),\\ninference(split_conjunct,[status(thm)],[c_0_4])).\n", "fillcolor": "lightskyblue1"}, {"id": 12, "label": "12", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "sr", "raw_label": "thf(c_0_12, negated_conjecture, (epred1_0 @ (esk3_1 @ epred1_0) @ (esk4_1 @ epred1_0)),\\ninference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_7, c_0_8]), c_0_9])).\n", "fillcolor": "lightskyblue1"}, {"id": 13, "label": "13", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "sr", "raw_label": "thf(c_0_13, negated_conjecture, ~((epred1_0 @ (s @ (s @ (esk3_1 @ epred1_0))) @ (esk4_1 @ epred1_0))),\\ninference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_11]), c_0_8]), c_0_12])]), c_0_9])).\n", "fillcolor": "lightskyblue1"}, {"id": 14, "label": "14", "color": {"background": "#000000", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#FFFFFF"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "cn", "raw_label": "thf(c_0_14, negated_conjecture, ($false),\\ninference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_6]), c_0_12])])).\n", "fillcolor": "darkorchid1"}], "edges": [{"from": 0, "to": 2, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 2, "to": 3, "label": "apply_def", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 1, "to": 3, "label": "apply_def", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 3, "to": 4, "label": "distribute", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 4, "to": 5, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 4, "to": 6, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 4, "to": 7, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 4, "to": 8, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 6, "to": 9, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 5, "to": 9, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 4, "to": 10, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 4, "to": 11, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 9, "to": 12, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 8, "to": 12, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 7, "to": 12, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 9, "to": 13, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 12, "to": 13, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 8, "to": 13, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 11, "to": 13, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 10, "to": 13, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 12, "to": 14, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 6, "to": 14, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 13, "to": 14, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}]}, {"problem_file": "PUZ003-1.p", "status": "# SZS status Unsatisfiable", "nodes": [{"id": 0, "label": "one_shaved_then_al", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "one_shaved_then_all_shaved", "role": "axiom", "formula": "shaved(members,X1)|~member(X1)|~member(X2)|~shaved(X1,X2)", "inference_rule": null, "raw_label": "cnf(one_shaved_then_all_shaved, axiom, (shaved(members,X1)|~member(X1)|~member(X2)|~shaved(X1,X2)),\\nfile('Problems/PUZ/PUZ003-1.p', one_shaved_then_all_shaved)).\n", "fillcolor": "forestgreen"}, {"id": 1, "label": "all_shaved_then_on", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "all_shaved_then_one_shaved", "role": "axiom", "formula": "shaved(X2,X1)|~shaved(members,X1)|~member(X2)", "inference_rule": null, "raw_label": "cnf(all_shaved_then_one_shaved, axiom, (shaved(X2,X1)|~shaved(members,X1)|~member(X2)),\\nfile('Problems/PUZ/PUZ003-1.p', all_shaved_then_one_shaved)).\n", "fillcolor": "forestgreen"}, {"id": 2, "label": "guido_has_shaved_c", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "guido_has_shaved_cesare", "role": "hypothesis", "formula": "shaved(guido,cesare)", "inference_rule": null, "raw_label": "cnf(guido_has_shaved_cesare, hypothesis, (shaved(guido,cesare)),\\nfile('Problems/PUZ/PUZ003-1.p', guido_has_shaved_cesare)).\n", "fillcolor": "forestgreen"}, {"id": 3, "label": "cesare", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "cesare", "role": "hypothesis", "formula": "member(cesare)", "inference_rule": null, "raw_label": "cnf(cesare, hypothesis, (member(cesare)),\\nfile('Problems/PUZ/PUZ003-1.p', cesare)).\n", "fillcolor": "forestgreen"}, {"id": 4, "label": "guido", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "guido", "role": "hypothesis", "formula": "member(guido)", "inference_rule": null, "raw_label": "cnf(guido, hypothesis, (member(guido)),\\nfile('Problems/PUZ/PUZ003-1.p', guido)).\n", "fillcolor": "forestgreen"}, {"id": 5, "label": "prove_petruchio_ha", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "prove_petruchio_has_shaved_lorenzo", "role": "negated_conjecture", "formula": "~shaved(petruchio,lorenzo)", "inference_rule": null, "raw_label": "cnf(prove_petruchio_has_shaved_lorenzo, negated_conjecture, (~shaved(petruchio,lorenzo)),\\nfile('Problems/PUZ/PUZ003-1.p', prove_petruchio_has_shaved_lorenzo)).\n", "fillcolor": "dodgerblue"}, {"id": 6, "label": "petruchio", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "petruchio", "role": "hypothesis", "formula": "member(petruchio)", "inference_rule": null, "raw_label": "cnf(petruchio, hypothesis, (member(petruchio)),\\nfile('Problems/PUZ/PUZ003-1.p', petruchio)).\n", "fillcolor": "forestgreen"}, {"id": 7, "label": "lorenzo", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "lorenzo", "role": "hypothesis", "formula": "member(lorenzo)", "inference_rule": null, "raw_label": "cnf(lorenzo, hypothesis, (member(lorenzo)),\\nfile('Problems/PUZ/PUZ003-1.p', lorenzo)).\n", "fillcolor": "forestgreen"}, {"id": 8, "label": "c_0_8", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_8", "role": "plain", "formula": "shaved(members,X1)|~member(X1)|~member(X2)|~shaved(X1,X2)", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_8, plain, (shaved(members,X1)|~member(X1)|~member(X2)|~shaved(X1,X2)),\\ninference(fof_simplification,[status(thm)],[one_shaved_then_all_shaved])).\n", "fillcolor": "palegreen"}, {"id": 9, "label": "c_0_9", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_9", "role": "plain", "formula": "shaved(X2,X1)|~shaved(members,X1)|~member(X2)", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_9, plain, (shaved(X2,X1)|~shaved(members,X1)|~member(X2)),\\ninference(fof_simplification,[status(thm)],[all_shaved_then_one_shaved])).\n", "fillcolor": "palegreen"}, {"id": 10, "label": "c_0_10", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_10", "role": "plain", "formula": "shaved(members,X1)|~member(X1)|~member(X2)|~shaved(X1,X2)", "inference_rule": null, "raw_label": "cnf(c_0_10, plain, (shaved(members,X1)|~member(X1)|~member(X2)|~shaved(X1,X2)),\\nc_0_8).\n", "fillcolor": "palegreen"}, {"id": 11, "label": "c_0_11", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_11", "role": "hypothesis", "formula": "shaved(guido,cesare)", "inference_rule": null, "raw_label": "cnf(c_0_11, hypothesis, (shaved(guido,cesare)),\\nguido_has_shaved_cesare).\n", "fillcolor": "palegreen"}, {"id": 12, "label": "c_0_12", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_12", "role": "hypothesis", "formula": "member(cesare)", "inference_rule": null, "raw_label": "cnf(c_0_12, hypothesis, (member(cesare)),\\ncesare).\n", "fillcolor": "palegreen"}, {"id": 13, "label": "c_0_13", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_13", "role": "hypothesis", "formula": "member(guido)", "inference_rule": null, "raw_label": "cnf(c_0_13, hypothesis, (member(guido)),\\nguido).\n", "fillcolor": "palegreen"}, {"id": 14, "label": "c_0_14", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_14", "role": "plain", "formula": "shaved(X2,X1)|~shaved(members,X1)|~member(X2)", "inference_rule": null, "raw_label": "cnf(c_0_14, plain, (shaved(X2,X1)|~shaved(members,X1)|~member(X2)),\\nc_0_9).\n", "fillcolor": "palegreen"}, {"id": 15, "label": "c_0_15", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_15", "role": "hypothesis", "formula": "shaved(members,guido)", "inference_rule": "cn", "raw_label": "cnf(c_0_15, hypothesis, (shaved(members,guido)),\\ninference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_11]), c_0_12]), c_0_13])])).\n", "fillcolor": "palegreen"}, {"id": 16, "label": "c_0_16", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_16", "role": "hypothesis", "formula": "shaved(X1,guido)|~member(X1)", "inference_rule": "spm", "raw_label": "cnf(c_0_16, hypothesis, (shaved(X1,guido)|~member(X1)),\\ninference(spm,[status(thm)],[c_0_14, c_0_15])).\n", "fillcolor": "palegreen"}, {"id": 17, "label": "c_0_17", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_17", "role": "negated_conjecture", "formula": "~shaved(petruchio,lorenzo)", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_17, negated_conjecture, (~shaved(petruchio,lorenzo)),\\ninference(fof_simplification,[status(thm)],[prove_petruchio_has_shaved_lorenzo])).\n", "fillcolor": "lightskyblue1"}, {"id": 18, "label": "c_0_18", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_18", "role": "hypothesis", "formula": "shaved(members,X1)|~member(X1)", "inference_rule": "cn", "raw_label": "cnf(c_0_18, hypothesis, (shaved(members,X1)|~member(X1)),\\ninference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_16]), c_0_13])])).\n", "fillcolor": "palegreen"}, {"id": 19, "label": "c_0_19", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_19", "role": "negated_conjecture", "formula": "~shaved(petruchio,lorenzo)", "inference_rule": null, "raw_label": "cnf(c_0_19, negated_conjecture, (~shaved(petruchio,lorenzo)),\\nc_0_17).\n", "fillcolor": "lightskyblue1"}, {"id": 20, "label": "c_0_20", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_20", "role": "hypothesis", "formula": "shaved(X1,X2)|~member(X1)|~member(X2)", "inference_rule": "spm", "raw_label": "cnf(c_0_20, hypothesis, (shaved(X1,X2)|~member(X1)|~member(X2)),\\ninference(spm,[status(thm)],[c_0_14, c_0_18])).\n", "fillcolor": "palegreen"}, {"id": 21, "label": "c_0_21", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_21", "role": "hypothesis", "formula": "member(petruchio)", "inference_rule": null, "raw_label": "cnf(c_0_21, hypothesis, (member(petruchio)),\\npetruchio).\n", "fillcolor": "palegreen"}, {"id": 22, "label": "c_0_22", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_22", "role": "hypothesis", "formula": "member(lorenzo)", "inference_rule": null, "raw_label": "cnf(c_0_22, hypothesis, (member(lorenzo)),\\nlorenzo).\n", "fillcolor": "palegreen"}, {"id": 23, "label": "c_0_23", "color": {"background": "#000000", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#FFFFFF"}, "cnf_name": "c_0_23", "role": "negated_conjecture", "formula": "$false", "inference_rule": "cn", "raw_label": "cnf(c_0_23, negated_conjecture, ($false),\\ninference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_21]), c_0_22])])).\n", "fillcolor": "darkorchid1"}], "edges": [{"from": 0, "to": 8, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 1, "to": 9, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 8, "to": 10, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 2, "to": 11, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 3, "to": 12, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 4, "to": 13, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 9, "to": 14, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 13, "to": 15, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 12, "to": 15, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 11, "to": 15, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 10, "to": 15, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 15, "to": 16, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 14, "to": 16, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 5, "to": 17, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 13, "to": 18, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 16, "to": 18, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 10, "to": 18, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 17, "to": 19, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 18, "to": 20, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 14, "to": 20, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 6, "to": 21, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 7, "to": 22, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 22, "to": 23, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 21, "to": 23, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 20, "to": 23, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 19, "to": 23, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}]}, {"problem_file": "PUZ005-1.p", "status": "# SZS status Unsatisfiable", "nodes": [{"id": 0, "label": "admissions3", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "admissions3", "role": "axiom", "formula": "~member(X1,lying_days(X2))|~admits(X2,X1,X3)|~member(X3,lying_days(X2))", "inference_rule": null, "raw_label": "cnf(admissions3, axiom, (~member(X1,lying_days(X2))|~admits(X2,X1,X3)|~member(X3,lying_days(X2))),\\nfile('Problems/PUZ/PUZ005-1.p', admissions3)).\n", "fillcolor": "forestgreen"}, {"id": 1, "label": "admissions1", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "admissions1", "role": "axiom", "formula": "member(X1,lying_days(X2))|member(X3,lying_days(X2))|~admits(X2,X1,X3)", "inference_rule": null, "raw_label": "cnf(admissions1, axiom, (member(X1,lying_days(X2))|member(X3,lying_days(X2))|~admits(X2,X1,X3)),\\nfile('Problems/PUZ/PUZ005-1.p', admissions1)).\n", "fillcolor": "forestgreen"}, {"id": 2, "label": "admissions6", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "admissions6", "role": "axiom", "formula": "admits(unicorn,today,yesterday(today))", "inference_rule": null, "raw_label": "cnf(admissions6, axiom, (admits(unicorn,today,yesterday(today))),\\nfile('Problems/PUZ/PUZ005-1.p', admissions6)).\n", "fillcolor": "forestgreen"}, {"id": 3, "label": "unicorn_lies_on_fr", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "unicorn_lies_on_friday", "role": "axiom", "formula": "member(X1,lying_days(unicorn))|~friday(X1)", "inference_rule": null, "raw_label": "cnf(unicorn_lies_on_friday, axiom, (member(X1,lying_days(unicorn))|~friday(X1)),\\nfile('Problems/PUZ/PUZ005-1.p', unicorn_lies_on_friday)).\n", "fillcolor": "forestgreen"}, {"id": 4, "label": "unicorn_lies_on_th", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "unicorn_lies_on_thursday", "role": "axiom", "formula": "member(X1,lying_days(unicorn))|~thursday(X1)", "inference_rule": null, "raw_label": "cnf(unicorn_lies_on_thursday, axiom, (member(X1,lying_days(unicorn))|~thursday(X1)),\\nfile('Problems/PUZ/PUZ005-1.p', unicorn_lies_on_thursday)).\n", "fillcolor": "forestgreen"}, {"id": 5, "label": "unicorns_lying_day", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "unicorns_lying_days", "role": "axiom", "formula": "thursday(X1)|friday(X1)|saturday(X1)|~member(X1,lying_days(unicorn))", "inference_rule": null, "raw_label": "cnf(unicorns_lying_days, axiom, (thursday(X1)|friday(X1)|saturday(X1)|~member(X1,lying_days(unicorn))),\\nfile('Problems/PUZ/PUZ005-1.p', unicorns_lying_days)).\n", "fillcolor": "forestgreen"}, {"id": 6, "label": "unicorn_lies_on_sa", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "unicorn_lies_on_saturday", "role": "axiom", "formula": "member(X1,lying_days(unicorn))|~saturday(X1)", "inference_rule": null, "raw_label": "cnf(unicorn_lies_on_saturday, axiom, (member(X1,lying_days(unicorn))|~saturday(X1)),\\nfile('Problems/PUZ/PUZ005-1.p', unicorn_lies_on_saturday)).\n", "fillcolor": "forestgreen"}, {"id": 7, "label": "yesterday_friday", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "yesterday_friday", "role": "axiom", "formula": "friday(yesterday(X1))|~saturday(X1)", "inference_rule": null, "raw_label": "cnf(yesterday_friday, axiom, (friday(yesterday(X1))|~saturday(X1)),\\nfile('Problems/PUZ/PUZ005-1.p', yesterday_friday)).\n", "fillcolor": "forestgreen"}, {"id": 8, "label": "yesterday_thursday", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "yesterday_thursday", "role": "axiom", "formula": "thursday(yesterday(X1))|~friday(X1)", "inference_rule": null, "raw_label": "cnf(yesterday_thursday, axiom, (thursday(yesterday(X1))|~friday(X1)),\\nfile('Problems/PUZ/PUZ005-1.p', yesterday_thursday)).\n", "fillcolor": "forestgreen"}, {"id": 9, "label": "prove_today_is_thu", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "prove_today_is_thursday", "role": "negated_conjecture", "formula": "~thursday(today)", "inference_rule": null, "raw_label": "cnf(prove_today_is_thursday, negated_conjecture, (~thursday(today)),\\nfile('Problems/PUZ/PUZ005-1.p', prove_today_is_thursday)).\n", "fillcolor": "dodgerblue"}, {"id": 10, "label": "wednesday_not_satu", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "wednesday_not_saturday", "role": "axiom", "formula": "~wednesday(X1)|~saturday(X1)", "inference_rule": null, "raw_label": "cnf(wednesday_not_saturday, axiom, (~wednesday(X1)|~saturday(X1)),\\nfile('Problems/PUZ/PUZ005-1.p', wednesday_not_saturday)).\n", "fillcolor": "forestgreen"}, {"id": 11, "label": "wednesday_not_thur", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "wednesday_not_thursday", "role": "axiom", "formula": "~wednesday(X1)|~thursday(X1)", "inference_rule": null, "raw_label": "cnf(wednesday_not_thursday, axiom, (~wednesday(X1)|~thursday(X1)),\\nfile('Problems/PUZ/PUZ005-1.p', wednesday_not_thursday)).\n", "fillcolor": "forestgreen"}, {"id": 12, "label": "wednesday_not_frid", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "wednesday_not_friday", "role": "axiom", "formula": "~wednesday(X1)|~friday(X1)", "inference_rule": null, "raw_label": "cnf(wednesday_not_friday, axiom, (~wednesday(X1)|~friday(X1)),\\nfile('Problems/PUZ/PUZ005-1.p', wednesday_not_friday)).\n", "fillcolor": "forestgreen"}, {"id": 13, "label": "tuesday_not_saturd", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "tuesday_not_saturday", "role": "axiom", "formula": "~tuesday(X1)|~saturday(X1)", "inference_rule": null, "raw_label": "cnf(tuesday_not_saturday, axiom, (~tuesday(X1)|~saturday(X1)),\\nfile('Problems/PUZ/PUZ005-1.p', tuesday_not_saturday)).\n", "fillcolor": "forestgreen"}, {"id": 14, "label": "tuesday_not_thursd", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "tuesday_not_thursday", "role": "axiom", "formula": "~tuesday(X1)|~thursday(X1)", "inference_rule": null, "raw_label": "cnf(tuesday_not_thursday, axiom, (~tuesday(X1)|~thursday(X1)),\\nfile('Problems/PUZ/PUZ005-1.p', tuesday_not_thursday)).\n", "fillcolor": "forestgreen"}, {"id": 15, "label": "tuesday_not_friday", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "tuesday_not_friday", "role": "axiom", "formula": "~tuesday(X1)|~friday(X1)", "inference_rule": null, "raw_label": "cnf(tuesday_not_friday, axiom, (~tuesday(X1)|~friday(X1)),\\nfile('Problems/PUZ/PUZ005-1.p', tuesday_not_friday)).\n", "fillcolor": "forestgreen"}, {"id": 16, "label": "monday_not_saturda", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "monday_not_saturday", "role": "axiom", "formula": "~monday(X1)|~saturday(X1)", "inference_rule": null, "raw_label": "cnf(monday_not_saturday, axiom, (~monday(X1)|~saturday(X1)),\\nfile('Problems/PUZ/PUZ005-1.p', monday_not_saturday)).\n", "fillcolor": "forestgreen"}, {"id": 17, "label": "monday_not_thursda", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "monday_not_thursday", "role": "axiom", "formula": "~monday(X1)|~thursday(X1)", "inference_rule": null, "raw_label": "cnf(monday_not_thursday, axiom, (~monday(X1)|~thursday(X1)),\\nfile('Problems/PUZ/PUZ005-1.p', monday_not_thursday)).\n", "fillcolor": "forestgreen"}, {"id": 18, "label": "monday_not_friday", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "monday_not_friday", "role": "axiom", "formula": "~monday(X1)|~friday(X1)", "inference_rule": null, "raw_label": "cnf(monday_not_friday, axiom, (~monday(X1)|~friday(X1)),\\nfile('Problems/PUZ/PUZ005-1.p', monday_not_friday)).\n", "fillcolor": "forestgreen"}, {"id": 19, "label": "admissions5", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "admissions5", "role": "axiom", "formula": "admits(lion,today,yesterday(today))", "inference_rule": null, "raw_label": "cnf(admissions5, axiom, (admits(lion,today,yesterday(today))),\\nfile('Problems/PUZ/PUZ005-1.p', admissions5)).\n", "fillcolor": "forestgreen"}, {"id": 20, "label": "lion_lies_on_tuesd", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "lion_lies_on_tuesday", "role": "axiom", "formula": "member(X1,lying_days(lion))|~tuesday(X1)", "inference_rule": null, "raw_label": "cnf(lion_lies_on_tuesday, axiom, (member(X1,lying_days(lion))|~tuesday(X1)),\\nfile('Problems/PUZ/PUZ005-1.p', lion_lies_on_tuesday)).\n", "fillcolor": "forestgreen"}, {"id": 21, "label": "lion_lies_on_monda", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "lion_lies_on_monday", "role": "axiom", "formula": "member(X1,lying_days(lion))|~monday(X1)", "inference_rule": null, "raw_label": "cnf(lion_lies_on_monday, axiom, (member(X1,lying_days(lion))|~monday(X1)),\\nfile('Problems/PUZ/PUZ005-1.p', lion_lies_on_monday)).\n", "fillcolor": "forestgreen"}, {"id": 22, "label": "saturday_not_sunda", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "saturday_not_sunday", "role": "axiom", "formula": "~saturday(X1)|~sunday(X1)", "inference_rule": null, "raw_label": "cnf(saturday_not_sunday, axiom, (~saturday(X1)|~sunday(X1)),\\nfile('Problems/PUZ/PUZ005-1.p', saturday_not_sunday)).\n", "fillcolor": "forestgreen"}, {"id": 23, "label": "yesterday_sunday", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "yesterday_sunday", "role": "axiom", "formula": "sunday(yesterday(X1))|~monday(X1)", "inference_rule": null, "raw_label": "cnf(yesterday_sunday, axiom, (sunday(yesterday(X1))|~monday(X1)),\\nfile('Problems/PUZ/PUZ005-1.p', yesterday_sunday)).\n", "fillcolor": "forestgreen"}, {"id": 24, "label": "thursday_not_sunda", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "thursday_not_sunday", "role": "axiom", "formula": "~thursday(X1)|~sunday(X1)", "inference_rule": null, "raw_label": "cnf(thursday_not_sunday, axiom, (~thursday(X1)|~sunday(X1)),\\nfile('Problems/PUZ/PUZ005-1.p', thursday_not_sunday)).\n", "fillcolor": "forestgreen"}, {"id": 25, "label": "friday_not_sunday", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "friday_not_sunday", "role": "axiom", "formula": "~friday(X1)|~sunday(X1)", "inference_rule": null, "raw_label": "cnf(friday_not_sunday, axiom, (~friday(X1)|~sunday(X1)),\\nfile('Problems/PUZ/PUZ005-1.p', friday_not_sunday)).\n", "fillcolor": "forestgreen"}, {"id": 26, "label": "lions_lying_days", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "lions_lying_days", "role": "axiom", "formula": "monday(X1)|tuesday(X1)|wednesday(X1)|~member(X1,lying_days(lion))", "inference_rule": null, "raw_label": "cnf(lions_lying_days, axiom, (monday(X1)|tuesday(X1)|wednesday(X1)|~member(X1,lying_days(lion))),\\nfile('Problems/PUZ/PUZ005-1.p', lions_lying_days)).\n", "fillcolor": "forestgreen"}, {"id": 27, "label": "lion_lies_on_wedne", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "lion_lies_on_wednesday", "role": "axiom", "formula": "member(X1,lying_days(lion))|~wednesday(X1)", "inference_rule": null, "raw_label": "cnf(lion_lies_on_wednesday, axiom, (member(X1,lying_days(lion))|~wednesday(X1)),\\nfile('Problems/PUZ/PUZ005-1.p', lion_lies_on_wednesday)).\n", "fillcolor": "forestgreen"}, {"id": 28, "label": "yesterday_tuesday", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "yesterday_tuesday", "role": "axiom", "formula": "tuesday(yesterday(X1))|~wednesday(X1)", "inference_rule": null, "raw_label": "cnf(yesterday_tuesday, axiom, (tuesday(yesterday(X1))|~wednesday(X1)),\\nfile('Problems/PUZ/PUZ005-1.p', yesterday_tuesday)).\n", "fillcolor": "forestgreen"}, {"id": 29, "label": "yesterday_monday", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "yesterday_monday", "role": "axiom", "formula": "monday(yesterday(X1))|~tuesday(X1)", "inference_rule": null, "raw_label": "cnf(yesterday_monday, axiom, (monday(yesterday(X1))|~tuesday(X1)),\\nfile('Problems/PUZ/PUZ005-1.p', yesterday_monday)).\n", "fillcolor": "forestgreen"}, {"id": 30, "label": "c_0_30", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_30", "role": "plain", "formula": "~member(X1,lying_days(X2))|~admits(X2,X1,X3)|~member(X3,lying_days(X2))", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_30, plain, (~member(X1,lying_days(X2))|~admits(X2,X1,X3)|~member(X3,lying_days(X2))),\\ninference(fof_simplification,[status(thm)],[admissions3])).\n", "fillcolor": "palegreen"}, {"id": 31, "label": "c_0_31", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_31", "role": "plain", "formula": "member(X1,lying_days(X2))|member(X3,lying_days(X2))|~admits(X2,X1,X3)", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_31, plain, (member(X1,lying_days(X2))|member(X3,lying_days(X2))|~admits(X2,X1,X3)),\\ninference(fof_simplification,[status(thm)],[admissions1])).\n", "fillcolor": "palegreen"}, {"id": 32, "label": "c_0_32", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_32", "role": "plain", "formula": "~member(X1,lying_days(X2))|~admits(X2,X1,X3)|~member(X3,lying_days(X2))", "inference_rule": null, "raw_label": "cnf(c_0_32, plain, (~member(X1,lying_days(X2))|~admits(X2,X1,X3)|~member(X3,lying_days(X2))),\\nc_0_30).\n", "fillcolor": "palegreen"}, {"id": 33, "label": "c_0_33", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_33", "role": "axiom", "formula": "admits(unicorn,today,yesterday(today))", "inference_rule": null, "raw_label": "cnf(c_0_33, axiom, (admits(unicorn,today,yesterday(today))),\\nadmissions6).\n", "fillcolor": "palegreen"}, {"id": 34, "label": "c_0_34", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_34", "role": "plain", "formula": "member(X1,lying_days(unicorn))|~friday(X1)", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_34, plain, (member(X1,lying_days(unicorn))|~friday(X1)),\\ninference(fof_simplification,[status(thm)],[unicorn_lies_on_friday])).\n", "fillcolor": "palegreen"}, {"id": 35, "label": "c_0_35", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_35", "role": "plain", "formula": "member(X1,lying_days(unicorn))|~thursday(X1)", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_35, plain, (member(X1,lying_days(unicorn))|~thursday(X1)),\\ninference(fof_simplification,[status(thm)],[unicorn_lies_on_thursday])).\n", "fillcolor": "palegreen"}, {"id": 36, "label": "c_0_36", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_36", "role": "plain", "formula": "thursday(X1)|friday(X1)|saturday(X1)|~member(X1,lying_days(unicorn))", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_36, plain, (thursday(X1)|friday(X1)|saturday(X1)|~member(X1,lying_days(unicorn))),\\ninference(fof_simplification,[status(thm)],[unicorns_lying_days])).\n", "fillcolor": "palegreen"}, {"id": 37, "label": "c_0_37", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_37", "role": "plain", "formula": "member(X1,lying_days(X2))|member(X3,lying_days(X2))|~admits(X2,X1,X3)", "inference_rule": null, "raw_label": "cnf(c_0_37, plain, (member(X1,lying_days(X2))|member(X3,lying_days(X2))|~admits(X2,X1,X3)),\\nc_0_31).\n", "fillcolor": "palegreen"}, {"id": 38, "label": "c_0_38", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_38", "role": "plain", "formula": "~member(yesterday(today", "inference_rule": "spm", "raw_label": "cnf(c_0_38, plain, (~member(yesterday(today),lying_days(unicorn))|~member(today,lying_days(unicorn))),\\ninference(spm,[status(thm)],[c_0_32, c_0_33])).\n", "fillcolor": "palegreen"}, {"id": 39, "label": "c_0_39", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_39", "role": "plain", "formula": "member(X1,lying_days(unicorn))|~friday(X1)", "inference_rule": null, "raw_label": "cnf(c_0_39, plain, (member(X1,lying_days(unicorn))|~friday(X1)),\\nc_0_34).\n", "fillcolor": "palegreen"}, {"id": 40, "label": "c_0_40", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_40", "role": "plain", "formula": "member(X1,lying_days(unicorn))|~saturday(X1)", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_40, plain, (member(X1,lying_days(unicorn))|~saturday(X1)),\\ninference(fof_simplification,[status(thm)],[unicorn_lies_on_saturday])).\n", "fillcolor": "palegreen"}, {"id": 41, "label": "c_0_41", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_41", "role": "plain", "formula": "friday(yesterday(X1))|~saturday(X1)", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_41, plain, (friday(yesterday(X1))|~saturday(X1)),\\ninference(fof_simplification,[status(thm)],[yesterday_friday])).\n", "fillcolor": "palegreen"}, {"id": 42, "label": "c_0_42", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_42", "role": "plain", "formula": "member(X1,lying_days(unicorn))|~thursday(X1)", "inference_rule": null, "raw_label": "cnf(c_0_42, plain, (member(X1,lying_days(unicorn))|~thursday(X1)),\\nc_0_35).\n", "fillcolor": "palegreen"}, {"id": 43, "label": "c_0_43", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_43", "role": "plain", "formula": "thursday(yesterday(X1))|~friday(X1)", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_43, plain, (thursday(yesterday(X1))|~friday(X1)),\\ninference(fof_simplification,[status(thm)],[yesterday_thursday])).\n", "fillcolor": "palegreen"}, {"id": 44, "label": "c_0_44", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_44", "role": "plain", "formula": "thursday(X1)|friday(X1)|saturday(X1)|~member(X1,lying_days(unicorn))", "inference_rule": null, "raw_label": "cnf(c_0_44, plain, (thursday(X1)|friday(X1)|saturday(X1)|~member(X1,lying_days(unicorn))),\\nc_0_36).\n", "fillcolor": "palegreen"}, {"id": 45, "label": "c_0_45", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_45", "role": "plain", "formula": "member(yesterday(today", "inference_rule": "spm", "raw_label": "cnf(c_0_45, plain, (member(yesterday(today),lying_days(unicorn))|member(today,lying_days(unicorn))),\\ninference(spm,[status(thm)],[c_0_37, c_0_33])).\n", "fillcolor": "palegreen"}, {"id": 46, "label": "c_0_46", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_46", "role": "plain", "formula": "~member(today,lying_days(unicorn))|~friday(yesterday(today))", "inference_rule": "spm", "raw_label": "cnf(c_0_46, plain, (~member(today,lying_days(unicorn))|~friday(yesterday(today))),\\ninference(spm,[status(thm)],[c_0_38, c_0_39])).\n", "fillcolor": "palegreen"}, {"id": 47, "label": "c_0_47", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_47", "role": "plain", "formula": "member(X1,lying_days(unicorn))|~saturday(X1)", "inference_rule": null, "raw_label": "cnf(c_0_47, plain, (member(X1,lying_days(unicorn))|~saturday(X1)),\\nc_0_40).\n", "fillcolor": "palegreen"}, {"id": 48, "label": "c_0_48", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_48", "role": "plain", "formula": "friday(yesterday(X1))|~saturday(X1)", "inference_rule": null, "raw_label": "cnf(c_0_48, plain, (friday(yesterday(X1))|~saturday(X1)),\\nc_0_41).\n", "fillcolor": "palegreen"}, {"id": 49, "label": "c_0_49", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_49", "role": "plain", "formula": "~member(today,lying_days(unicorn))|~thursday(yesterday(today))", "inference_rule": "spm", "raw_label": "cnf(c_0_49, plain, (~member(today,lying_days(unicorn))|~thursday(yesterday(today))),\\ninference(spm,[status(thm)],[c_0_38, c_0_42])).\n", "fillcolor": "palegreen"}, {"id": 50, "label": "c_0_50", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_50", "role": "plain", "formula": "thursday(yesterday(X1))|~friday(X1)", "inference_rule": null, "raw_label": "cnf(c_0_50, plain, (thursday(yesterday(X1))|~friday(X1)),\\nc_0_43).\n", "fillcolor": "palegreen"}, {"id": 51, "label": "c_0_51", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_51", "role": "negated_conjecture", "formula": "~thursday(today)", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_51, negated_conjecture, (~thursday(today)),\\ninference(fof_simplification,[status(thm)],[prove_today_is_thursday])).\n", "fillcolor": "lightskyblue1"}, {"id": 52, "label": "c_0_52", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_52", "role": "plain", "formula": "~wednesday(X1)|~saturday(X1)", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_52, plain, (~wednesday(X1)|~saturday(X1)),\\ninference(fof_simplification,[status(thm)],[wednesday_not_saturday])).\n", "fillcolor": "palegreen"}, {"id": 53, "label": "c_0_53", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_53", "role": "plain", "formula": "member(today,lying_days(unicorn))|saturday(yesterday(today))|friday(yesterday(today))|thursday(yesterday(today))", "inference_rule": "spm", "raw_label": "cnf(c_0_53, plain, (member(today,lying_days(unicorn))|saturday(yesterday(today))|friday(yesterday(today))|thursday(yesterday(today))),\\ninference(spm,[status(thm)],[c_0_44, c_0_45])).\n", "fillcolor": "palegreen"}, {"id": 54, "label": "c_0_54", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_54", "role": "plain", "formula": "~saturday(today)", "inference_rule": "csr", "raw_label": "cnf(c_0_54, plain, (~saturday(today)),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_48])).\n", "fillcolor": "palegreen"}, {"id": 55, "label": "c_0_55", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_55", "role": "plain", "formula": "~friday(today)", "inference_rule": "csr", "raw_label": "cnf(c_0_55, plain, (~friday(today)),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_39]), c_0_50])).\n", "fillcolor": "palegreen"}, {"id": 56, "label": "c_0_56", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_56", "role": "negated_conjecture", "formula": "~thursday(today)", "inference_rule": null, "raw_label": "cnf(c_0_56, negated_conjecture, (~thursday(today)),\\nc_0_51).\n", "fillcolor": "lightskyblue1"}, {"id": 57, "label": "c_0_57", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_57", "role": "plain", "formula": "~wednesday(X1)|~thursday(X1)", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_57, plain, (~wednesday(X1)|~thursday(X1)),\\ninference(fof_simplification,[status(thm)],[wednesday_not_thursday])).\n", "fillcolor": "palegreen"}, {"id": 58, "label": "c_0_58", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_58", "role": "plain", "formula": "~wednesday(X1)|~friday(X1)", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_58, plain, (~wednesday(X1)|~friday(X1)),\\ninference(fof_simplification,[status(thm)],[wednesday_not_friday])).\n", "fillcolor": "palegreen"}, {"id": 59, "label": "c_0_59", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_59", "role": "plain", "formula": "~tuesday(X1)|~saturday(X1)", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_59, plain, (~tuesday(X1)|~saturday(X1)),\\ninference(fof_simplification,[status(thm)],[tuesday_not_saturday])).\n", "fillcolor": "palegreen"}, {"id": 60, "label": "c_0_60", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_60", "role": "plain", "formula": "~tuesday(X1)|~thursday(X1)", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_60, plain, (~tuesday(X1)|~thursday(X1)),\\ninference(fof_simplification,[status(thm)],[tuesday_not_thursday])).\n", "fillcolor": "palegreen"}, {"id": 61, "label": "c_0_61", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_61", "role": "plain", "formula": "~tuesday(X1)|~friday(X1)", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_61, plain, (~tuesday(X1)|~friday(X1)),\\ninference(fof_simplification,[status(thm)],[tuesday_not_friday])).\n", "fillcolor": "palegreen"}, {"id": 62, "label": "c_0_62", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_62", "role": "plain", "formula": "~monday(X1)|~saturday(X1)", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_62, plain, (~monday(X1)|~saturday(X1)),\\ninference(fof_simplification,[status(thm)],[monday_not_saturday])).\n", "fillcolor": "palegreen"}, {"id": 63, "label": "c_0_63", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_63", "role": "plain", "formula": "~monday(X1)|~thursday(X1)", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_63, plain, (~monday(X1)|~thursday(X1)),\\ninference(fof_simplification,[status(thm)],[monday_not_thursday])).\n", "fillcolor": "palegreen"}, {"id": 64, "label": "c_0_64", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_64", "role": "plain", "formula": "~monday(X1)|~friday(X1)", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_64, plain, (~monday(X1)|~friday(X1)),\\ninference(fof_simplification,[status(thm)],[monday_not_friday])).\n", "fillcolor": "palegreen"}, {"id": 65, "label": "c_0_65", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_65", "role": "axiom", "formula": "admits(lion,today,yesterday(today))", "inference_rule": null, "raw_label": "cnf(c_0_65, axiom, (admits(lion,today,yesterday(today))),\\nadmissions5).\n", "fillcolor": "palegreen"}, {"id": 66, "label": "c_0_66", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_66", "role": "plain", "formula": "member(X1,lying_days(lion))|~tuesday(X1)", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_66, plain, (member(X1,lying_days(lion))|~tuesday(X1)),\\ninference(fof_simplification,[status(thm)],[lion_lies_on_tuesday])).\n", "fillcolor": "palegreen"}, {"id": 67, "label": "c_0_67", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_67", "role": "plain", "formula": "member(X1,lying_days(lion))|~monday(X1)", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_67, plain, (member(X1,lying_days(lion))|~monday(X1)),\\ninference(fof_simplification,[status(thm)],[lion_lies_on_monday])).\n", "fillcolor": "palegreen"}, {"id": 68, "label": "c_0_68", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_68", "role": "plain", "formula": "~saturday(X1)|~sunday(X1)", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_68, plain, (~saturday(X1)|~sunday(X1)),\\ninference(fof_simplification,[status(thm)],[saturday_not_sunday])).\n", "fillcolor": "palegreen"}, {"id": 69, "label": "c_0_69", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_69", "role": "plain", "formula": "sunday(yesterday(X1))|~monday(X1)", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_69, plain, (sunday(yesterday(X1))|~monday(X1)),\\ninference(fof_simplification,[status(thm)],[yesterday_sunday])).\n", "fillcolor": "palegreen"}, {"id": 70, "label": "c_0_70", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_70", "role": "plain", "formula": "~thursday(X1)|~sunday(X1)", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_70, plain, (~thursday(X1)|~sunday(X1)),\\ninference(fof_simplification,[status(thm)],[thursday_not_sunday])).\n", "fillcolor": "palegreen"}, {"id": 71, "label": "c_0_71", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_71", "role": "plain", "formula": "~friday(X1)|~sunday(X1)", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_71, plain, (~friday(X1)|~sunday(X1)),\\ninference(fof_simplification,[status(thm)],[friday_not_sunday])).\n", "fillcolor": "palegreen"}, {"id": 72, "label": "c_0_72", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_72", "role": "plain", "formula": "monday(X1)|tuesday(X1)|wednesday(X1)|~member(X1,lying_days(lion))", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_72, plain, (monday(X1)|tuesday(X1)|wednesday(X1)|~member(X1,lying_days(lion))),\\ninference(fof_simplification,[status(thm)],[lions_lying_days])).\n", "fillcolor": "palegreen"}, {"id": 73, "label": "c_0_73", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_73", "role": "plain", "formula": "~wednesday(X1)|~saturday(X1)", "inference_rule": null, "raw_label": "cnf(c_0_73, plain, (~wednesday(X1)|~saturday(X1)),\\nc_0_52).\n", "fillcolor": "palegreen"}, {"id": 74, "label": "c_0_74", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_74", "role": "plain", "formula": "saturday(yesterday(today))|friday(yesterday(today))|thursday(yesterday(today))", "inference_rule": "sr", "raw_label": "cnf(c_0_74, plain, (saturday(yesterday(today))|friday(yesterday(today))|thursday(yesterday(today))),\\ninference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_53]), c_0_54]), c_0_55]), c_0_56])).\n", "fillcolor": "palegreen"}, {"id": 75, "label": "c_0_75", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_75", "role": "plain", "formula": "~wednesday(X1)|~thursday(X1)", "inference_rule": null, "raw_label": "cnf(c_0_75, plain, (~wednesday(X1)|~thursday(X1)),\\nc_0_57).\n", "fillcolor": "palegreen"}, {"id": 76, "label": "c_0_76", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_76", "role": "plain", "formula": "~wednesday(X1)|~friday(X1)", "inference_rule": null, "raw_label": "cnf(c_0_76, plain, (~wednesday(X1)|~friday(X1)),\\nc_0_58).\n", "fillcolor": "palegreen"}, {"id": 77, "label": "c_0_77", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_77", "role": "plain", "formula": "~tuesday(X1)|~saturday(X1)", "inference_rule": null, "raw_label": "cnf(c_0_77, plain, (~tuesday(X1)|~saturday(X1)),\\nc_0_59).\n", "fillcolor": "palegreen"}, {"id": 78, "label": "c_0_78", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_78", "role": "plain", "formula": "~tuesday(X1)|~thursday(X1)", "inference_rule": null, "raw_label": "cnf(c_0_78, plain, (~tuesday(X1)|~thursday(X1)),\\nc_0_60).\n", "fillcolor": "palegreen"}, {"id": 79, "label": "c_0_79", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_79", "role": "plain", "formula": "~tuesday(X1)|~friday(X1)", "inference_rule": null, "raw_label": "cnf(c_0_79, plain, (~tuesday(X1)|~friday(X1)),\\nc_0_61).\n", "fillcolor": "palegreen"}, {"id": 80, "label": "c_0_80", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_80", "role": "plain", "formula": "~monday(X1)|~saturday(X1)", "inference_rule": null, "raw_label": "cnf(c_0_80, plain, (~monday(X1)|~saturday(X1)),\\nc_0_62).\n", "fillcolor": "palegreen"}, {"id": 81, "label": "c_0_81", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_81", "role": "plain", "formula": "~monday(X1)|~thursday(X1)", "inference_rule": null, "raw_label": "cnf(c_0_81, plain, (~monday(X1)|~thursday(X1)),\\nc_0_63).\n", "fillcolor": "palegreen"}, {"id": 82, "label": "c_0_82", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_82", "role": "plain", "formula": "~monday(X1)|~friday(X1)", "inference_rule": null, "raw_label": "cnf(c_0_82, plain, (~monday(X1)|~friday(X1)),\\nc_0_64).\n", "fillcolor": "palegreen"}, {"id": 83, "label": "c_0_83", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_83", "role": "plain", "formula": "~member(yesterday(today", "inference_rule": "spm", "raw_label": "cnf(c_0_83, plain, (~member(yesterday(today),lying_days(lion))|~member(today,lying_days(lion))),\\ninference(spm,[status(thm)],[c_0_32, c_0_65])).\n", "fillcolor": "palegreen"}, {"id": 84, "label": "c_0_84", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_84", "role": "plain", "formula": "member(X1,lying_days(lion))|~tuesday(X1)", "inference_rule": null, "raw_label": "cnf(c_0_84, plain, (member(X1,lying_days(lion))|~tuesday(X1)),\\nc_0_66).\n", "fillcolor": "palegreen"}, {"id": 85, "label": "c_0_85", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_85", "role": "plain", "formula": "member(X1,lying_days(lion))|~wednesday(X1)", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_85, plain, (member(X1,lying_days(lion))|~wednesday(X1)),\\ninference(fof_simplification,[status(thm)],[lion_lies_on_wednesday])).\n", "fillcolor": "palegreen"}, {"id": 86, "label": "c_0_86", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_86", "role": "plain", "formula": "tuesday(yesterday(X1))|~wednesday(X1)", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_86, plain, (tuesday(yesterday(X1))|~wednesday(X1)),\\ninference(fof_simplification,[status(thm)],[yesterday_tuesday])).\n", "fillcolor": "palegreen"}, {"id": 87, "label": "c_0_87", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_87", "role": "plain", "formula": "member(X1,lying_days(lion))|~monday(X1)", "inference_rule": null, "raw_label": "cnf(c_0_87, plain, (member(X1,lying_days(lion))|~monday(X1)),\\nc_0_67).\n", "fillcolor": "palegreen"}, {"id": 88, "label": "c_0_88", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_88", "role": "plain", "formula": "monday(yesterday(X1))|~tuesday(X1)", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_88, plain, (monday(yesterday(X1))|~tuesday(X1)),\\ninference(fof_simplification,[status(thm)],[yesterday_monday])).\n", "fillcolor": "palegreen"}, {"id": 89, "label": "c_0_89", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_89", "role": "plain", "formula": "~saturday(X1)|~sunday(X1)", "inference_rule": null, "raw_label": "cnf(c_0_89, plain, (~saturday(X1)|~sunday(X1)),\\nc_0_68).\n", "fillcolor": "palegreen"}, {"id": 90, "label": "c_0_90", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_90", "role": "plain", "formula": "sunday(yesterday(X1))|~monday(X1)", "inference_rule": null, "raw_label": "cnf(c_0_90, plain, (sunday(yesterday(X1))|~monday(X1)),\\nc_0_69).\n", "fillcolor": "palegreen"}, {"id": 91, "label": "c_0_91", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_91", "role": "plain", "formula": "~thursday(X1)|~sunday(X1)", "inference_rule": null, "raw_label": "cnf(c_0_91, plain, (~thursday(X1)|~sunday(X1)),\\nc_0_70).\n", "fillcolor": "palegreen"}, {"id": 92, "label": "c_0_92", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_92", "role": "plain", "formula": "~friday(X1)|~sunday(X1)", "inference_rule": null, "raw_label": "cnf(c_0_92, plain, (~friday(X1)|~sunday(X1)),\\nc_0_71).\n", "fillcolor": "palegreen"}, {"id": 93, "label": "c_0_93", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_93", "role": "plain", "formula": "monday(X1)|tuesday(X1)|wednesday(X1)|~member(X1,lying_days(lion))", "inference_rule": null, "raw_label": "cnf(c_0_93, plain, (monday(X1)|tuesday(X1)|wednesday(X1)|~member(X1,lying_days(lion))),\\nc_0_72).\n", "fillcolor": "palegreen"}, {"id": 94, "label": "c_0_94", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_94", "role": "plain", "formula": "member(yesterday(today", "inference_rule": "spm", "raw_label": "cnf(c_0_94, plain, (member(yesterday(today),lying_days(lion))|member(today,lying_days(lion))),\\ninference(spm,[status(thm)],[c_0_37, c_0_65])).\n", "fillcolor": "palegreen"}, {"id": 95, "label": "c_0_95", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_95", "role": "plain", "formula": "~wednesday(yesterday(today))", "inference_rule": "csr", "raw_label": "cnf(c_0_95, plain, (~wednesday(yesterday(today))),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_74]), c_0_75]), c_0_76])).\n", "fillcolor": "palegreen"}, {"id": 96, "label": "c_0_96", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_96", "role": "plain", "formula": "~tuesday(yesterday(today))", "inference_rule": "csr", "raw_label": "cnf(c_0_96, plain, (~tuesday(yesterday(today))),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_74]), c_0_78]), c_0_79])).\n", "fillcolor": "palegreen"}, {"id": 97, "label": "c_0_97", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_97", "role": "plain", "formula": "~monday(yesterday(today))", "inference_rule": "csr", "raw_label": "cnf(c_0_97, plain, (~monday(yesterday(today))),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_74]), c_0_81]), c_0_82])).\n", "fillcolor": "palegreen"}, {"id": 98, "label": "c_0_98", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_98", "role": "plain", "formula": "~member(today,lying_days(lion))|~tuesday(yesterday(today))", "inference_rule": "spm", "raw_label": "cnf(c_0_98, plain, (~member(today,lying_days(lion))|~tuesday(yesterday(today))),\\ninference(spm,[status(thm)],[c_0_83, c_0_84])).\n", "fillcolor": "palegreen"}, {"id": 99, "label": "c_0_99", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_99", "role": "plain", "formula": "member(X1,lying_days(lion))|~wednesday(X1)", "inference_rule": null, "raw_label": "cnf(c_0_99, plain, (member(X1,lying_days(lion))|~wednesday(X1)),\\nc_0_85).\n", "fillcolor": "palegreen"}, {"id": 100, "label": "c_0_100", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_100", "role": "plain", "formula": "tuesday(yesterday(X1))|~wednesday(X1)", "inference_rule": null, "raw_label": "cnf(c_0_100, plain, (tuesday(yesterday(X1))|~wednesday(X1)),\\nc_0_86).\n", "fillcolor": "palegreen"}, {"id": 101, "label": "c_0_101", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_101", "role": "plain", "formula": "~member(today,lying_days(lion))|~monday(yesterday(today))", "inference_rule": "spm", "raw_label": "cnf(c_0_101, plain, (~member(today,lying_days(lion))|~monday(yesterday(today))),\\ninference(spm,[status(thm)],[c_0_83, c_0_87])).\n", "fillcolor": "palegreen"}, {"id": 102, "label": "c_0_102", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_102", "role": "plain", "formula": "monday(yesterday(X1))|~tuesday(X1)", "inference_rule": null, "raw_label": "cnf(c_0_102, plain, (monday(yesterday(X1))|~tuesday(X1)),\\nc_0_88).\n", "fillcolor": "palegreen"}, {"id": 103, "label": "c_0_103", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_103", "role": "plain", "formula": "~saturday(yesterday(X1))|~monday(X1)", "inference_rule": "spm", "raw_label": "cnf(c_0_103, plain, (~saturday(yesterday(X1))|~monday(X1)),\\ninference(spm,[status(thm)],[c_0_89, c_0_90])).\n", "fillcolor": "palegreen"}, {"id": 104, "label": "c_0_104", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_104", "role": "plain", "formula": "~thursday(yesterday(X1))|~monday(X1)", "inference_rule": "spm", "raw_label": "cnf(c_0_104, plain, (~thursday(yesterday(X1))|~monday(X1)),\\ninference(spm,[status(thm)],[c_0_91, c_0_90])).\n", "fillcolor": "palegreen"}, {"id": 105, "label": "c_0_105", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_105", "role": "plain", "formula": "~friday(yesterday(X1))|~monday(X1)", "inference_rule": "spm", "raw_label": "cnf(c_0_105, plain, (~friday(yesterday(X1))|~monday(X1)),\\ninference(spm,[status(thm)],[c_0_92, c_0_90])).\n", "fillcolor": "palegreen"}, {"id": 106, "label": "c_0_106", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_106", "role": "plain", "formula": "member(today,lying_days(lion))", "inference_rule": "sr", "raw_label": "cnf(c_0_106, plain, (member(today,lying_days(lion))),\\ninference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_94]), c_0_95]), c_0_96]), c_0_97])).\n", "fillcolor": "palegreen"}, {"id": 107, "label": "c_0_107", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_107", "role": "plain", "formula": "~wednesday(today)", "inference_rule": "csr", "raw_label": "cnf(c_0_107, plain, (~wednesday(today)),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_99]), c_0_100])).\n", "fillcolor": "palegreen"}, {"id": 108, "label": "c_0_108", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_108", "role": "plain", "formula": "~tuesday(today)", "inference_rule": "csr", "raw_label": "cnf(c_0_108, plain, (~tuesday(today)),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_84]), c_0_102])).\n", "fillcolor": "palegreen"}, {"id": 109, "label": "c_0_109", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_109", "role": "plain", "formula": "~monday(today)", "inference_rule": "csr", "raw_label": "cnf(c_0_109, plain, (~monday(today)),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_74]), c_0_104]), c_0_105])).\n", "fillcolor": "palegreen"}, {"id": 110, "label": "c_0_110", "color": {"background": "#000000", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#FFFFFF"}, "cnf_name": "c_0_110", "role": "plain", "formula": "$false", "inference_rule": "sr", "raw_label": "cnf(c_0_110, plain, ($false),\\ninference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_106]), c_0_107]), c_0_108]), c_0_109])).\n", "fillcolor": "darkorchid1"}], "edges": [{"from": 0, "to": 30, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 1, "to": 31, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 30, "to": 32, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 2, "to": 33, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 3, "to": 34, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 4, "to": 35, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 5, "to": 36, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 31, "to": 37, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 33, "to": 38, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 32, "to": 38, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 34, "to": 39, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 6, "to": 40, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 7, "to": 41, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 35, "to": 42, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 8, "to": 43, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 36, "to": 44, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 33, "to": 45, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 37, "to": 45, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 39, "to": 46, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 38, "to": 46, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 40, "to": 47, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 41, "to": 48, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 42, "to": 49, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 38, "to": 49, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 43, "to": 50, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 9, "to": 51, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 10, "to": 52, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 45, "to": 53, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 44, "to": 53, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 48, "to": 54, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 47, "to": 54, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 46, "to": 54, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 50, "to": 55, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 39, "to": 55, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 49, "to": 55, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 51, "to": 56, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 11, "to": 57, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 12, "to": 58, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 13, "to": 59, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 14, "to": 60, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 15, "to": 61, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 16, "to": 62, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 17, "to": 63, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 18, "to": 64, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 19, "to": 65, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 20, "to": 66, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 21, "to": 67, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 22, "to": 68, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 23, "to": 69, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 24, "to": 70, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 25, "to": 71, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 26, "to": 72, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 52, "to": 73, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 56, "to": 74, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 55, "to": 74, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 54, "to": 74, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 53, "to": 74, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 44, "to": 74, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 57, "to": 75, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 58, "to": 76, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 59, "to": 77, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 60, "to": 78, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 61, "to": 79, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 62, "to": 80, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 63, "to": 81, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 64, "to": 82, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 65, "to": 83, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 32, "to": 83, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 66, "to": 84, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 27, "to": 85, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 28, "to": 86, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 67, "to": 87, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 29, "to": 88, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 68, "to": 89, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 69, "to": 90, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 70, "to": 91, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 71, "to": 92, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 72, "to": 93, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 65, "to": 94, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 37, "to": 94, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 76, "to": 95, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 75, "to": 95, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 74, "to": 95, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 73, "to": 95, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 79, "to": 96, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 78, "to": 96, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 74, "to": 96, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 77, "to": 96, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 82, "to": 97, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 81, "to": 97, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 74, "to": 97, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 80, "to": 97, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 84, "to": 98, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 83, "to": 98, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 85, "to": 99, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 86, "to": 100, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 87, "to": 101, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 83, "to": 101, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 88, "to": 102, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 90, "to": 103, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 89, "to": 103, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 90, "to": 104, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 91, "to": 104, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 90, "to": 105, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 92, "to": 105, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 97, "to": 106, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 96, "to": 106, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 95, "to": 106, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 94, "to": 106, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 93, "to": 106, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 100, "to": 107, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 99, "to": 107, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 98, "to": 107, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 102, "to": 108, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 84, "to": 108, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 101, "to": 108, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 105, "to": 109, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 104, "to": 109, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 74, "to": 109, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 103, "to": 109, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 109, "to": 110, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 108, "to": 110, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 107, "to": 110, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 106, "to": 110, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 93, "to": 110, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}]}, {"problem_file": "PUZ015-2.006.p", "status": "# SZS status Unsatisfiable", "nodes": [{"id": 0, "label": "unique_cover2_6_2", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover2_6_2", "role": "negated_conjecture", "formula": "~horizontal_6_2|~horizontal_6_1", "inference_rule": null, "raw_label": "cnf(unique_cover2_6_2, negated_conjecture, (~horizontal_6_2|~horizontal_6_1),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover2_6_2)).\n", "fillcolor": "dodgerblue"}, {"id": 1, "label": "unique_cover5_5_1", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover5_5_1", "role": "negated_conjecture", "formula": "~vertical_5_1|~vertical_4_1", "inference_rule": null, "raw_label": "cnf(unique_cover5_5_1, negated_conjecture, (~vertical_5_1|~vertical_4_1),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover5_5_1)).\n", "fillcolor": "dodgerblue"}, {"id": 2, "label": "covered_6_1", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "covered_6_1", "role": "negated_conjecture", "formula": "horizontal_6_1|vertical_5_1", "inference_rule": null, "raw_label": "cnf(covered_6_1, negated_conjecture, (horizontal_6_1|vertical_5_1),\\nfile('Problems/PUZ/PUZ015-2.006.p', covered_6_1)).\n", "fillcolor": "dodgerblue"}, {"id": 3, "label": "unique_cover2_6_4", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover2_6_4", "role": "negated_conjecture", "formula": "~horizontal_6_4|~horizontal_6_3", "inference_rule": null, "raw_label": "cnf(unique_cover2_6_4, negated_conjecture, (~horizontal_6_4|~horizontal_6_3),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover2_6_4)).\n", "fillcolor": "dodgerblue"}, {"id": 4, "label": "covered_6_3", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "covered_6_3", "role": "negated_conjecture", "formula": "horizontal_6_3|horizontal_6_2|vertical_5_3", "inference_rule": null, "raw_label": "cnf(covered_6_3, negated_conjecture, (horizontal_6_3|horizontal_6_2|vertical_5_3),\\nfile('Problems/PUZ/PUZ015-2.006.p', covered_6_3)).\n", "fillcolor": "dodgerblue"}, {"id": 5, "label": "unique_cover1_5_1", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover1_5_1", "role": "negated_conjecture", "formula": "~horizontal_5_1|~vertical_5_1", "inference_rule": null, "raw_label": "cnf(unique_cover1_5_1, negated_conjecture, (~horizontal_5_1|~vertical_5_1),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover1_5_1)).\n", "fillcolor": "dodgerblue"}, {"id": 6, "label": "unique_cover6_6_2", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover6_6_2", "role": "negated_conjecture", "formula": "~horizontal_6_1|~vertical_5_2", "inference_rule": null, "raw_label": "cnf(unique_cover6_6_2, negated_conjecture, (~horizontal_6_1|~vertical_5_2),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover6_6_2)).\n", "fillcolor": "dodgerblue"}, {"id": 7, "label": "unique_cover5_3_1", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover5_3_1", "role": "negated_conjecture", "formula": "~vertical_3_1|~vertical_2_1", "inference_rule": null, "raw_label": "cnf(unique_cover5_3_1, negated_conjecture, (~vertical_3_1|~vertical_2_1),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover5_3_1)).\n", "fillcolor": "dodgerblue"}, {"id": 8, "label": "covered_4_1", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "covered_4_1", "role": "negated_conjecture", "formula": "horizontal_4_1|vertical_4_1|vertical_3_1", "inference_rule": null, "raw_label": "cnf(covered_4_1, negated_conjecture, (horizontal_4_1|vertical_4_1|vertical_3_1),\\nfile('Problems/PUZ/PUZ015-2.006.p', covered_4_1)).\n", "fillcolor": "dodgerblue"}, {"id": 9, "label": "uncovered3_6_6", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "uncovered3_6_6", "role": "negated_conjecture", "formula": "~horizontal_6_5", "inference_rule": null, "raw_label": "cnf(uncovered3_6_6, negated_conjecture, (~horizontal_6_5),\\nfile('Problems/PUZ/PUZ015-2.006.p', uncovered3_6_6)).\n", "fillcolor": "dodgerblue"}, {"id": 10, "label": "uncovered4_6_6", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "uncovered4_6_6", "role": "negated_conjecture", "formula": "~vertical_5_6", "inference_rule": null, "raw_label": "cnf(uncovered4_6_6, negated_conjecture, (~vertical_5_6),\\nfile('Problems/PUZ/PUZ015-2.006.p', uncovered4_6_6)).\n", "fillcolor": "dodgerblue"}, {"id": 11, "label": "unique_cover3_6_2", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover3_6_2", "role": "negated_conjecture", "formula": "~horizontal_6_2|~vertical_5_2", "inference_rule": null, "raw_label": "cnf(unique_cover3_6_2, negated_conjecture, (~horizontal_6_2|~vertical_5_2),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover3_6_2)).\n", "fillcolor": "dodgerblue"}, {"id": 12, "label": "unique_cover2_1_5", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover2_1_5", "role": "negated_conjecture", "formula": "~horizontal_1_5|~horizontal_1_4", "inference_rule": null, "raw_label": "cnf(unique_cover2_1_5, negated_conjecture, (~horizontal_1_5|~horizontal_1_4),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover2_1_5)).\n", "fillcolor": "dodgerblue"}, {"id": 13, "label": "covered_6_5", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "covered_6_5", "role": "negated_conjecture", "formula": "horizontal_6_5|horizontal_6_4|vertical_5_5", "inference_rule": null, "raw_label": "cnf(covered_6_5, negated_conjecture, (horizontal_6_5|horizontal_6_4|vertical_5_5),\\nfile('Problems/PUZ/PUZ015-2.006.p', covered_6_5)).\n", "fillcolor": "dodgerblue"}, {"id": 14, "label": "uncovered2_1_1", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "uncovered2_1_1", "role": "negated_conjecture", "formula": "~vertical_1_1", "inference_rule": null, "raw_label": "cnf(uncovered2_1_1, negated_conjecture, (~vertical_1_1),\\nfile('Problems/PUZ/PUZ015-2.006.p', uncovered2_1_1)).\n", "fillcolor": "dodgerblue"}, {"id": 15, "label": "unique_cover4_4_6", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover4_4_6", "role": "negated_conjecture", "formula": "~vertical_4_6|~horizontal_4_5", "inference_rule": null, "raw_label": "cnf(unique_cover4_4_6, negated_conjecture, (~vertical_4_6|~horizontal_4_5),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover4_4_6)).\n", "fillcolor": "dodgerblue"}, {"id": 16, "label": "covered_5_6", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "covered_5_6", "role": "negated_conjecture", "formula": "vertical_5_6|horizontal_5_5|vertical_4_6", "inference_rule": null, "raw_label": "cnf(covered_5_6, negated_conjecture, (vertical_5_6|horizontal_5_5|vertical_4_6),\\nfile('Problems/PUZ/PUZ015-2.006.p', covered_5_6)).\n", "fillcolor": "dodgerblue"}, {"id": 17, "label": "unique_cover3_5_1", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover3_5_1", "role": "negated_conjecture", "formula": "~horizontal_5_1|~vertical_4_1", "inference_rule": null, "raw_label": "cnf(unique_cover3_5_1, negated_conjecture, (~horizontal_5_1|~vertical_4_1),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover3_5_1)).\n", "fillcolor": "dodgerblue"}, {"id": 18, "label": "unique_cover5_2_6", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover5_2_6", "role": "negated_conjecture", "formula": "~vertical_2_6|~vertical_1_6", "inference_rule": null, "raw_label": "cnf(unique_cover5_2_6, negated_conjecture, (~vertical_2_6|~vertical_1_6),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover5_2_6)).\n", "fillcolor": "dodgerblue"}, {"id": 19, "label": "covered_1_6", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "covered_1_6", "role": "negated_conjecture", "formula": "vertical_1_6|horizontal_1_5", "inference_rule": null, "raw_label": "cnf(covered_1_6, negated_conjecture, (vertical_1_6|horizontal_1_5),\\nfile('Problems/PUZ/PUZ015-2.006.p', covered_1_6)).\n", "fillcolor": "dodgerblue"}, {"id": 20, "label": "unique_cover5_4_6", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover5_4_6", "role": "negated_conjecture", "formula": "~vertical_4_6|~vertical_3_6", "inference_rule": null, "raw_label": "cnf(unique_cover5_4_6, negated_conjecture, (~vertical_4_6|~vertical_3_6),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover5_4_6)).\n", "fillcolor": "dodgerblue"}, {"id": 21, "label": "covered_2_1", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "covered_2_1", "role": "negated_conjecture", "formula": "horizontal_2_1|vertical_2_1|vertical_1_1", "inference_rule": null, "raw_label": "cnf(covered_2_1, negated_conjecture, (horizontal_2_1|vertical_2_1|vertical_1_1),\\nfile('Problems/PUZ/PUZ015-2.006.p', covered_2_1)).\n", "fillcolor": "dodgerblue"}, {"id": 22, "label": "unique_cover1_5_3", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover1_5_3", "role": "negated_conjecture", "formula": "~horizontal_5_3|~vertical_5_3", "inference_rule": null, "raw_label": "cnf(unique_cover1_5_3, negated_conjecture, (~horizontal_5_3|~vertical_5_3),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover1_5_3)).\n", "fillcolor": "dodgerblue"}, {"id": 23, "label": "covered_3_6", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "covered_3_6", "role": "negated_conjecture", "formula": "vertical_3_6|horizontal_3_5|vertical_2_6", "inference_rule": null, "raw_label": "cnf(covered_3_6, negated_conjecture, (vertical_3_6|horizontal_3_5|vertical_2_6),\\nfile('Problems/PUZ/PUZ015-2.006.p', covered_3_6)).\n", "fillcolor": "dodgerblue"}, {"id": 24, "label": "unique_cover1_5_5", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover1_5_5", "role": "negated_conjecture", "formula": "~horizontal_5_5|~vertical_5_5", "inference_rule": null, "raw_label": "cnf(unique_cover1_5_5, negated_conjecture, (~horizontal_5_5|~vertical_5_5),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover1_5_5)).\n", "fillcolor": "dodgerblue"}, {"id": 25, "label": "covered_4_5", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "covered_4_5", "role": "negated_conjecture", "formula": "horizontal_4_5|vertical_4_5|horizontal_4_4|vertical_3_5", "inference_rule": null, "raw_label": "cnf(covered_4_5, negated_conjecture, (horizontal_4_5|vertical_4_5|horizontal_4_4|vertical_3_5),\\nfile('Problems/PUZ/PUZ015-2.006.p', covered_4_5)).\n", "fillcolor": "dodgerblue"}, {"id": 26, "label": "unique_cover5_5_5", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover5_5_5", "role": "negated_conjecture", "formula": "~vertical_5_5|~vertical_4_5", "inference_rule": null, "raw_label": "cnf(unique_cover5_5_5, negated_conjecture, (~vertical_5_5|~vertical_4_5),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover5_5_5)).\n", "fillcolor": "dodgerblue"}, {"id": 27, "label": "unique_cover2_5_3", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover2_5_3", "role": "negated_conjecture", "formula": "~horizontal_5_3|~horizontal_5_2", "inference_rule": null, "raw_label": "cnf(unique_cover2_5_3, negated_conjecture, (~horizontal_5_3|~horizontal_5_2),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover2_5_3)).\n", "fillcolor": "dodgerblue"}, {"id": 28, "label": "unique_cover4_4_2", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover4_4_2", "role": "negated_conjecture", "formula": "~vertical_4_2|~horizontal_4_1", "inference_rule": null, "raw_label": "cnf(unique_cover4_4_2, negated_conjecture, (~vertical_4_2|~horizontal_4_1),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover4_4_2)).\n", "fillcolor": "dodgerblue"}, {"id": 29, "label": "covered_5_2", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "covered_5_2", "role": "negated_conjecture", "formula": "horizontal_5_2|vertical_5_2|horizontal_5_1|vertical_4_2", "inference_rule": null, "raw_label": "cnf(covered_5_2, negated_conjecture, (horizontal_5_2|vertical_5_2|horizontal_5_1|vertical_4_2),\\nfile('Problems/PUZ/PUZ015-2.006.p', covered_5_2)).\n", "fillcolor": "dodgerblue"}, {"id": 30, "label": "unique_cover3_6_4", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover3_6_4", "role": "negated_conjecture", "formula": "~horizontal_6_4|~vertical_5_4", "inference_rule": null, "raw_label": "cnf(unique_cover3_6_4, negated_conjecture, (~horizontal_6_4|~vertical_5_4),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover3_6_4)).\n", "fillcolor": "dodgerblue"}, {"id": 31, "label": "unique_cover2_1_3", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover2_1_3", "role": "negated_conjecture", "formula": "~horizontal_1_3|~horizontal_1_2", "inference_rule": null, "raw_label": "cnf(unique_cover2_1_3, negated_conjecture, (~horizontal_1_3|~horizontal_1_2),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover2_1_3)).\n", "fillcolor": "dodgerblue"}, {"id": 32, "label": "covered_1_4", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "covered_1_4", "role": "negated_conjecture", "formula": "horizontal_1_4|vertical_1_4|horizontal_1_3", "inference_rule": null, "raw_label": "cnf(covered_1_4, negated_conjecture, (horizontal_1_4|vertical_1_4|horizontal_1_3),\\nfile('Problems/PUZ/PUZ015-2.006.p', covered_1_4)).\n", "fillcolor": "dodgerblue"}, {"id": 33, "label": "unique_cover6_2_6", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover6_2_6", "role": "negated_conjecture", "formula": "~horizontal_2_5|~vertical_1_6", "inference_rule": null, "raw_label": "cnf(unique_cover6_2_6, negated_conjecture, (~horizontal_2_5|~vertical_1_6),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover6_2_6)).\n", "fillcolor": "dodgerblue"}, {"id": 34, "label": "unique_cover2_5_5", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover2_5_5", "role": "negated_conjecture", "formula": "~horizontal_5_5|~horizontal_5_4", "inference_rule": null, "raw_label": "cnf(unique_cover2_5_5, negated_conjecture, (~horizontal_5_5|~horizontal_5_4),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover2_5_5)).\n", "fillcolor": "dodgerblue"}, {"id": 35, "label": "unique_cover6_5_5", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover6_5_5", "role": "negated_conjecture", "formula": "~horizontal_5_4|~vertical_4_5", "inference_rule": null, "raw_label": "cnf(unique_cover6_5_5, negated_conjecture, (~horizontal_5_4|~vertical_4_5),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover6_5_5)).\n", "fillcolor": "dodgerblue"}, {"id": 36, "label": "unique_cover6_4_2", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover6_4_2", "role": "negated_conjecture", "formula": "~horizontal_4_1|~vertical_3_2", "inference_rule": null, "raw_label": "cnf(unique_cover6_4_2, negated_conjecture, (~horizontal_4_1|~vertical_3_2),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover6_4_2)).\n", "fillcolor": "dodgerblue"}, {"id": 37, "label": "unique_cover5_4_2", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover5_4_2", "role": "negated_conjecture", "formula": "~vertical_4_2|~vertical_3_2", "inference_rule": null, "raw_label": "cnf(unique_cover5_4_2, negated_conjecture, (~vertical_4_2|~vertical_3_2),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover5_4_2)).\n", "fillcolor": "dodgerblue"}, {"id": 38, "label": "uncovered1_1_1", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "uncovered1_1_1", "role": "negated_conjecture", "formula": "~horizontal_1_1", "inference_rule": null, "raw_label": "cnf(uncovered1_1_1, negated_conjecture, (~horizontal_1_1),\\nfile('Problems/PUZ/PUZ015-2.006.p', uncovered1_1_1)).\n", "fillcolor": "dodgerblue"}, {"id": 39, "label": "unique_cover4_1_5", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover4_1_5", "role": "negated_conjecture", "formula": "~vertical_1_5|~horizontal_1_4", "inference_rule": null, "raw_label": "cnf(unique_cover4_1_5, negated_conjecture, (~vertical_1_5|~horizontal_1_4),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover4_1_5)).\n", "fillcolor": "dodgerblue"}, {"id": 40, "label": "unique_cover1_1_5", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover1_1_5", "role": "negated_conjecture", "formula": "~horizontal_1_5|~vertical_1_5", "inference_rule": null, "raw_label": "cnf(unique_cover1_1_5, negated_conjecture, (~horizontal_1_5|~vertical_1_5),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover1_1_5)).\n", "fillcolor": "dodgerblue"}, {"id": 41, "label": "unique_cover2_2_2", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover2_2_2", "role": "negated_conjecture", "formula": "~horizontal_2_2|~horizontal_2_1", "inference_rule": null, "raw_label": "cnf(unique_cover2_2_2, negated_conjecture, (~horizontal_2_2|~horizontal_2_1),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover2_2_2)).\n", "fillcolor": "dodgerblue"}, {"id": 42, "label": "covered_5_4", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "covered_5_4", "role": "negated_conjecture", "formula": "horizontal_5_4|vertical_5_4|horizontal_5_3|vertical_4_4", "inference_rule": null, "raw_label": "cnf(covered_5_4, negated_conjecture, (horizontal_5_4|vertical_5_4|horizontal_5_3|vertical_4_4),\\nfile('Problems/PUZ/PUZ015-2.006.p', covered_5_4)).\n", "fillcolor": "dodgerblue"}, {"id": 43, "label": "covered_1_2", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "covered_1_2", "role": "negated_conjecture", "formula": "horizontal_1_2|vertical_1_2|horizontal_1_1", "inference_rule": null, "raw_label": "cnf(covered_1_2, negated_conjecture, (horizontal_1_2|vertical_1_2|horizontal_1_1),\\nfile('Problems/PUZ/PUZ015-2.006.p', covered_1_2)).\n", "fillcolor": "dodgerblue"}, {"id": 44, "label": "unique_cover6_5_3", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover6_5_3", "role": "negated_conjecture", "formula": "~horizontal_5_2|~vertical_4_3", "inference_rule": null, "raw_label": "cnf(unique_cover6_5_3, negated_conjecture, (~horizontal_5_2|~vertical_4_3),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover6_5_3)).\n", "fillcolor": "dodgerblue"}, {"id": 45, "label": "unique_cover3_3_5", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover3_3_5", "role": "negated_conjecture", "formula": "~horizontal_3_5|~vertical_2_5", "inference_rule": null, "raw_label": "cnf(unique_cover3_3_5, negated_conjecture, (~horizontal_3_5|~vertical_2_5),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover3_3_5)).\n", "fillcolor": "dodgerblue"}, {"id": 46, "label": "unique_cover4_2_6", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover4_2_6", "role": "negated_conjecture", "formula": "~vertical_2_6|~horizontal_2_5", "inference_rule": null, "raw_label": "cnf(unique_cover4_2_6, negated_conjecture, (~vertical_2_6|~horizontal_2_5),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover4_2_6)).\n", "fillcolor": "dodgerblue"}, {"id": 47, "label": "unique_cover4_4_4", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover4_4_4", "role": "negated_conjecture", "formula": "~vertical_4_4|~horizontal_4_3", "inference_rule": null, "raw_label": "cnf(unique_cover4_4_4, negated_conjecture, (~vertical_4_4|~horizontal_4_3),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover4_4_4)).\n", "fillcolor": "dodgerblue"}, {"id": 48, "label": "unique_cover2_4_4", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover2_4_4", "role": "negated_conjecture", "formula": "~horizontal_4_4|~horizontal_4_3", "inference_rule": null, "raw_label": "cnf(unique_cover2_4_4, negated_conjecture, (~horizontal_4_4|~horizontal_4_3),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover2_4_4)).\n", "fillcolor": "dodgerblue"}, {"id": 49, "label": "unique_cover3_4_2", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover3_4_2", "role": "negated_conjecture", "formula": "~horizontal_4_2|~vertical_3_2", "inference_rule": null, "raw_label": "cnf(unique_cover3_4_2, negated_conjecture, (~horizontal_4_2|~vertical_3_2),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover3_4_2)).\n", "fillcolor": "dodgerblue"}, {"id": 50, "label": "unique_cover5_4_4", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover5_4_4", "role": "negated_conjecture", "formula": "~vertical_4_4|~vertical_3_4", "inference_rule": null, "raw_label": "cnf(unique_cover5_4_4, negated_conjecture, (~vertical_4_4|~vertical_3_4),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover5_4_4)).\n", "fillcolor": "dodgerblue"}, {"id": 51, "label": "unique_cover3_4_4", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover3_4_4", "role": "negated_conjecture", "formula": "~horizontal_4_4|~vertical_3_4", "inference_rule": null, "raw_label": "cnf(unique_cover3_4_4, negated_conjecture, (~horizontal_4_4|~vertical_3_4),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover3_4_4)).\n", "fillcolor": "dodgerblue"}, {"id": 52, "label": "unique_cover5_2_4", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover5_2_4", "role": "negated_conjecture", "formula": "~vertical_2_4|~vertical_1_4", "inference_rule": null, "raw_label": "cnf(unique_cover5_2_4, negated_conjecture, (~vertical_2_4|~vertical_1_4),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover5_2_4)).\n", "fillcolor": "dodgerblue"}, {"id": 53, "label": "unique_cover4_5_3", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover4_5_3", "role": "negated_conjecture", "formula": "~vertical_5_3|~horizontal_5_2", "inference_rule": null, "raw_label": "cnf(unique_cover4_5_3, negated_conjecture, (~vertical_5_3|~horizontal_5_2),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover4_5_3)).\n", "fillcolor": "dodgerblue"}, {"id": 54, "label": "covered_4_3", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "covered_4_3", "role": "negated_conjecture", "formula": "horizontal_4_3|vertical_4_3|horizontal_4_2|vertical_3_3", "inference_rule": null, "raw_label": "cnf(covered_4_3, negated_conjecture, (horizontal_4_3|vertical_4_3|horizontal_4_2|vertical_3_3),\\nfile('Problems/PUZ/PUZ015-2.006.p', covered_4_3)).\n", "fillcolor": "dodgerblue"}, {"id": 55, "label": "covered_2_5", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "covered_2_5", "role": "negated_conjecture", "formula": "horizontal_2_5|vertical_2_5|horizontal_2_4|vertical_1_5", "inference_rule": null, "raw_label": "cnf(covered_2_5, negated_conjecture, (horizontal_2_5|vertical_2_5|horizontal_2_4|vertical_1_5),\\nfile('Problems/PUZ/PUZ015-2.006.p', covered_2_5)).\n", "fillcolor": "dodgerblue"}, {"id": 56, "label": "unique_cover1_2_4", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover1_2_4", "role": "negated_conjecture", "formula": "~horizontal_2_4|~vertical_2_4", "inference_rule": null, "raw_label": "cnf(unique_cover1_2_4, negated_conjecture, (~horizontal_2_4|~vertical_2_4),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover1_2_4)).\n", "fillcolor": "dodgerblue"}, {"id": 57, "label": "unique_cover1_3_5", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover1_3_5", "role": "negated_conjecture", "formula": "~horizontal_3_5|~vertical_3_5", "inference_rule": null, "raw_label": "cnf(unique_cover1_3_5, negated_conjecture, (~horizontal_3_5|~vertical_3_5),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover1_3_5)).\n", "fillcolor": "dodgerblue"}, {"id": 58, "label": "unique_cover5_3_5", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover5_3_5", "role": "negated_conjecture", "formula": "~vertical_3_5|~vertical_2_5", "inference_rule": null, "raw_label": "cnf(unique_cover5_3_5, negated_conjecture, (~vertical_3_5|~vertical_2_5),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover5_3_5)).\n", "fillcolor": "dodgerblue"}, {"id": 59, "label": "unique_cover3_3_1", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover3_3_1", "role": "negated_conjecture", "formula": "~horizontal_3_1|~vertical_2_1", "inference_rule": null, "raw_label": "cnf(unique_cover3_3_1, negated_conjecture, (~horizontal_3_1|~vertical_2_1),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover3_3_1)).\n", "fillcolor": "dodgerblue"}, {"id": 60, "label": "unique_cover2_3_5", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover2_3_5", "role": "negated_conjecture", "formula": "~horizontal_3_5|~horizontal_3_4", "inference_rule": null, "raw_label": "cnf(unique_cover2_3_5, negated_conjecture, (~horizontal_3_5|~horizontal_3_4),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover2_3_5)).\n", "fillcolor": "dodgerblue"}, {"id": 61, "label": "unique_cover1_3_3", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover1_3_3", "role": "negated_conjecture", "formula": "~horizontal_3_3|~vertical_3_3", "inference_rule": null, "raw_label": "cnf(unique_cover1_3_3, negated_conjecture, (~horizontal_3_3|~vertical_3_3),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover1_3_3)).\n", "fillcolor": "dodgerblue"}, {"id": 62, "label": "unique_cover6_3_5", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover6_3_5", "role": "negated_conjecture", "formula": "~horizontal_3_4|~vertical_2_5", "inference_rule": null, "raw_label": "cnf(unique_cover6_3_5, negated_conjecture, (~horizontal_3_4|~vertical_2_5),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover6_3_5)).\n", "fillcolor": "dodgerblue"}, {"id": 63, "label": "unique_cover3_2_2", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover3_2_2", "role": "negated_conjecture", "formula": "~horizontal_2_2|~vertical_1_2", "inference_rule": null, "raw_label": "cnf(unique_cover3_2_2, negated_conjecture, (~horizontal_2_2|~vertical_1_2),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover3_2_2)).\n", "fillcolor": "dodgerblue"}, {"id": 64, "label": "unique_cover5_5_3", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover5_5_3", "role": "negated_conjecture", "formula": "~vertical_5_3|~vertical_4_3", "inference_rule": null, "raw_label": "cnf(unique_cover5_5_3, negated_conjecture, (~vertical_5_3|~vertical_4_3),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover5_5_3)).\n", "fillcolor": "dodgerblue"}, {"id": 65, "label": "unique_cover3_2_4", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover3_2_4", "role": "negated_conjecture", "formula": "~horizontal_2_4|~vertical_1_4", "inference_rule": null, "raw_label": "cnf(unique_cover3_2_4, negated_conjecture, (~horizontal_2_4|~vertical_1_4),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover3_2_4)).\n", "fillcolor": "dodgerblue"}, {"id": 66, "label": "covered_3_4", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "covered_3_4", "role": "negated_conjecture", "formula": "horizontal_3_4|vertical_3_4|horizontal_3_3|vertical_2_4", "inference_rule": null, "raw_label": "cnf(covered_3_4, negated_conjecture, (horizontal_3_4|vertical_3_4|horizontal_3_3|vertical_2_4),\\nfile('Problems/PUZ/PUZ015-2.006.p', covered_3_4)).\n", "fillcolor": "dodgerblue"}, {"id": 67, "label": "covered_3_2", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "covered_3_2", "role": "negated_conjecture", "formula": "horizontal_3_2|vertical_3_2|horizontal_3_1|vertical_2_2", "inference_rule": null, "raw_label": "cnf(covered_3_2, negated_conjecture, (horizontal_3_2|vertical_3_2|horizontal_3_1|vertical_2_2),\\nfile('Problems/PUZ/PUZ015-2.006.p', covered_3_2)).\n", "fillcolor": "dodgerblue"}, {"id": 68, "label": "unique_cover1_2_2", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover1_2_2", "role": "negated_conjecture", "formula": "~horizontal_2_2|~vertical_2_2", "inference_rule": null, "raw_label": "cnf(unique_cover1_2_2, negated_conjecture, (~horizontal_2_2|~vertical_2_2),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover1_2_2)).\n", "fillcolor": "dodgerblue"}, {"id": 69, "label": "unique_cover2_4_2", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover2_4_2", "role": "negated_conjecture", "formula": "~horizontal_4_2|~horizontal_4_1", "inference_rule": null, "raw_label": "cnf(unique_cover2_4_2, negated_conjecture, (~horizontal_4_2|~horizontal_4_1),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover2_4_2)).\n", "fillcolor": "dodgerblue"}, {"id": 70, "label": "unique_cover1_4_2", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover1_4_2", "role": "negated_conjecture", "formula": "~horizontal_4_2|~vertical_4_2", "inference_rule": null, "raw_label": "cnf(unique_cover1_4_2, negated_conjecture, (~horizontal_4_2|~vertical_4_2),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover1_4_2)).\n", "fillcolor": "dodgerblue"}, {"id": 71, "label": "unique_cover2_3_3", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover2_3_3", "role": "negated_conjecture", "formula": "~horizontal_3_3|~horizontal_3_2", "inference_rule": null, "raw_label": "cnf(unique_cover2_3_3, negated_conjecture, (~horizontal_3_3|~horizontal_3_2),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover2_3_3)).\n", "fillcolor": "dodgerblue"}, {"id": 72, "label": "unique_cover6_2_4", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover6_2_4", "role": "negated_conjecture", "formula": "~horizontal_2_3|~vertical_1_4", "inference_rule": null, "raw_label": "cnf(unique_cover6_2_4, negated_conjecture, (~horizontal_2_3|~vertical_1_4),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover6_2_4)).\n", "fillcolor": "dodgerblue"}, {"id": 73, "label": "unique_cover2_2_4", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover2_2_4", "role": "negated_conjecture", "formula": "~horizontal_2_4|~horizontal_2_3", "inference_rule": null, "raw_label": "cnf(unique_cover2_2_4, negated_conjecture, (~horizontal_2_4|~horizontal_2_3),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover2_2_4)).\n", "fillcolor": "dodgerblue"}, {"id": 74, "label": "unique_cover6_2_2", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover6_2_2", "role": "negated_conjecture", "formula": "~horizontal_2_1|~vertical_1_2", "inference_rule": null, "raw_label": "cnf(unique_cover6_2_2, negated_conjecture, (~horizontal_2_1|~vertical_1_2),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover6_2_2)).\n", "fillcolor": "dodgerblue"}, {"id": 75, "label": "unique_cover4_3_3", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover4_3_3", "role": "negated_conjecture", "formula": "~vertical_3_3|~horizontal_3_2", "inference_rule": null, "raw_label": "cnf(unique_cover4_3_3, negated_conjecture, (~vertical_3_3|~horizontal_3_2),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover4_3_3)).\n", "fillcolor": "dodgerblue"}, {"id": 76, "label": "unique_cover4_2_2", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover4_2_2", "role": "negated_conjecture", "formula": "~vertical_2_2|~horizontal_2_1", "inference_rule": null, "raw_label": "cnf(unique_cover4_2_2, negated_conjecture, (~vertical_2_2|~horizontal_2_1),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover4_2_2)).\n", "fillcolor": "dodgerblue"}, {"id": 77, "label": "unique_cover5_2_2", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover5_2_2", "role": "negated_conjecture", "formula": "~vertical_2_2|~vertical_1_2", "inference_rule": null, "raw_label": "cnf(unique_cover5_2_2, negated_conjecture, (~vertical_2_2|~vertical_1_2),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover5_2_2)).\n", "fillcolor": "dodgerblue"}, {"id": 78, "label": "unique_cover1_1_3", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover1_1_3", "role": "negated_conjecture", "formula": "~horizontal_1_3|~vertical_1_3", "inference_rule": null, "raw_label": "cnf(unique_cover1_1_3, negated_conjecture, (~horizontal_1_3|~vertical_1_3),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover1_1_3)).\n", "fillcolor": "dodgerblue"}, {"id": 79, "label": "unique_cover3_5_3", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover3_5_3", "role": "negated_conjecture", "formula": "~horizontal_5_3|~vertical_4_3", "inference_rule": null, "raw_label": "cnf(unique_cover3_5_3, negated_conjecture, (~horizontal_5_3|~vertical_4_3),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover3_5_3)).\n", "fillcolor": "dodgerblue"}, {"id": 80, "label": "unique_cover4_2_4", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover4_2_4", "role": "negated_conjecture", "formula": "~vertical_2_4|~horizontal_2_3", "inference_rule": null, "raw_label": "cnf(unique_cover4_2_4, negated_conjecture, (~vertical_2_4|~horizontal_2_3),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover4_2_4)).\n", "fillcolor": "dodgerblue"}, {"id": 81, "label": "unique_cover4_1_3", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover4_1_3", "role": "negated_conjecture", "formula": "~vertical_1_3|~horizontal_1_2", "inference_rule": null, "raw_label": "cnf(unique_cover4_1_3, negated_conjecture, (~vertical_1_3|~horizontal_1_2),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover4_1_3)).\n", "fillcolor": "dodgerblue"}, {"id": 82, "label": "unique_cover4_3_5", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover4_3_5", "role": "negated_conjecture", "formula": "~vertical_3_5|~horizontal_3_4", "inference_rule": null, "raw_label": "cnf(unique_cover4_3_5, negated_conjecture, (~vertical_3_5|~horizontal_3_4),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover4_3_5)).\n", "fillcolor": "dodgerblue"}, {"id": 83, "label": "unique_cover5_3_3", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover5_3_3", "role": "negated_conjecture", "formula": "~vertical_3_3|~vertical_2_3", "inference_rule": null, "raw_label": "cnf(unique_cover5_3_3, negated_conjecture, (~vertical_3_3|~vertical_2_3),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover5_3_3)).\n", "fillcolor": "dodgerblue"}, {"id": 84, "label": "covered_2_3", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "covered_2_3", "role": "negated_conjecture", "formula": "horizontal_2_3|vertical_2_3|horizontal_2_2|vertical_1_3", "inference_rule": null, "raw_label": "cnf(covered_2_3, negated_conjecture, (horizontal_2_3|vertical_2_3|horizontal_2_2|vertical_1_3),\\nfile('Problems/PUZ/PUZ015-2.006.p', covered_2_3)).\n", "fillcolor": "dodgerblue"}, {"id": 85, "label": "unique_cover1_3_1", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover1_3_1", "role": "negated_conjecture", "formula": "~horizontal_3_1|~vertical_3_1", "inference_rule": null, "raw_label": "cnf(unique_cover1_3_1, negated_conjecture, (~horizontal_3_1|~vertical_3_1),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover1_3_1)).\n", "fillcolor": "dodgerblue"}, {"id": 86, "label": "unique_cover6_4_4", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover6_4_4", "role": "negated_conjecture", "formula": "~horizontal_4_3|~vertical_3_4", "inference_rule": null, "raw_label": "cnf(unique_cover6_4_4, negated_conjecture, (~horizontal_4_3|~vertical_3_4),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover6_4_4)).\n", "fillcolor": "dodgerblue"}, {"id": 87, "label": "unique_cover3_3_3", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover3_3_3", "role": "negated_conjecture", "formula": "~horizontal_3_3|~vertical_2_3", "inference_rule": null, "raw_label": "cnf(unique_cover3_3_3, negated_conjecture, (~horizontal_3_3|~vertical_2_3),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover3_3_3)).\n", "fillcolor": "dodgerblue"}, {"id": 88, "label": "unique_cover6_3_3", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover6_3_3", "role": "negated_conjecture", "formula": "~horizontal_3_2|~vertical_2_3", "inference_rule": null, "raw_label": "cnf(unique_cover6_3_3, negated_conjecture, (~horizontal_3_2|~vertical_2_3),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover6_3_3)).\n", "fillcolor": "dodgerblue"}, {"id": 89, "label": "unique_cover6_4_6", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover6_4_6", "role": "negated_conjecture", "formula": "~horizontal_4_5|~vertical_3_6", "inference_rule": null, "raw_label": "cnf(unique_cover6_4_6, negated_conjecture, (~horizontal_4_5|~vertical_3_6),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover6_4_6)).\n", "fillcolor": "dodgerblue"}, {"id": 90, "label": "unique_cover6_6_4", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover6_6_4", "role": "negated_conjecture", "formula": "~horizontal_6_3|~vertical_5_4", "inference_rule": null, "raw_label": "cnf(unique_cover6_6_4, negated_conjecture, (~horizontal_6_3|~vertical_5_4),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover6_6_4)).\n", "fillcolor": "dodgerblue"}, {"id": 91, "label": "unique_cover4_5_5", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover4_5_5", "role": "negated_conjecture", "formula": "~vertical_5_5|~horizontal_5_4", "inference_rule": null, "raw_label": "cnf(unique_cover4_5_5, negated_conjecture, (~vertical_5_5|~horizontal_5_4),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover4_5_5)).\n", "fillcolor": "dodgerblue"}, {"id": 92, "label": "unique_cover3_5_5", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover3_5_5", "role": "negated_conjecture", "formula": "~horizontal_5_5|~vertical_4_5", "inference_rule": null, "raw_label": "cnf(unique_cover3_5_5, negated_conjecture, (~horizontal_5_5|~vertical_4_5),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover3_5_5)).\n", "fillcolor": "dodgerblue"}, {"id": 93, "label": "unique_cover1_4_4", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover1_4_4", "role": "negated_conjecture", "formula": "~horizontal_4_4|~vertical_4_4", "inference_rule": null, "raw_label": "cnf(unique_cover1_4_4, negated_conjecture, (~horizontal_4_4|~vertical_4_4),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover1_4_4)).\n", "fillcolor": "dodgerblue"}, {"id": 94, "label": "unique_cover6_4_5", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "unique_cover6_4_5", "role": "negated_conjecture", "formula": "~horizontal_4_4|~vertical_3_5", "inference_rule": null, "raw_label": "cnf(unique_cover6_4_5, negated_conjecture, (~horizontal_4_4|~vertical_3_5),\\nfile('Problems/PUZ/PUZ015-2.006.p', unique_cover6_4_5)).\n", "fillcolor": "dodgerblue"}, {"id": 95, "label": "c_0_95", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_95", "role": "negated_conjecture", "formula": "~horizontal_6_2|~horizontal_6_1", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_95, negated_conjecture, (~horizontal_6_2|~horizontal_6_1),\\ninference(fof_simplification,[status(thm)],[unique_cover2_6_2])).\n", "fillcolor": "lightskyblue1"}, {"id": 96, "label": "c_0_96", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_96", "role": "negated_conjecture", "formula": "~vertical_5_1|~vertical_4_1", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_96, negated_conjecture, (~vertical_5_1|~vertical_4_1),\\ninference(fof_simplification,[status(thm)],[unique_cover5_5_1])).\n", "fillcolor": "lightskyblue1"}, {"id": 97, "label": "c_0_97", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_97", "role": "negated_conjecture", "formula": "~horizontal_6_2|~horizontal_6_1", "inference_rule": null, "raw_label": "cnf(c_0_97, negated_conjecture, (~horizontal_6_2|~horizontal_6_1),\\nc_0_95).\n", "fillcolor": "lightskyblue1"}, {"id": 98, "label": "c_0_98", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_98", "role": "negated_conjecture", "formula": "horizontal_6_1|vertical_5_1", "inference_rule": null, "raw_label": "cnf(c_0_98, negated_conjecture, (horizontal_6_1|vertical_5_1),\\ncovered_6_1).\n", "fillcolor": "lightskyblue1"}, {"id": 99, "label": "c_0_99", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_99", "role": "negated_conjecture", "formula": "~horizontal_6_4|~horizontal_6_3", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_99, negated_conjecture, (~horizontal_6_4|~horizontal_6_3),\\ninference(fof_simplification,[status(thm)],[unique_cover2_6_4])).\n", "fillcolor": "lightskyblue1"}, {"id": 100, "label": "c_0_100", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_100", "role": "negated_conjecture", "formula": "~vertical_5_1|~vertical_4_1", "inference_rule": null, "raw_label": "cnf(c_0_100, negated_conjecture, (~vertical_5_1|~vertical_4_1),\\nc_0_96).\n", "fillcolor": "lightskyblue1"}, {"id": 101, "label": "c_0_101", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_101", "role": "negated_conjecture", "formula": "vertical_5_1|~horizontal_6_2", "inference_rule": "spm", "raw_label": "cnf(c_0_101, negated_conjecture, (vertical_5_1|~horizontal_6_2),\\ninference(spm,[status(thm)],[c_0_97, c_0_98])).\n", "fillcolor": "lightskyblue1"}, {"id": 102, "label": "c_0_102", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_102", "role": "negated_conjecture", "formula": "~horizontal_6_4|~horizontal_6_3", "inference_rule": null, "raw_label": "cnf(c_0_102, negated_conjecture, (~horizontal_6_4|~horizontal_6_3),\\nc_0_99).\n", "fillcolor": "lightskyblue1"}, {"id": 103, "label": "c_0_103", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_103", "role": "negated_conjecture", "formula": "horizontal_6_3|horizontal_6_2|vertical_5_3", "inference_rule": null, "raw_label": "cnf(c_0_103, negated_conjecture, (horizontal_6_3|horizontal_6_2|vertical_5_3),\\ncovered_6_3).\n", "fillcolor": "lightskyblue1"}, {"id": 104, "label": "c_0_104", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_104", "role": "negated_conjecture", "formula": "~horizontal_6_2|~vertical_4_1", "inference_rule": "spm", "raw_label": "cnf(c_0_104, negated_conjecture, (~horizontal_6_2|~vertical_4_1),\\ninference(spm,[status(thm)],[c_0_100, c_0_101])).\n", "fillcolor": "lightskyblue1"}, {"id": 105, "label": "c_0_105", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_105", "role": "negated_conjecture", "formula": "horizontal_6_2|vertical_5_3|~horizontal_6_4", "inference_rule": "spm", "raw_label": "cnf(c_0_105, negated_conjecture, (horizontal_6_2|vertical_5_3|~horizontal_6_4),\\ninference(spm,[status(thm)],[c_0_102, c_0_103])).\n", "fillcolor": "lightskyblue1"}, {"id": 106, "label": "c_0_106", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_106", "role": "negated_conjecture", "formula": "~horizontal_5_1|~vertical_5_1", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_106, negated_conjecture, (~horizontal_5_1|~vertical_5_1),\\ninference(fof_simplification,[status(thm)],[unique_cover1_5_1])).\n", "fillcolor": "lightskyblue1"}, {"id": 107, "label": "c_0_107", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_107", "role": "negated_conjecture", "formula": "~horizontal_6_1|~vertical_5_2", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_107, negated_conjecture, (~horizontal_6_1|~vertical_5_2),\\ninference(fof_simplification,[status(thm)],[unique_cover6_6_2])).\n", "fillcolor": "lightskyblue1"}, {"id": 108, "label": "c_0_108", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_108", "role": "negated_conjecture", "formula": "~vertical_3_1|~vertical_2_1", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_108, negated_conjecture, (~vertical_3_1|~vertical_2_1),\\ninference(fof_simplification,[status(thm)],[unique_cover5_3_1])).\n", "fillcolor": "lightskyblue1"}, {"id": 109, "label": "c_0_109", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_109", "role": "negated_conjecture", "formula": "vertical_5_3|~horizontal_6_4|~vertical_4_1", "inference_rule": "spm", "raw_label": "cnf(c_0_109, negated_conjecture, (vertical_5_3|~horizontal_6_4|~vertical_4_1),\\ninference(spm,[status(thm)],[c_0_104, c_0_105])).\n", "fillcolor": "lightskyblue1"}, {"id": 110, "label": "c_0_110", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_110", "role": "negated_conjecture", "formula": "horizontal_4_1|vertical_4_1|vertical_3_1", "inference_rule": null, "raw_label": "cnf(c_0_110, negated_conjecture, (horizontal_4_1|vertical_4_1|vertical_3_1),\\ncovered_4_1).\n", "fillcolor": "lightskyblue1"}, {"id": 111, "label": "c_0_111", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_111", "role": "negated_conjecture", "formula": "~horizontal_6_5", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_111, negated_conjecture, (~horizontal_6_5),\\ninference(fof_simplification,[status(thm)],[uncovered3_6_6])).\n", "fillcolor": "lightskyblue1"}, {"id": 112, "label": "c_0_112", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_112", "role": "negated_conjecture", "formula": "~vertical_5_6", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_112, negated_conjecture, (~vertical_5_6),\\ninference(fof_simplification,[status(thm)],[uncovered4_6_6])).\n", "fillcolor": "lightskyblue1"}, {"id": 113, "label": "c_0_113", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_113", "role": "negated_conjecture", "formula": "~horizontal_6_2|~vertical_5_2", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_113, negated_conjecture, (~horizontal_6_2|~vertical_5_2),\\ninference(fof_simplification,[status(thm)],[unique_cover3_6_2])).\n", "fillcolor": "lightskyblue1"}, {"id": 114, "label": "c_0_114", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_114", "role": "negated_conjecture", "formula": "~horizontal_5_1|~vertical_5_1", "inference_rule": null, "raw_label": "cnf(c_0_114, negated_conjecture, (~horizontal_5_1|~vertical_5_1),\\nc_0_106).\n", "fillcolor": "lightskyblue1"}, {"id": 115, "label": "c_0_115", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_115", "role": "negated_conjecture", "formula": "~horizontal_6_1|~vertical_5_2", "inference_rule": null, "raw_label": "cnf(c_0_115, negated_conjecture, (~horizontal_6_1|~vertical_5_2),\\nc_0_107).\n", "fillcolor": "lightskyblue1"}, {"id": 116, "label": "c_0_116", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_116", "role": "negated_conjecture", "formula": "~horizontal_1_5|~horizontal_1_4", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_116, negated_conjecture, (~horizontal_1_5|~horizontal_1_4),\\ninference(fof_simplification,[status(thm)],[unique_cover2_1_5])).\n", "fillcolor": "lightskyblue1"}, {"id": 117, "label": "c_0_117", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_117", "role": "negated_conjecture", "formula": "~vertical_3_1|~vertical_2_1", "inference_rule": null, "raw_label": "cnf(c_0_117, negated_conjecture, (~vertical_3_1|~vertical_2_1),\\nc_0_108).\n", "fillcolor": "lightskyblue1"}, {"id": 118, "label": "c_0_118", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_118", "role": "negated_conjecture", "formula": "vertical_5_3|horizontal_4_1|vertical_3_1|~horizontal_6_4", "inference_rule": "spm", "raw_label": "cnf(c_0_118, negated_conjecture, (vertical_5_3|horizontal_4_1|vertical_3_1|~horizontal_6_4),\\ninference(spm,[status(thm)],[c_0_109, c_0_110])).\n", "fillcolor": "lightskyblue1"}, {"id": 119, "label": "c_0_119", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_119", "role": "negated_conjecture", "formula": "horizontal_6_5|horizontal_6_4|vertical_5_5", "inference_rule": null, "raw_label": "cnf(c_0_119, negated_conjecture, (horizontal_6_5|horizontal_6_4|vertical_5_5),\\ncovered_6_5).\n", "fillcolor": "lightskyblue1"}, {"id": 120, "label": "c_0_120", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_120", "role": "negated_conjecture", "formula": "~horizontal_6_5", "inference_rule": null, "raw_label": "cnf(c_0_120, negated_conjecture, (~horizontal_6_5),\\nc_0_111).\n", "fillcolor": "lightskyblue1"}, {"id": 121, "label": "c_0_121", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_121", "role": "negated_conjecture", "formula": "~vertical_1_1", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_121, negated_conjecture, (~vertical_1_1),\\ninference(fof_simplification,[status(thm)],[uncovered2_1_1])).\n", "fillcolor": "lightskyblue1"}, {"id": 122, "label": "c_0_122", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_122", "role": "negated_conjecture", "formula": "~vertical_4_6|~horizontal_4_5", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_122, negated_conjecture, (~vertical_4_6|~horizontal_4_5),\\ninference(fof_simplification,[status(thm)],[unique_cover4_4_6])).\n", "fillcolor": "lightskyblue1"}, {"id": 123, "label": "c_0_123", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_123", "role": "negated_conjecture", "formula": "vertical_5_6|horizontal_5_5|vertical_4_6", "inference_rule": null, "raw_label": "cnf(c_0_123, negated_conjecture, (vertical_5_6|horizontal_5_5|vertical_4_6),\\ncovered_5_6).\n", "fillcolor": "lightskyblue1"}, {"id": 124, "label": "c_0_124", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_124", "role": "negated_conjecture", "formula": "~vertical_5_6", "inference_rule": null, "raw_label": "cnf(c_0_124, negated_conjecture, (~vertical_5_6),\\nc_0_112).\n", "fillcolor": "lightskyblue1"}, {"id": 125, "label": "c_0_125", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_125", "role": "negated_conjecture", "formula": "~horizontal_6_2|~vertical_5_2", "inference_rule": null, "raw_label": "cnf(c_0_125, negated_conjecture, (~horizontal_6_2|~vertical_5_2),\\nc_0_113).\n", "fillcolor": "lightskyblue1"}, {"id": 126, "label": "c_0_126", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_126", "role": "negated_conjecture", "formula": "~horizontal_6_2|~horizontal_5_1", "inference_rule": "spm", "raw_label": "cnf(c_0_126, negated_conjecture, (~horizontal_6_2|~horizontal_5_1),\\ninference(spm,[status(thm)],[c_0_114, c_0_101])).\n", "fillcolor": "lightskyblue1"}, {"id": 127, "label": "c_0_127", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_127", "role": "negated_conjecture", "formula": "vertical_5_1|~vertical_5_2", "inference_rule": "spm", "raw_label": "cnf(c_0_127, negated_conjecture, (vertical_5_1|~vertical_5_2),\\ninference(spm,[status(thm)],[c_0_115, c_0_98])).\n", "fillcolor": "lightskyblue1"}, {"id": 128, "label": "c_0_128", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_128", "role": "negated_conjecture", "formula": "~horizontal_5_1|~vertical_4_1", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_128, negated_conjecture, (~horizontal_5_1|~vertical_4_1),\\ninference(fof_simplification,[status(thm)],[unique_cover3_5_1])).\n", "fillcolor": "lightskyblue1"}, {"id": 129, "label": "c_0_129", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_129", "role": "negated_conjecture", "formula": "~vertical_2_6|~vertical_1_6", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_129, negated_conjecture, (~vertical_2_6|~vertical_1_6),\\ninference(fof_simplification,[status(thm)],[unique_cover5_2_6])).\n", "fillcolor": "lightskyblue1"}, {"id": 130, "label": "c_0_130", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_130", "role": "negated_conjecture", "formula": "~horizontal_1_5|~horizontal_1_4", "inference_rule": null, "raw_label": "cnf(c_0_130, negated_conjecture, (~horizontal_1_5|~horizontal_1_4),\\nc_0_116).\n", "fillcolor": "lightskyblue1"}, {"id": 131, "label": "c_0_131", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_131", "role": "negated_conjecture", "formula": "vertical_1_6|horizontal_1_5", "inference_rule": null, "raw_label": "cnf(c_0_131, negated_conjecture, (vertical_1_6|horizontal_1_5),\\ncovered_1_6).\n", "fillcolor": "lightskyblue1"}, {"id": 132, "label": "c_0_132", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_132", "role": "negated_conjecture", "formula": "~vertical_4_6|~vertical_3_6", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_132, negated_conjecture, (~vertical_4_6|~vertical_3_6),\\ninference(fof_simplification,[status(thm)],[unique_cover5_4_6])).\n", "fillcolor": "lightskyblue1"}, {"id": 133, "label": "c_0_133", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_133", "role": "negated_conjecture", "formula": "vertical_5_3|horizontal_4_1|~horizontal_6_4|~vertical_2_1", "inference_rule": "spm", "raw_label": "cnf(c_0_133, negated_conjecture, (vertical_5_3|horizontal_4_1|~horizontal_6_4|~vertical_2_1),\\ninference(spm,[status(thm)],[c_0_117, c_0_118])).\n", "fillcolor": "lightskyblue1"}, {"id": 134, "label": "c_0_134", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_134", "role": "negated_conjecture", "formula": "horizontal_6_4|vertical_5_5", "inference_rule": "sr", "raw_label": "cnf(c_0_134, negated_conjecture, (horizontal_6_4|vertical_5_5),\\ninference(sr,[status(thm)],[c_0_119, c_0_120])).\n", "fillcolor": "lightskyblue1"}, {"id": 135, "label": "c_0_135", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_135", "role": "negated_conjecture", "formula": "horizontal_2_1|vertical_2_1|vertical_1_1", "inference_rule": null, "raw_label": "cnf(c_0_135, negated_conjecture, (horizontal_2_1|vertical_2_1|vertical_1_1),\\ncovered_2_1).\n", "fillcolor": "lightskyblue1"}, {"id": 136, "label": "c_0_136", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_136", "role": "negated_conjecture", "formula": "~vertical_1_1", "inference_rule": null, "raw_label": "cnf(c_0_136, negated_conjecture, (~vertical_1_1),\\nc_0_121).\n", "fillcolor": "lightskyblue1"}, {"id": 137, "label": "c_0_137", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_137", "role": "negated_conjecture", "formula": "~vertical_4_6|~horizontal_4_5", "inference_rule": null, "raw_label": "cnf(c_0_137, negated_conjecture, (~vertical_4_6|~horizontal_4_5),\\nc_0_122).\n", "fillcolor": "lightskyblue1"}, {"id": 138, "label": "c_0_138", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_138", "role": "negated_conjecture", "formula": "horizontal_5_5|vertical_4_6", "inference_rule": "sr", "raw_label": "cnf(c_0_138, negated_conjecture, (horizontal_5_5|vertical_4_6),\\ninference(sr,[status(thm)],[c_0_123, c_0_124])).\n", "fillcolor": "lightskyblue1"}, {"id": 139, "label": "c_0_139", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_139", "role": "negated_conjecture", "formula": "~horizontal_5_3|~vertical_5_3", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_139, negated_conjecture, (~horizontal_5_3|~vertical_5_3),\\ninference(fof_simplification,[status(thm)],[unique_cover1_5_3])).\n", "fillcolor": "lightskyblue1"}, {"id": 140, "label": "c_0_140", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_140", "role": "negated_conjecture", "formula": "vertical_5_3|~horizontal_6_4|~vertical_5_2", "inference_rule": "spm", "raw_label": "cnf(c_0_140, negated_conjecture, (vertical_5_3|~horizontal_6_4|~vertical_5_2),\\ninference(spm,[status(thm)],[c_0_125, c_0_105])).\n", "fillcolor": "lightskyblue1"}, {"id": 141, "label": "c_0_141", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_141", "role": "negated_conjecture", "formula": "vertical_5_3|~horizontal_6_4|~horizontal_5_1", "inference_rule": "spm", "raw_label": "cnf(c_0_141, negated_conjecture, (vertical_5_3|~horizontal_6_4|~horizontal_5_1),\\ninference(spm,[status(thm)],[c_0_126, c_0_105])).\n", "fillcolor": "lightskyblue1"}, {"id": 142, "label": "c_0_142", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_142", "role": "negated_conjecture", "formula": "~vertical_5_2|~vertical_4_1", "inference_rule": "spm", "raw_label": "cnf(c_0_142, negated_conjecture, (~vertical_5_2|~vertical_4_1),\\ninference(spm,[status(thm)],[c_0_100, c_0_127])).\n", "fillcolor": "lightskyblue1"}, {"id": 143, "label": "c_0_143", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_143", "role": "negated_conjecture", "formula": "~horizontal_5_1|~vertical_4_1", "inference_rule": null, "raw_label": "cnf(c_0_143, negated_conjecture, (~horizontal_5_1|~vertical_4_1),\\nc_0_128).\n", "fillcolor": "lightskyblue1"}, {"id": 144, "label": "c_0_144", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_144", "role": "negated_conjecture", "formula": "~vertical_2_6|~vertical_1_6", "inference_rule": null, "raw_label": "cnf(c_0_144, negated_conjecture, (~vertical_2_6|~vertical_1_6),\\nc_0_129).\n", "fillcolor": "lightskyblue1"}, {"id": 145, "label": "c_0_145", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_145", "role": "negated_conjecture", "formula": "vertical_1_6|~horizontal_1_4", "inference_rule": "spm", "raw_label": "cnf(c_0_145, negated_conjecture, (vertical_1_6|~horizontal_1_4),\\ninference(spm,[status(thm)],[c_0_130, c_0_131])).\n", "fillcolor": "lightskyblue1"}, {"id": 146, "label": "c_0_146", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_146", "role": "negated_conjecture", "formula": "~vertical_4_6|~vertical_3_6", "inference_rule": null, "raw_label": "cnf(c_0_146, negated_conjecture, (~vertical_4_6|~vertical_3_6),\\nc_0_132).\n", "fillcolor": "lightskyblue1"}, {"id": 147, "label": "c_0_147", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_147", "role": "negated_conjecture", "formula": "vertical_3_6|horizontal_3_5|vertical_2_6", "inference_rule": null, "raw_label": "cnf(c_0_147, negated_conjecture, (vertical_3_6|horizontal_3_5|vertical_2_6),\\ncovered_3_6).\n", "fillcolor": "lightskyblue1"}, {"id": 148, "label": "c_0_148", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_148", "role": "negated_conjecture", "formula": "vertical_5_5|vertical_5_3|horizontal_4_1|~vertical_2_1", "inference_rule": "spm", "raw_label": "cnf(c_0_148, negated_conjecture, (vertical_5_5|vertical_5_3|horizontal_4_1|~vertical_2_1),\\ninference(spm,[status(thm)],[c_0_133, c_0_134])).\n", "fillcolor": "lightskyblue1"}, {"id": 149, "label": "c_0_149", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_149", "role": "negated_conjecture", "formula": "vertical_2_1|horizontal_2_1", "inference_rule": "sr", "raw_label": "cnf(c_0_149, negated_conjecture, (vertical_2_1|horizontal_2_1),\\ninference(sr,[status(thm)],[c_0_135, c_0_136])).\n", "fillcolor": "lightskyblue1"}, {"id": 150, "label": "c_0_150", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_150", "role": "negated_conjecture", "formula": "~horizontal_5_5|~vertical_5_5", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_150, negated_conjecture, (~horizontal_5_5|~vertical_5_5),\\ninference(fof_simplification,[status(thm)],[unique_cover1_5_5])).\n", "fillcolor": "lightskyblue1"}, {"id": 151, "label": "c_0_151", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_151", "role": "negated_conjecture", "formula": "horizontal_5_5|~horizontal_4_5", "inference_rule": "spm", "raw_label": "cnf(c_0_151, negated_conjecture, (horizontal_5_5|~horizontal_4_5),\\ninference(spm,[status(thm)],[c_0_137, c_0_138])).\n", "fillcolor": "lightskyblue1"}, {"id": 152, "label": "c_0_152", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_152", "role": "negated_conjecture", "formula": "horizontal_4_5|vertical_4_5|horizontal_4_4|vertical_3_5", "inference_rule": null, "raw_label": "cnf(c_0_152, negated_conjecture, (horizontal_4_5|vertical_4_5|horizontal_4_4|vertical_3_5),\\ncovered_4_5).\n", "fillcolor": "lightskyblue1"}, {"id": 153, "label": "c_0_153", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_153", "role": "negated_conjecture", "formula": "~vertical_5_5|~vertical_4_5", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_153, negated_conjecture, (~vertical_5_5|~vertical_4_5),\\ninference(fof_simplification,[status(thm)],[unique_cover5_5_5])).\n", "fillcolor": "lightskyblue1"}, {"id": 154, "label": "c_0_154", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_154", "role": "negated_conjecture", "formula": "~horizontal_5_3|~vertical_5_3", "inference_rule": null, "raw_label": "cnf(c_0_154, negated_conjecture, (~horizontal_5_3|~vertical_5_3),\\nc_0_139).\n", "fillcolor": "lightskyblue1"}, {"id": 155, "label": "c_0_155", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_155", "role": "negated_conjecture", "formula": "vertical_5_5|vertical_5_3|~vertical_5_2", "inference_rule": "spm", "raw_label": "cnf(c_0_155, negated_conjecture, (vertical_5_5|vertical_5_3|~vertical_5_2),\\ninference(spm,[status(thm)],[c_0_140, c_0_134])).\n", "fillcolor": "lightskyblue1"}, {"id": 156, "label": "c_0_156", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_156", "role": "negated_conjecture", "formula": "vertical_5_5|vertical_5_3|~horizontal_5_1", "inference_rule": "spm", "raw_label": "cnf(c_0_156, negated_conjecture, (vertical_5_5|vertical_5_3|~horizontal_5_1),\\ninference(spm,[status(thm)],[c_0_141, c_0_134])).\n", "fillcolor": "lightskyblue1"}, {"id": 157, "label": "c_0_157", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_157", "role": "negated_conjecture", "formula": "~horizontal_5_3|~horizontal_5_2", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_157, negated_conjecture, (~horizontal_5_3|~horizontal_5_2),\\ninference(fof_simplification,[status(thm)],[unique_cover2_5_3])).\n", "fillcolor": "lightskyblue1"}, {"id": 158, "label": "c_0_158", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_158", "role": "negated_conjecture", "formula": "horizontal_4_1|vertical_3_1|~vertical_5_2", "inference_rule": "spm", "raw_label": "cnf(c_0_158, negated_conjecture, (horizontal_4_1|vertical_3_1|~vertical_5_2),\\ninference(spm,[status(thm)],[c_0_142, c_0_110])).\n", "fillcolor": "lightskyblue1"}, {"id": 159, "label": "c_0_159", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_159", "role": "negated_conjecture", "formula": "horizontal_4_1|vertical_3_1|~horizontal_5_1", "inference_rule": "spm", "raw_label": "cnf(c_0_159, negated_conjecture, (horizontal_4_1|vertical_3_1|~horizontal_5_1),\\ninference(spm,[status(thm)],[c_0_143, c_0_110])).\n", "fillcolor": "lightskyblue1"}, {"id": 160, "label": "c_0_160", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_160", "role": "negated_conjecture", "formula": "~vertical_2_6|~horizontal_1_4", "inference_rule": "spm", "raw_label": "cnf(c_0_160, negated_conjecture, (~vertical_2_6|~horizontal_1_4),\\ninference(spm,[status(thm)],[c_0_144, c_0_145])).\n", "fillcolor": "lightskyblue1"}, {"id": 161, "label": "c_0_161", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_161", "role": "negated_conjecture", "formula": "horizontal_3_5|vertical_2_6|~vertical_4_6", "inference_rule": "spm", "raw_label": "cnf(c_0_161, negated_conjecture, (horizontal_3_5|vertical_2_6|~vertical_4_6),\\ninference(spm,[status(thm)],[c_0_146, c_0_147])).\n", "fillcolor": "lightskyblue1"}, {"id": 162, "label": "c_0_162", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_162", "role": "negated_conjecture", "formula": "~vertical_4_2|~horizontal_4_1", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_162, negated_conjecture, (~vertical_4_2|~horizontal_4_1),\\ninference(fof_simplification,[status(thm)],[unique_cover4_4_2])).\n", "fillcolor": "lightskyblue1"}, {"id": 163, "label": "c_0_163", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_163", "role": "negated_conjecture", "formula": "vertical_5_5|vertical_5_3|horizontal_4_1|horizontal_2_1", "inference_rule": "spm", "raw_label": "cnf(c_0_163, negated_conjecture, (vertical_5_5|vertical_5_3|horizontal_4_1|horizontal_2_1),\\ninference(spm,[status(thm)],[c_0_148, c_0_149])).\n", "fillcolor": "lightskyblue1"}, {"id": 164, "label": "c_0_164", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_164", "role": "negated_conjecture", "formula": "~horizontal_5_5|~vertical_5_5", "inference_rule": null, "raw_label": "cnf(c_0_164, negated_conjecture, (~horizontal_5_5|~vertical_5_5),\\nc_0_150).\n", "fillcolor": "lightskyblue1"}, {"id": 165, "label": "c_0_165", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_165", "role": "negated_conjecture", "formula": "horizontal_5_5|vertical_4_5|horizontal_4_4|vertical_3_5", "inference_rule": "spm", "raw_label": "cnf(c_0_165, negated_conjecture, (horizontal_5_5|vertical_4_5|horizontal_4_4|vertical_3_5),\\ninference(spm,[status(thm)],[c_0_151, c_0_152])).\n", "fillcolor": "lightskyblue1"}, {"id": 166, "label": "c_0_166", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_166", "role": "negated_conjecture", "formula": "~vertical_5_5|~vertical_4_5", "inference_rule": null, "raw_label": "cnf(c_0_166, negated_conjecture, (~vertical_5_5|~vertical_4_5),\\nc_0_153).\n", "fillcolor": "lightskyblue1"}, {"id": 167, "label": "c_0_167", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_167", "role": "negated_conjecture", "formula": "vertical_5_5|~horizontal_5_3|~vertical_5_2", "inference_rule": "spm", "raw_label": "cnf(c_0_167, negated_conjecture, (vertical_5_5|~horizontal_5_3|~vertical_5_2),\\ninference(spm,[status(thm)],[c_0_154, c_0_155])).\n", "fillcolor": "lightskyblue1"}, {"id": 168, "label": "c_0_168", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_168", "role": "negated_conjecture", "formula": "horizontal_5_2|vertical_5_2|horizontal_5_1|vertical_4_2", "inference_rule": null, "raw_label": "cnf(c_0_168, negated_conjecture, (horizontal_5_2|vertical_5_2|horizontal_5_1|vertical_4_2),\\ncovered_5_2).\n", "fillcolor": "lightskyblue1"}, {"id": 169, "label": "c_0_169", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_169", "role": "negated_conjecture", "formula": "vertical_5_5|~horizontal_5_3|~horizontal_5_1", "inference_rule": "spm", "raw_label": "cnf(c_0_169, negated_conjecture, (vertical_5_5|~horizontal_5_3|~horizontal_5_1),\\ninference(spm,[status(thm)],[c_0_154, c_0_156])).\n", "fillcolor": "lightskyblue1"}, {"id": 170, "label": "c_0_170", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_170", "role": "negated_conjecture", "formula": "~horizontal_5_3|~horizontal_5_2", "inference_rule": null, "raw_label": "cnf(c_0_170, negated_conjecture, (~horizontal_5_3|~horizontal_5_2),\\nc_0_157).\n", "fillcolor": "lightskyblue1"}, {"id": 171, "label": "c_0_171", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_171", "role": "negated_conjecture", "formula": "~horizontal_6_4|~vertical_5_4", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_171, negated_conjecture, (~horizontal_6_4|~vertical_5_4),\\ninference(fof_simplification,[status(thm)],[unique_cover3_6_4])).\n", "fillcolor": "lightskyblue1"}, {"id": 172, "label": "c_0_172", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_172", "role": "negated_conjecture", "formula": "horizontal_4_1|~vertical_5_2|~vertical_2_1", "inference_rule": "spm", "raw_label": "cnf(c_0_172, negated_conjecture, (horizontal_4_1|~vertical_5_2|~vertical_2_1),\\ninference(spm,[status(thm)],[c_0_117, c_0_158])).\n", "fillcolor": "lightskyblue1"}, {"id": 173, "label": "c_0_173", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_173", "role": "negated_conjecture", "formula": "horizontal_4_1|~horizontal_5_1|~vertical_2_1", "inference_rule": "spm", "raw_label": "cnf(c_0_173, negated_conjecture, (horizontal_4_1|~horizontal_5_1|~vertical_2_1),\\ninference(spm,[status(thm)],[c_0_117, c_0_159])).\n", "fillcolor": "lightskyblue1"}, {"id": 174, "label": "c_0_174", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_174", "role": "negated_conjecture", "formula": "~horizontal_1_3|~horizontal_1_2", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_174, negated_conjecture, (~horizontal_1_3|~horizontal_1_2),\\ninference(fof_simplification,[status(thm)],[unique_cover2_1_3])).\n", "fillcolor": "lightskyblue1"}, {"id": 175, "label": "c_0_175", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_175", "role": "negated_conjecture", "formula": "horizontal_3_5|~vertical_4_6|~horizontal_1_4", "inference_rule": "spm", "raw_label": "cnf(c_0_175, negated_conjecture, (horizontal_3_5|~vertical_4_6|~horizontal_1_4),\\ninference(spm,[status(thm)],[c_0_160, c_0_161])).\n", "fillcolor": "lightskyblue1"}, {"id": 176, "label": "c_0_176", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_176", "role": "negated_conjecture", "formula": "horizontal_1_4|vertical_1_4|horizontal_1_3", "inference_rule": null, "raw_label": "cnf(c_0_176, negated_conjecture, (horizontal_1_4|vertical_1_4|horizontal_1_3),\\ncovered_1_4).\n", "fillcolor": "lightskyblue1"}, {"id": 177, "label": "c_0_177", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_177", "role": "negated_conjecture", "formula": "~horizontal_2_5|~vertical_1_6", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_177, negated_conjecture, (~horizontal_2_5|~vertical_1_6),\\ninference(fof_simplification,[status(thm)],[unique_cover6_2_6])).\n", "fillcolor": "lightskyblue1"}, {"id": 178, "label": "c_0_178", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_178", "role": "negated_conjecture", "formula": "~vertical_4_2|~horizontal_4_1", "inference_rule": null, "raw_label": "cnf(c_0_178, negated_conjecture, (~vertical_4_2|~horizontal_4_1),\\nc_0_162).\n", "fillcolor": "lightskyblue1"}, {"id": 179, "label": "c_0_179", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_179", "role": "negated_conjecture", "formula": "vertical_5_5|horizontal_4_1|horizontal_2_1|~horizontal_5_3", "inference_rule": "spm", "raw_label": "cnf(c_0_179, negated_conjecture, (vertical_5_5|horizontal_4_1|horizontal_2_1|~horizontal_5_3),\\ninference(spm,[status(thm)],[c_0_154, c_0_163])).\n", "fillcolor": "lightskyblue1"}, {"id": 180, "label": "c_0_180", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_180", "role": "negated_conjecture", "formula": "horizontal_4_4|vertical_3_5|~vertical_5_5", "inference_rule": "csr", "raw_label": "cnf(c_0_180, negated_conjecture, (horizontal_4_4|vertical_3_5|~vertical_5_5),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_164, c_0_165]), c_0_166])).\n", "fillcolor": "lightskyblue1"}, {"id": 181, "label": "c_0_181", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_181", "role": "negated_conjecture", "formula": "vertical_5_5|vertical_4_2|~horizontal_5_3", "inference_rule": "csr", "raw_label": "cnf(c_0_181, negated_conjecture, (vertical_5_5|vertical_4_2|~horizontal_5_3),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_167, c_0_168]), c_0_169]), c_0_170])).\n", "fillcolor": "lightskyblue1"}, {"id": 182, "label": "c_0_182", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_182", "role": "negated_conjecture", "formula": "~horizontal_6_4|~vertical_5_4", "inference_rule": null, "raw_label": "cnf(c_0_182, negated_conjecture, (~horizontal_6_4|~vertical_5_4),\\nc_0_171).\n", "fillcolor": "lightskyblue1"}, {"id": 183, "label": "c_0_183", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_183", "role": "negated_conjecture", "formula": "~horizontal_5_5|~horizontal_5_4", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_183, negated_conjecture, (~horizontal_5_5|~horizontal_5_4),\\ninference(fof_simplification,[status(thm)],[unique_cover2_5_5])).\n", "fillcolor": "lightskyblue1"}, {"id": 184, "label": "c_0_184", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_184", "role": "negated_conjecture", "formula": "~horizontal_5_4|~vertical_4_5", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_184, negated_conjecture, (~horizontal_5_4|~vertical_4_5),\\ninference(fof_simplification,[status(thm)],[unique_cover6_5_5])).\n", "fillcolor": "lightskyblue1"}, {"id": 185, "label": "c_0_185", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_185", "role": "negated_conjecture", "formula": "~horizontal_4_1|~vertical_3_2", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_185, negated_conjecture, (~horizontal_4_1|~vertical_3_2),\\ninference(fof_simplification,[status(thm)],[unique_cover6_4_2])).\n", "fillcolor": "lightskyblue1"}, {"id": 186, "label": "c_0_186", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_186", "role": "negated_conjecture", "formula": "horizontal_4_1|horizontal_2_1|~vertical_5_2", "inference_rule": "spm", "raw_label": "cnf(c_0_186, negated_conjecture, (horizontal_4_1|horizontal_2_1|~vertical_5_2),\\ninference(spm,[status(thm)],[c_0_172, c_0_149])).\n", "fillcolor": "lightskyblue1"}, {"id": 187, "label": "c_0_187", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_187", "role": "negated_conjecture", "formula": "horizontal_4_1|horizontal_2_1|~horizontal_5_1", "inference_rule": "spm", "raw_label": "cnf(c_0_187, negated_conjecture, (horizontal_4_1|horizontal_2_1|~horizontal_5_1),\\ninference(spm,[status(thm)],[c_0_173, c_0_149])).\n", "fillcolor": "lightskyblue1"}, {"id": 188, "label": "c_0_188", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_188", "role": "negated_conjecture", "formula": "~vertical_4_2|~vertical_3_2", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_188, negated_conjecture, (~vertical_4_2|~vertical_3_2),\\ninference(fof_simplification,[status(thm)],[unique_cover5_4_2])).\n", "fillcolor": "lightskyblue1"}, {"id": 189, "label": "c_0_189", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_189", "role": "negated_conjecture", "formula": "~horizontal_1_3|~horizontal_1_2", "inference_rule": null, "raw_label": "cnf(c_0_189, negated_conjecture, (~horizontal_1_3|~horizontal_1_2),\\nc_0_174).\n", "fillcolor": "lightskyblue1"}, {"id": 190, "label": "c_0_190", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_190", "role": "negated_conjecture", "formula": "horizontal_3_5|vertical_1_4|horizontal_1_3|~vertical_4_6", "inference_rule": "spm", "raw_label": "cnf(c_0_190, negated_conjecture, (horizontal_3_5|vertical_1_4|horizontal_1_3|~vertical_4_6),\\ninference(spm,[status(thm)],[c_0_175, c_0_176])).\n", "fillcolor": "lightskyblue1"}, {"id": 191, "label": "c_0_191", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_191", "role": "negated_conjecture", "formula": "~horizontal_1_1", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_191, negated_conjecture, (~horizontal_1_1),\\ninference(fof_simplification,[status(thm)],[uncovered1_1_1])).\n", "fillcolor": "lightskyblue1"}, {"id": 192, "label": "c_0_192", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_192", "role": "negated_conjecture", "formula": "~horizontal_2_5|~vertical_1_6", "inference_rule": null, "raw_label": "cnf(c_0_192, negated_conjecture, (~horizontal_2_5|~vertical_1_6),\\nc_0_177).\n", "fillcolor": "lightskyblue1"}, {"id": 193, "label": "c_0_193", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_193", "role": "negated_conjecture", "formula": "~vertical_1_5|~horizontal_1_4", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_193, negated_conjecture, (~vertical_1_5|~horizontal_1_4),\\ninference(fof_simplification,[status(thm)],[unique_cover4_1_5])).\n", "fillcolor": "lightskyblue1"}, {"id": 194, "label": "c_0_194", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_194", "role": "negated_conjecture", "formula": "~horizontal_1_5|~vertical_1_5", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_194, negated_conjecture, (~horizontal_1_5|~vertical_1_5),\\ninference(fof_simplification,[status(thm)],[unique_cover1_1_5])).\n", "fillcolor": "lightskyblue1"}, {"id": 195, "label": "c_0_195", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_195", "role": "negated_conjecture", "formula": "~horizontal_2_2|~horizontal_2_1", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_195, negated_conjecture, (~horizontal_2_2|~horizontal_2_1),\\ninference(fof_simplification,[status(thm)],[unique_cover2_2_2])).\n", "fillcolor": "lightskyblue1"}, {"id": 196, "label": "c_0_196", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_196", "role": "negated_conjecture", "formula": "vertical_5_5|horizontal_2_1|~horizontal_5_3|~vertical_4_2", "inference_rule": "spm", "raw_label": "cnf(c_0_196, negated_conjecture, (vertical_5_5|horizontal_2_1|~horizontal_5_3|~vertical_4_2),\\ninference(spm,[status(thm)],[c_0_178, c_0_179])).\n", "fillcolor": "lightskyblue1"}, {"id": 197, "label": "c_0_197", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_197", "role": "negated_conjecture", "formula": "horizontal_4_4|vertical_4_2|vertical_3_5|~horizontal_5_3", "inference_rule": "spm", "raw_label": "cnf(c_0_197, negated_conjecture, (horizontal_4_4|vertical_4_2|vertical_3_5|~horizontal_5_3),\\ninference(spm,[status(thm)],[c_0_180, c_0_181])).\n", "fillcolor": "lightskyblue1"}, {"id": 198, "label": "c_0_198", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_198", "role": "negated_conjecture", "formula": "vertical_5_5|~vertical_5_4", "inference_rule": "spm", "raw_label": "cnf(c_0_198, negated_conjecture, (vertical_5_5|~vertical_5_4),\\ninference(spm,[status(thm)],[c_0_182, c_0_134])).\n", "fillcolor": "lightskyblue1"}, {"id": 199, "label": "c_0_199", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_199", "role": "negated_conjecture", "formula": "horizontal_5_4|vertical_5_4|horizontal_5_3|vertical_4_4", "inference_rule": null, "raw_label": "cnf(c_0_199, negated_conjecture, (horizontal_5_4|vertical_5_4|horizontal_5_3|vertical_4_4),\\ncovered_5_4).\n", "fillcolor": "lightskyblue1"}, {"id": 200, "label": "c_0_200", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_200", "role": "negated_conjecture", "formula": "~horizontal_5_5|~horizontal_5_4", "inference_rule": null, "raw_label": "cnf(c_0_200, negated_conjecture, (~horizontal_5_5|~horizontal_5_4),\\nc_0_183).\n", "fillcolor": "lightskyblue1"}, {"id": 201, "label": "c_0_201", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_201", "role": "negated_conjecture", "formula": "~horizontal_5_4|~vertical_4_5", "inference_rule": null, "raw_label": "cnf(c_0_201, negated_conjecture, (~horizontal_5_4|~vertical_4_5),\\nc_0_184).\n", "fillcolor": "lightskyblue1"}, {"id": 202, "label": "c_0_202", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_202", "role": "negated_conjecture", "formula": "~horizontal_4_1|~vertical_3_2", "inference_rule": null, "raw_label": "cnf(c_0_202, negated_conjecture, (~horizontal_4_1|~vertical_3_2),\\nc_0_185).\n", "fillcolor": "lightskyblue1"}, {"id": 203, "label": "c_0_203", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_203", "role": "negated_conjecture", "formula": "horizontal_5_2|vertical_4_2|horizontal_4_1|horizontal_2_1", "inference_rule": "csr", "raw_label": "cnf(c_0_203, negated_conjecture, (horizontal_5_2|vertical_4_2|horizontal_4_1|horizontal_2_1),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_186, c_0_168]), c_0_187])).\n", "fillcolor": "lightskyblue1"}, {"id": 204, "label": "c_0_204", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_204", "role": "negated_conjecture", "formula": "~vertical_4_2|~vertical_3_2", "inference_rule": null, "raw_label": "cnf(c_0_204, negated_conjecture, (~vertical_4_2|~vertical_3_2),\\nc_0_188).\n", "fillcolor": "lightskyblue1"}, {"id": 205, "label": "c_0_205", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_205", "role": "negated_conjecture", "formula": "horizontal_3_5|vertical_1_4|~vertical_4_6|~horizontal_1_2", "inference_rule": "spm", "raw_label": "cnf(c_0_205, negated_conjecture, (horizontal_3_5|vertical_1_4|~vertical_4_6|~horizontal_1_2),\\ninference(spm,[status(thm)],[c_0_189, c_0_190])).\n", "fillcolor": "lightskyblue1"}, {"id": 206, "label": "c_0_206", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_206", "role": "negated_conjecture", "formula": "horizontal_1_2|vertical_1_2|horizontal_1_1", "inference_rule": null, "raw_label": "cnf(c_0_206, negated_conjecture, (horizontal_1_2|vertical_1_2|horizontal_1_1),\\ncovered_1_2).\n", "fillcolor": "lightskyblue1"}, {"id": 207, "label": "c_0_207", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_207", "role": "negated_conjecture", "formula": "~horizontal_1_1", "inference_rule": null, "raw_label": "cnf(c_0_207, negated_conjecture, (~horizontal_1_1),\\nc_0_191).\n", "fillcolor": "lightskyblue1"}, {"id": 208, "label": "c_0_208", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_208", "role": "negated_conjecture", "formula": "~horizontal_2_5|~horizontal_1_4", "inference_rule": "spm", "raw_label": "cnf(c_0_208, negated_conjecture, (~horizontal_2_5|~horizontal_1_4),\\ninference(spm,[status(thm)],[c_0_192, c_0_145])).\n", "fillcolor": "lightskyblue1"}, {"id": 209, "label": "c_0_209", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_209", "role": "negated_conjecture", "formula": "~vertical_1_5|~horizontal_1_4", "inference_rule": null, "raw_label": "cnf(c_0_209, negated_conjecture, (~vertical_1_5|~horizontal_1_4),\\nc_0_193).\n", "fillcolor": "lightskyblue1"}, {"id": 210, "label": "c_0_210", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_210", "role": "negated_conjecture", "formula": "~horizontal_1_5|~vertical_1_5", "inference_rule": null, "raw_label": "cnf(c_0_210, negated_conjecture, (~horizontal_1_5|~vertical_1_5),\\nc_0_194).\n", "fillcolor": "lightskyblue1"}, {"id": 211, "label": "c_0_211", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_211", "role": "negated_conjecture", "formula": "~horizontal_2_2|~horizontal_2_1", "inference_rule": null, "raw_label": "cnf(c_0_211, negated_conjecture, (~horizontal_2_2|~horizontal_2_1),\\nc_0_195).\n", "fillcolor": "lightskyblue1"}, {"id": 212, "label": "c_0_212", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_212", "role": "negated_conjecture", "formula": "horizontal_4_4|vertical_3_5|horizontal_2_1|~horizontal_5_3", "inference_rule": "csr", "raw_label": "cnf(c_0_212, negated_conjecture, (horizontal_4_4|vertical_3_5|horizontal_2_1|~horizontal_5_3),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_180, c_0_196]), c_0_197])).\n", "fillcolor": "lightskyblue1"}, {"id": 213, "label": "c_0_213", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_213", "role": "negated_conjecture", "formula": "vertical_5_5|horizontal_5_4|horizontal_5_3|vertical_4_4", "inference_rule": "spm", "raw_label": "cnf(c_0_213, negated_conjecture, (vertical_5_5|horizontal_5_4|horizontal_5_3|vertical_4_4),\\ninference(spm,[status(thm)],[c_0_198, c_0_199])).\n", "fillcolor": "lightskyblue1"}, {"id": 214, "label": "c_0_214", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_214", "role": "negated_conjecture", "formula": "horizontal_4_4|vertical_3_5|~horizontal_5_4", "inference_rule": "csr", "raw_label": "cnf(c_0_214, negated_conjecture, (horizontal_4_4|vertical_3_5|~horizontal_5_4),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_200, c_0_165]), c_0_201])).\n", "fillcolor": "lightskyblue1"}, {"id": 215, "label": "c_0_215", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_215", "role": "negated_conjecture", "formula": "~horizontal_5_2|~vertical_4_3", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_215, negated_conjecture, (~horizontal_5_2|~vertical_4_3),\\ninference(fof_simplification,[status(thm)],[unique_cover6_5_3])).\n", "fillcolor": "lightskyblue1"}, {"id": 216, "label": "c_0_216", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_216", "role": "negated_conjecture", "formula": "horizontal_5_2|horizontal_2_1|~vertical_3_2", "inference_rule": "csr", "raw_label": "cnf(c_0_216, negated_conjecture, (horizontal_5_2|horizontal_2_1|~vertical_3_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_202, c_0_203]), c_0_204])).\n", "fillcolor": "lightskyblue1"}, {"id": 217, "label": "c_0_217", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_217", "role": "negated_conjecture", "formula": "~horizontal_3_5|~vertical_2_5", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_217, negated_conjecture, (~horizontal_3_5|~vertical_2_5),\\ninference(fof_simplification,[status(thm)],[unique_cover3_3_5])).\n", "fillcolor": "lightskyblue1"}, {"id": 218, "label": "c_0_218", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_218", "role": "negated_conjecture", "formula": "horizontal_5_5|horizontal_3_5|vertical_1_4|~horizontal_1_2", "inference_rule": "spm", "raw_label": "cnf(c_0_218, negated_conjecture, (horizontal_5_5|horizontal_3_5|vertical_1_4|~horizontal_1_2),\\ninference(spm,[status(thm)],[c_0_205, c_0_138])).\n", "fillcolor": "lightskyblue1"}, {"id": 219, "label": "c_0_219", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_219", "role": "negated_conjecture", "formula": "vertical_1_2|horizontal_1_2", "inference_rule": "sr", "raw_label": "cnf(c_0_219, negated_conjecture, (vertical_1_2|horizontal_1_2),\\ninference(sr,[status(thm)],[c_0_206, c_0_207])).\n", "fillcolor": "lightskyblue1"}, {"id": 220, "label": "c_0_220", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_220", "role": "negated_conjecture", "formula": "vertical_1_4|horizontal_1_3|~horizontal_2_5", "inference_rule": "spm", "raw_label": "cnf(c_0_220, negated_conjecture, (vertical_1_4|horizontal_1_3|~horizontal_2_5),\\ninference(spm,[status(thm)],[c_0_208, c_0_176])).\n", "fillcolor": "lightskyblue1"}, {"id": 221, "label": "c_0_221", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_221", "role": "negated_conjecture", "formula": "vertical_1_4|horizontal_1_3|~vertical_1_5", "inference_rule": "spm", "raw_label": "cnf(c_0_221, negated_conjecture, (vertical_1_4|horizontal_1_3|~vertical_1_5),\\ninference(spm,[status(thm)],[c_0_209, c_0_176])).\n", "fillcolor": "lightskyblue1"}, {"id": 222, "label": "c_0_222", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_222", "role": "negated_conjecture", "formula": "~vertical_2_6|~horizontal_2_5", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_222, negated_conjecture, (~vertical_2_6|~horizontal_2_5),\\ninference(fof_simplification,[status(thm)],[unique_cover4_2_6])).\n", "fillcolor": "lightskyblue1"}, {"id": 223, "label": "c_0_223", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_223", "role": "negated_conjecture", "formula": "vertical_1_6|~vertical_1_5", "inference_rule": "spm", "raw_label": "cnf(c_0_223, negated_conjecture, (vertical_1_6|~vertical_1_5),\\ninference(spm,[status(thm)],[c_0_210, c_0_131])).\n", "fillcolor": "lightskyblue1"}, {"id": 224, "label": "c_0_224", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_224", "role": "negated_conjecture", "formula": "~vertical_4_4|~horizontal_4_3", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_224, negated_conjecture, (~vertical_4_4|~horizontal_4_3),\\ninference(fof_simplification,[status(thm)],[unique_cover4_4_4])).\n", "fillcolor": "lightskyblue1"}, {"id": 225, "label": "c_0_225", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_225", "role": "negated_conjecture", "formula": "horizontal_4_4|vertical_3_5|~horizontal_5_3|~horizontal_2_2", "inference_rule": "spm", "raw_label": "cnf(c_0_225, negated_conjecture, (horizontal_4_4|vertical_3_5|~horizontal_5_3|~horizontal_2_2),\\ninference(spm,[status(thm)],[c_0_211, c_0_212])).\n", "fillcolor": "lightskyblue1"}, {"id": 226, "label": "c_0_226", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_226", "role": "negated_conjecture", "formula": "horizontal_5_3|vertical_4_4|horizontal_4_4|vertical_3_5", "inference_rule": "csr", "raw_label": "cnf(c_0_226, negated_conjecture, (horizontal_5_3|vertical_4_4|horizontal_4_4|vertical_3_5),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_180, c_0_213]), c_0_214])).\n", "fillcolor": "lightskyblue1"}, {"id": 227, "label": "c_0_227", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_227", "role": "negated_conjecture", "formula": "~horizontal_4_4|~horizontal_4_3", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_227, negated_conjecture, (~horizontal_4_4|~horizontal_4_3),\\ninference(fof_simplification,[status(thm)],[unique_cover2_4_4])).\n", "fillcolor": "lightskyblue1"}, {"id": 228, "label": "c_0_228", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_228", "role": "negated_conjecture", "formula": "~horizontal_5_2|~vertical_4_3", "inference_rule": null, "raw_label": "cnf(c_0_228, negated_conjecture, (~horizontal_5_2|~vertical_4_3),\\nc_0_215).\n", "fillcolor": "lightskyblue1"}, {"id": 229, "label": "c_0_229", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_229", "role": "negated_conjecture", "formula": "horizontal_5_2|~vertical_3_2|~horizontal_2_2", "inference_rule": "spm", "raw_label": "cnf(c_0_229, negated_conjecture, (horizontal_5_2|~vertical_3_2|~horizontal_2_2),\\ninference(spm,[status(thm)],[c_0_211, c_0_216])).\n", "fillcolor": "lightskyblue1"}, {"id": 230, "label": "c_0_230", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_230", "role": "negated_conjecture", "formula": "~horizontal_4_2|~vertical_3_2", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_230, negated_conjecture, (~horizontal_4_2|~vertical_3_2),\\ninference(fof_simplification,[status(thm)],[unique_cover3_4_2])).\n", "fillcolor": "lightskyblue1"}, {"id": 231, "label": "c_0_231", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_231", "role": "negated_conjecture", "formula": "~vertical_4_4|~vertical_3_4", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_231, negated_conjecture, (~vertical_4_4|~vertical_3_4),\\ninference(fof_simplification,[status(thm)],[unique_cover5_4_4])).\n", "fillcolor": "lightskyblue1"}, {"id": 232, "label": "c_0_232", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_232", "role": "negated_conjecture", "formula": "~horizontal_4_4|~vertical_3_4", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_232, negated_conjecture, (~horizontal_4_4|~vertical_3_4),\\ninference(fof_simplification,[status(thm)],[unique_cover3_4_4])).\n", "fillcolor": "lightskyblue1"}, {"id": 233, "label": "c_0_233", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_233", "role": "negated_conjecture", "formula": "~vertical_2_4|~vertical_1_4", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_233, negated_conjecture, (~vertical_2_4|~vertical_1_4),\\ninference(fof_simplification,[status(thm)],[unique_cover5_2_4])).\n", "fillcolor": "lightskyblue1"}, {"id": 234, "label": "c_0_234", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_234", "role": "negated_conjecture", "formula": "~horizontal_3_5|~vertical_2_5", "inference_rule": null, "raw_label": "cnf(c_0_234, negated_conjecture, (~horizontal_3_5|~vertical_2_5),\\nc_0_217).\n", "fillcolor": "lightskyblue1"}, {"id": 235, "label": "c_0_235", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_235", "role": "negated_conjecture", "formula": "horizontal_5_5|horizontal_3_5|vertical_1_4|vertical_1_2", "inference_rule": "spm", "raw_label": "cnf(c_0_235, negated_conjecture, (horizontal_5_5|horizontal_3_5|vertical_1_4|vertical_1_2),\\ninference(spm,[status(thm)],[c_0_218, c_0_219])).\n", "fillcolor": "lightskyblue1"}, {"id": 236, "label": "c_0_236", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_236", "role": "negated_conjecture", "formula": "~vertical_5_3|~horizontal_5_2", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_236, negated_conjecture, (~vertical_5_3|~horizontal_5_2),\\ninference(fof_simplification,[status(thm)],[unique_cover4_5_3])).\n", "fillcolor": "lightskyblue1"}, {"id": 237, "label": "c_0_237", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_237", "role": "negated_conjecture", "formula": "vertical_1_4|~horizontal_2_5|~horizontal_1_2", "inference_rule": "spm", "raw_label": "cnf(c_0_237, negated_conjecture, (vertical_1_4|~horizontal_2_5|~horizontal_1_2),\\ninference(spm,[status(thm)],[c_0_189, c_0_220])).\n", "fillcolor": "lightskyblue1"}, {"id": 238, "label": "c_0_238", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_238", "role": "negated_conjecture", "formula": "vertical_1_4|~vertical_1_5|~horizontal_1_2", "inference_rule": "spm", "raw_label": "cnf(c_0_238, negated_conjecture, (vertical_1_4|~vertical_1_5|~horizontal_1_2),\\ninference(spm,[status(thm)],[c_0_189, c_0_221])).\n", "fillcolor": "lightskyblue1"}, {"id": 239, "label": "c_0_239", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_239", "role": "negated_conjecture", "formula": "~vertical_2_6|~horizontal_2_5", "inference_rule": null, "raw_label": "cnf(c_0_239, negated_conjecture, (~vertical_2_6|~horizontal_2_5),\\nc_0_222).\n", "fillcolor": "lightskyblue1"}, {"id": 240, "label": "c_0_240", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_240", "role": "negated_conjecture", "formula": "~vertical_2_6|~vertical_1_5", "inference_rule": "spm", "raw_label": "cnf(c_0_240, negated_conjecture, (~vertical_2_6|~vertical_1_5),\\ninference(spm,[status(thm)],[c_0_144, c_0_223])).\n", "fillcolor": "lightskyblue1"}, {"id": 241, "label": "c_0_241", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_241", "role": "negated_conjecture", "formula": "~vertical_4_4|~horizontal_4_3", "inference_rule": null, "raw_label": "cnf(c_0_241, negated_conjecture, (~vertical_4_4|~horizontal_4_3),\\nc_0_224).\n", "fillcolor": "lightskyblue1"}, {"id": 242, "label": "c_0_242", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_242", "role": "negated_conjecture", "formula": "vertical_4_4|horizontal_4_4|vertical_3_5|~horizontal_2_2", "inference_rule": "spm", "raw_label": "cnf(c_0_242, negated_conjecture, (vertical_4_4|horizontal_4_4|vertical_3_5|~horizontal_2_2),\\ninference(spm,[status(thm)],[c_0_225, c_0_226])).\n", "fillcolor": "lightskyblue1"}, {"id": 243, "label": "c_0_243", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_243", "role": "negated_conjecture", "formula": "~horizontal_4_4|~horizontal_4_3", "inference_rule": null, "raw_label": "cnf(c_0_243, negated_conjecture, (~horizontal_4_4|~horizontal_4_3),\\nc_0_227).\n", "fillcolor": "lightskyblue1"}, {"id": 244, "label": "c_0_244", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_244", "role": "negated_conjecture", "formula": "~vertical_4_3|~vertical_3_2|~horizontal_2_2", "inference_rule": "spm", "raw_label": "cnf(c_0_244, negated_conjecture, (~vertical_4_3|~vertical_3_2|~horizontal_2_2),\\ninference(spm,[status(thm)],[c_0_228, c_0_229])).\n", "fillcolor": "lightskyblue1"}, {"id": 245, "label": "c_0_245", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_245", "role": "negated_conjecture", "formula": "horizontal_4_3|vertical_4_3|horizontal_4_2|vertical_3_3", "inference_rule": null, "raw_label": "cnf(c_0_245, negated_conjecture, (horizontal_4_3|vertical_4_3|horizontal_4_2|vertical_3_3),\\ncovered_4_3).\n", "fillcolor": "lightskyblue1"}, {"id": 246, "label": "c_0_246", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_246", "role": "negated_conjecture", "formula": "~horizontal_4_2|~vertical_3_2", "inference_rule": null, "raw_label": "cnf(c_0_246, negated_conjecture, (~horizontal_4_2|~vertical_3_2),\\nc_0_230).\n", "fillcolor": "lightskyblue1"}, {"id": 247, "label": "c_0_247", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_247", "role": "negated_conjecture", "formula": "~vertical_4_4|~vertical_3_4", "inference_rule": null, "raw_label": "cnf(c_0_247, negated_conjecture, (~vertical_4_4|~vertical_3_4),\\nc_0_231).\n", "fillcolor": "lightskyblue1"}, {"id": 248, "label": "c_0_248", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_248", "role": "negated_conjecture", "formula": "vertical_4_4|horizontal_4_4|vertical_4_2|vertical_3_5", "inference_rule": "spm", "raw_label": "cnf(c_0_248, negated_conjecture, (vertical_4_4|horizontal_4_4|vertical_4_2|vertical_3_5),\\ninference(spm,[status(thm)],[c_0_197, c_0_226])).\n", "fillcolor": "lightskyblue1"}, {"id": 249, "label": "c_0_249", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_249", "role": "negated_conjecture", "formula": "~horizontal_4_4|~vertical_3_4", "inference_rule": null, "raw_label": "cnf(c_0_249, negated_conjecture, (~horizontal_4_4|~vertical_3_4),\\nc_0_232).\n", "fillcolor": "lightskyblue1"}, {"id": 250, "label": "c_0_250", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_250", "role": "negated_conjecture", "formula": "~vertical_2_4|~vertical_1_4", "inference_rule": null, "raw_label": "cnf(c_0_250, negated_conjecture, (~vertical_2_4|~vertical_1_4),\\nc_0_233).\n", "fillcolor": "lightskyblue1"}, {"id": 251, "label": "c_0_251", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_251", "role": "negated_conjecture", "formula": "horizontal_5_5|vertical_1_4|vertical_1_2|~vertical_2_5", "inference_rule": "spm", "raw_label": "cnf(c_0_251, negated_conjecture, (horizontal_5_5|vertical_1_4|vertical_1_2|~vertical_2_5),\\ninference(spm,[status(thm)],[c_0_234, c_0_235])).\n", "fillcolor": "lightskyblue1"}, {"id": 252, "label": "c_0_252", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_252", "role": "negated_conjecture", "formula": "~vertical_5_3|~horizontal_5_2", "inference_rule": null, "raw_label": "cnf(c_0_252, negated_conjecture, (~vertical_5_3|~horizontal_5_2),\\nc_0_236).\n", "fillcolor": "lightskyblue1"}, {"id": 253, "label": "c_0_253", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_253", "role": "negated_conjecture", "formula": "vertical_1_4|vertical_1_2|~horizontal_2_5", "inference_rule": "spm", "raw_label": "cnf(c_0_253, negated_conjecture, (vertical_1_4|vertical_1_2|~horizontal_2_5),\\ninference(spm,[status(thm)],[c_0_237, c_0_219])).\n", "fillcolor": "lightskyblue1"}, {"id": 254, "label": "c_0_254", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_254", "role": "negated_conjecture", "formula": "horizontal_2_5|vertical_2_5|horizontal_2_4|vertical_1_5", "inference_rule": null, "raw_label": "cnf(c_0_254, negated_conjecture, (horizontal_2_5|vertical_2_5|horizontal_2_4|vertical_1_5),\\ncovered_2_5).\n", "fillcolor": "lightskyblue1"}, {"id": 255, "label": "c_0_255", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_255", "role": "negated_conjecture", "formula": "vertical_1_4|vertical_1_2|~vertical_1_5", "inference_rule": "spm", "raw_label": "cnf(c_0_255, negated_conjecture, (vertical_1_4|vertical_1_2|~vertical_1_5),\\ninference(spm,[status(thm)],[c_0_238, c_0_219])).\n", "fillcolor": "lightskyblue1"}, {"id": 256, "label": "c_0_256", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_256", "role": "negated_conjecture", "formula": "~horizontal_2_4|~vertical_2_4", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_256, negated_conjecture, (~horizontal_2_4|~vertical_2_4),\\ninference(fof_simplification,[status(thm)],[unique_cover1_2_4])).\n", "fillcolor": "lightskyblue1"}, {"id": 257, "label": "c_0_257", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_257", "role": "negated_conjecture", "formula": "~horizontal_3_5|~vertical_3_5", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_257, negated_conjecture, (~horizontal_3_5|~vertical_3_5),\\ninference(fof_simplification,[status(thm)],[unique_cover1_3_5])).\n", "fillcolor": "lightskyblue1"}, {"id": 258, "label": "c_0_258", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_258", "role": "negated_conjecture", "formula": "horizontal_3_5|~vertical_4_6|~horizontal_2_5", "inference_rule": "spm", "raw_label": "cnf(c_0_258, negated_conjecture, (horizontal_3_5|~vertical_4_6|~horizontal_2_5),\\ninference(spm,[status(thm)],[c_0_239, c_0_161])).\n", "fillcolor": "lightskyblue1"}, {"id": 259, "label": "c_0_259", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_259", "role": "negated_conjecture", "formula": "horizontal_3_5|~vertical_4_6|~vertical_1_5", "inference_rule": "spm", "raw_label": "cnf(c_0_259, negated_conjecture, (horizontal_3_5|~vertical_4_6|~vertical_1_5),\\ninference(spm,[status(thm)],[c_0_240, c_0_161])).\n", "fillcolor": "lightskyblue1"}, {"id": 260, "label": "c_0_260", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_260", "role": "negated_conjecture", "formula": "~vertical_3_5|~vertical_2_5", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_260, negated_conjecture, (~vertical_3_5|~vertical_2_5),\\ninference(fof_simplification,[status(thm)],[unique_cover5_3_5])).\n", "fillcolor": "lightskyblue1"}, {"id": 261, "label": "c_0_261", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_261", "role": "negated_conjecture", "formula": "vertical_3_5|~horizontal_4_3|~horizontal_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_261, negated_conjecture, (vertical_3_5|~horizontal_4_3|~horizontal_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_241, c_0_242]), c_0_243])).\n", "fillcolor": "lightskyblue1"}, {"id": 262, "label": "c_0_262", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_262", "role": "negated_conjecture", "formula": "horizontal_4_3|vertical_3_3|~vertical_3_2|~horizontal_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_262, negated_conjecture, (horizontal_4_3|vertical_3_3|~vertical_3_2|~horizontal_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_244, c_0_245]), c_0_246])).\n", "fillcolor": "lightskyblue1"}, {"id": 263, "label": "c_0_263", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_263", "role": "negated_conjecture", "formula": "vertical_4_2|vertical_3_5|~vertical_3_4", "inference_rule": "csr", "raw_label": "cnf(c_0_263, negated_conjecture, (vertical_4_2|vertical_3_5|~vertical_3_4),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_247, c_0_248]), c_0_249])).\n", "fillcolor": "lightskyblue1"}, {"id": 264, "label": "c_0_264", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_264", "role": "negated_conjecture", "formula": "horizontal_5_5|vertical_1_2|~vertical_2_5|~vertical_2_4", "inference_rule": "spm", "raw_label": "cnf(c_0_264, negated_conjecture, (horizontal_5_5|vertical_1_2|~vertical_2_5|~vertical_2_4),\\ninference(spm,[status(thm)],[c_0_250, c_0_251])).\n", "fillcolor": "lightskyblue1"}, {"id": 265, "label": "c_0_265", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_265", "role": "negated_conjecture", "formula": "vertical_5_5|horizontal_4_1|horizontal_2_1|~horizontal_5_2", "inference_rule": "spm", "raw_label": "cnf(c_0_265, negated_conjecture, (vertical_5_5|horizontal_4_1|horizontal_2_1|~horizontal_5_2),\\ninference(spm,[status(thm)],[c_0_252, c_0_163])).\n", "fillcolor": "lightskyblue1"}, {"id": 266, "label": "c_0_266", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_266", "role": "negated_conjecture", "formula": "vertical_2_5|horizontal_2_4|vertical_1_4|vertical_1_2", "inference_rule": "csr", "raw_label": "cnf(c_0_266, negated_conjecture, (vertical_2_5|horizontal_2_4|vertical_1_4|vertical_1_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_253, c_0_254]), c_0_255])).\n", "fillcolor": "lightskyblue1"}, {"id": 267, "label": "c_0_267", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_267", "role": "negated_conjecture", "formula": "~horizontal_2_4|~vertical_2_4", "inference_rule": null, "raw_label": "cnf(c_0_267, negated_conjecture, (~horizontal_2_4|~vertical_2_4),\\nc_0_256).\n", "fillcolor": "lightskyblue1"}, {"id": 268, "label": "c_0_268", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_268", "role": "negated_conjecture", "formula": "~horizontal_3_1|~vertical_2_1", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_268, negated_conjecture, (~horizontal_3_1|~vertical_2_1),\\ninference(fof_simplification,[status(thm)],[unique_cover3_3_1])).\n", "fillcolor": "lightskyblue1"}, {"id": 269, "label": "c_0_269", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_269", "role": "negated_conjecture", "formula": "~horizontal_3_5|~horizontal_3_4", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_269, negated_conjecture, (~horizontal_3_5|~horizontal_3_4),\\ninference(fof_simplification,[status(thm)],[unique_cover2_3_5])).\n", "fillcolor": "lightskyblue1"}, {"id": 270, "label": "c_0_270", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_270", "role": "negated_conjecture", "formula": "~horizontal_3_5|~vertical_3_5", "inference_rule": null, "raw_label": "cnf(c_0_270, negated_conjecture, (~horizontal_3_5|~vertical_3_5),\\nc_0_257).\n", "fillcolor": "lightskyblue1"}, {"id": 271, "label": "c_0_271", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_271", "role": "negated_conjecture", "formula": "horizontal_5_5|horizontal_3_5|~horizontal_2_5", "inference_rule": "spm", "raw_label": "cnf(c_0_271, negated_conjecture, (horizontal_5_5|horizontal_3_5|~horizontal_2_5),\\ninference(spm,[status(thm)],[c_0_258, c_0_138])).\n", "fillcolor": "lightskyblue1"}, {"id": 272, "label": "c_0_272", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_272", "role": "negated_conjecture", "formula": "horizontal_5_5|horizontal_3_5|~vertical_1_5", "inference_rule": "spm", "raw_label": "cnf(c_0_272, negated_conjecture, (horizontal_5_5|horizontal_3_5|~vertical_1_5),\\ninference(spm,[status(thm)],[c_0_259, c_0_138])).\n", "fillcolor": "lightskyblue1"}, {"id": 273, "label": "c_0_273", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_273", "role": "negated_conjecture", "formula": "~horizontal_3_3|~vertical_3_3", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_273, negated_conjecture, (~horizontal_3_3|~vertical_3_3),\\ninference(fof_simplification,[status(thm)],[unique_cover1_3_3])).\n", "fillcolor": "lightskyblue1"}, {"id": 274, "label": "c_0_274", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_274", "role": "negated_conjecture", "formula": "~vertical_3_5|~vertical_2_5", "inference_rule": null, "raw_label": "cnf(c_0_274, negated_conjecture, (~vertical_3_5|~vertical_2_5),\\nc_0_260).\n", "fillcolor": "lightskyblue1"}, {"id": 275, "label": "c_0_275", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_275", "role": "negated_conjecture", "formula": "vertical_3_5|vertical_3_3|~vertical_3_2|~horizontal_2_2", "inference_rule": "spm", "raw_label": "cnf(c_0_275, negated_conjecture, (vertical_3_5|vertical_3_3|~vertical_3_2|~horizontal_2_2),\\ninference(spm,[status(thm)],[c_0_261, c_0_262])).\n", "fillcolor": "lightskyblue1"}, {"id": 276, "label": "c_0_276", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_276", "role": "negated_conjecture", "formula": "vertical_3_5|~vertical_3_4|~vertical_3_2", "inference_rule": "spm", "raw_label": "cnf(c_0_276, negated_conjecture, (vertical_3_5|~vertical_3_4|~vertical_3_2),\\ninference(spm,[status(thm)],[c_0_204, c_0_263])).\n", "fillcolor": "lightskyblue1"}, {"id": 277, "label": "c_0_277", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_277", "role": "negated_conjecture", "formula": "~horizontal_3_4|~vertical_2_5", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_277, negated_conjecture, (~horizontal_3_4|~vertical_2_5),\\ninference(fof_simplification,[status(thm)],[unique_cover6_3_5])).\n", "fillcolor": "lightskyblue1"}, {"id": 278, "label": "c_0_278", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_278", "role": "negated_conjecture", "formula": "vertical_1_2|~vertical_5_5|~vertical_2_5|~vertical_2_4", "inference_rule": "spm", "raw_label": "cnf(c_0_278, negated_conjecture, (vertical_1_2|~vertical_5_5|~vertical_2_5|~vertical_2_4),\\ninference(spm,[status(thm)],[c_0_164, c_0_264])).\n", "fillcolor": "lightskyblue1"}, {"id": 279, "label": "c_0_279", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_279", "role": "negated_conjecture", "formula": "vertical_5_5|horizontal_2_1|~horizontal_5_2|~vertical_3_2", "inference_rule": "spm", "raw_label": "cnf(c_0_279, negated_conjecture, (vertical_5_5|horizontal_2_1|~horizontal_5_2|~vertical_3_2),\\ninference(spm,[status(thm)],[c_0_202, c_0_265])).\n", "fillcolor": "lightskyblue1"}, {"id": 280, "label": "c_0_280", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_280", "role": "negated_conjecture", "formula": "vertical_2_5|vertical_1_2|~vertical_2_4", "inference_rule": "csr", "raw_label": "cnf(c_0_280, negated_conjecture, (vertical_2_5|vertical_1_2|~vertical_2_4),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_250, c_0_266]), c_0_267])).\n", "fillcolor": "lightskyblue1"}, {"id": 281, "label": "c_0_281", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_281", "role": "negated_conjecture", "formula": "~horizontal_2_2|~vertical_1_2", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_281, negated_conjecture, (~horizontal_2_2|~vertical_1_2),\\ninference(fof_simplification,[status(thm)],[unique_cover3_2_2])).\n", "fillcolor": "lightskyblue1"}, {"id": 282, "label": "c_0_282", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_282", "role": "negated_conjecture", "formula": "~horizontal_3_1|~vertical_2_1", "inference_rule": null, "raw_label": "cnf(c_0_282, negated_conjecture, (~horizontal_3_1|~vertical_2_1),\\nc_0_268).\n", "fillcolor": "lightskyblue1"}, {"id": 283, "label": "c_0_283", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_283", "role": "negated_conjecture", "formula": "~vertical_5_3|~vertical_4_3", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_283, negated_conjecture, (~vertical_5_3|~vertical_4_3),\\ninference(fof_simplification,[status(thm)],[unique_cover5_5_3])).\n", "fillcolor": "lightskyblue1"}, {"id": 284, "label": "c_0_284", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_284", "role": "negated_conjecture", "formula": "~horizontal_3_5|~horizontal_3_4", "inference_rule": null, "raw_label": "cnf(c_0_284, negated_conjecture, (~horizontal_3_5|~horizontal_3_4),\\nc_0_269).\n", "fillcolor": "lightskyblue1"}, {"id": 285, "label": "c_0_285", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_285", "role": "negated_conjecture", "formula": "horizontal_5_5|~vertical_3_5|~horizontal_2_5", "inference_rule": "spm", "raw_label": "cnf(c_0_285, negated_conjecture, (horizontal_5_5|~vertical_3_5|~horizontal_2_5),\\ninference(spm,[status(thm)],[c_0_270, c_0_271])).\n", "fillcolor": "lightskyblue1"}, {"id": 286, "label": "c_0_286", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_286", "role": "negated_conjecture", "formula": "horizontal_5_5|~vertical_3_5|~vertical_1_5", "inference_rule": "spm", "raw_label": "cnf(c_0_286, negated_conjecture, (horizontal_5_5|~vertical_3_5|~vertical_1_5),\\ninference(spm,[status(thm)],[c_0_270, c_0_272])).\n", "fillcolor": "lightskyblue1"}, {"id": 287, "label": "c_0_287", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_287", "role": "negated_conjecture", "formula": "~horizontal_2_4|~vertical_1_4", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_287, negated_conjecture, (~horizontal_2_4|~vertical_1_4),\\ninference(fof_simplification,[status(thm)],[unique_cover3_2_4])).\n", "fillcolor": "lightskyblue1"}, {"id": 288, "label": "c_0_288", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_288", "role": "negated_conjecture", "formula": "~horizontal_3_3|~vertical_3_3", "inference_rule": null, "raw_label": "cnf(c_0_288, negated_conjecture, (~horizontal_3_3|~vertical_3_3),\\nc_0_273).\n", "fillcolor": "lightskyblue1"}, {"id": 289, "label": "c_0_289", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_289", "role": "negated_conjecture", "formula": "vertical_3_3|~vertical_3_2|~vertical_2_5|~horizontal_2_2", "inference_rule": "spm", "raw_label": "cnf(c_0_289, negated_conjecture, (vertical_3_3|~vertical_3_2|~vertical_2_5|~horizontal_2_2),\\ninference(spm,[status(thm)],[c_0_274, c_0_275])).\n", "fillcolor": "lightskyblue1"}, {"id": 290, "label": "c_0_290", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_290", "role": "negated_conjecture", "formula": "~vertical_3_4|~vertical_3_2|~vertical_2_5", "inference_rule": "spm", "raw_label": "cnf(c_0_290, negated_conjecture, (~vertical_3_4|~vertical_3_2|~vertical_2_5),\\ninference(spm,[status(thm)],[c_0_274, c_0_276])).\n", "fillcolor": "lightskyblue1"}, {"id": 291, "label": "c_0_291", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_291", "role": "negated_conjecture", "formula": "horizontal_3_4|vertical_3_4|horizontal_3_3|vertical_2_4", "inference_rule": null, "raw_label": "cnf(c_0_291, negated_conjecture, (horizontal_3_4|vertical_3_4|horizontal_3_3|vertical_2_4),\\ncovered_3_4).\n", "fillcolor": "lightskyblue1"}, {"id": 292, "label": "c_0_292", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_292", "role": "negated_conjecture", "formula": "~horizontal_3_4|~vertical_2_5", "inference_rule": null, "raw_label": "cnf(c_0_292, negated_conjecture, (~horizontal_3_4|~vertical_2_5),\\nc_0_277).\n", "fillcolor": "lightskyblue1"}, {"id": 293, "label": "c_0_293", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_293", "role": "negated_conjecture", "formula": "horizontal_2_1|vertical_1_2|~vertical_3_2|~vertical_2_4", "inference_rule": "csr", "raw_label": "cnf(c_0_293, negated_conjecture, (horizontal_2_1|vertical_1_2|~vertical_3_2|~vertical_2_4),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_278, c_0_279]), c_0_280]), c_0_216])).\n", "fillcolor": "lightskyblue1"}, {"id": 294, "label": "c_0_294", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_294", "role": "negated_conjecture", "formula": "~horizontal_2_2|~vertical_1_2", "inference_rule": null, "raw_label": "cnf(c_0_294, negated_conjecture, (~horizontal_2_2|~vertical_1_2),\\nc_0_281).\n", "fillcolor": "lightskyblue1"}, {"id": 295, "label": "c_0_295", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_295", "role": "negated_conjecture", "formula": "horizontal_2_1|~horizontal_3_1", "inference_rule": "spm", "raw_label": "cnf(c_0_295, negated_conjecture, (horizontal_2_1|~horizontal_3_1),\\ninference(spm,[status(thm)],[c_0_282, c_0_149])).\n", "fillcolor": "lightskyblue1"}, {"id": 296, "label": "c_0_296", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_296", "role": "negated_conjecture", "formula": "horizontal_3_2|vertical_3_2|horizontal_3_1|vertical_2_2", "inference_rule": null, "raw_label": "cnf(c_0_296, negated_conjecture, (horizontal_3_2|vertical_3_2|horizontal_3_1|vertical_2_2),\\ncovered_3_2).\n", "fillcolor": "lightskyblue1"}, {"id": 297, "label": "c_0_297", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_297", "role": "negated_conjecture", "formula": "~horizontal_2_2|~vertical_2_2", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_297, negated_conjecture, (~horizontal_2_2|~vertical_2_2),\\ninference(fof_simplification,[status(thm)],[unique_cover1_2_2])).\n", "fillcolor": "lightskyblue1"}, {"id": 298, "label": "c_0_298", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_298", "role": "negated_conjecture", "formula": "~vertical_5_3|~vertical_4_3", "inference_rule": null, "raw_label": "cnf(c_0_298, negated_conjecture, (~vertical_5_3|~vertical_4_3),\\nc_0_283).\n", "fillcolor": "lightskyblue1"}, {"id": 299, "label": "c_0_299", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_299", "role": "negated_conjecture", "formula": "horizontal_5_5|~horizontal_3_4|~horizontal_2_5", "inference_rule": "spm", "raw_label": "cnf(c_0_299, negated_conjecture, (horizontal_5_5|~horizontal_3_4|~horizontal_2_5),\\ninference(spm,[status(thm)],[c_0_284, c_0_271])).\n", "fillcolor": "lightskyblue1"}, {"id": 300, "label": "c_0_300", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_300", "role": "negated_conjecture", "formula": "horizontal_5_5|~horizontal_3_4|~vertical_1_5", "inference_rule": "spm", "raw_label": "cnf(c_0_300, negated_conjecture, (horizontal_5_5|~horizontal_3_4|~vertical_1_5),\\ninference(spm,[status(thm)],[c_0_284, c_0_272])).\n", "fillcolor": "lightskyblue1"}, {"id": 301, "label": "c_0_301", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_301", "role": "negated_conjecture", "formula": "horizontal_5_5|horizontal_2_4|~vertical_3_5", "inference_rule": "csr", "raw_label": "cnf(c_0_301, negated_conjecture, (horizontal_5_5|horizontal_2_4|~vertical_3_5),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_285, c_0_254]), c_0_286]), c_0_274])).\n", "fillcolor": "lightskyblue1"}, {"id": 302, "label": "c_0_302", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_302", "role": "negated_conjecture", "formula": "~horizontal_2_4|~vertical_1_4", "inference_rule": null, "raw_label": "cnf(c_0_302, negated_conjecture, (~horizontal_2_4|~vertical_1_4),\\nc_0_287).\n", "fillcolor": "lightskyblue1"}, {"id": 303, "label": "c_0_303", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_303", "role": "negated_conjecture", "formula": "horizontal_5_5|vertical_1_4|vertical_1_2|~vertical_3_5", "inference_rule": "spm", "raw_label": "cnf(c_0_303, negated_conjecture, (horizontal_5_5|vertical_1_4|vertical_1_2|~vertical_3_5),\\ninference(spm,[status(thm)],[c_0_270, c_0_235])).\n", "fillcolor": "lightskyblue1"}, {"id": 304, "label": "c_0_304", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_304", "role": "negated_conjecture", "formula": "vertical_4_2|vertical_1_2|~horizontal_5_3|~vertical_2_4", "inference_rule": "csr", "raw_label": "cnf(c_0_304, negated_conjecture, (vertical_4_2|vertical_1_2|~horizontal_5_3|~vertical_2_4),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_278, c_0_181]), c_0_280])).\n", "fillcolor": "lightskyblue1"}, {"id": 305, "label": "c_0_305", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_305", "role": "negated_conjecture", "formula": "vertical_3_5|~vertical_3_4|~horizontal_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_305, negated_conjecture, (vertical_3_5|~vertical_3_4|~horizontal_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_247, c_0_242]), c_0_249])).\n", "fillcolor": "lightskyblue1"}, {"id": 306, "label": "c_0_306", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_306", "role": "negated_conjecture", "formula": "~horizontal_3_3|~vertical_3_2|~vertical_2_5|~horizontal_2_2", "inference_rule": "spm", "raw_label": "cnf(c_0_306, negated_conjecture, (~horizontal_3_3|~vertical_3_2|~vertical_2_5|~horizontal_2_2),\\ninference(spm,[status(thm)],[c_0_288, c_0_289])).\n", "fillcolor": "lightskyblue1"}, {"id": 307, "label": "c_0_307", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_307", "role": "negated_conjecture", "formula": "horizontal_3_3|vertical_2_4|~vertical_3_2|~vertical_2_5", "inference_rule": "csr", "raw_label": "cnf(c_0_307, negated_conjecture, (horizontal_3_3|vertical_2_4|~vertical_3_2|~vertical_2_5),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_290, c_0_291]), c_0_292])).\n", "fillcolor": "lightskyblue1"}, {"id": 308, "label": "c_0_308", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_308", "role": "negated_conjecture", "formula": "~vertical_3_2|~vertical_2_4|~horizontal_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_308, negated_conjecture, (~vertical_3_2|~vertical_2_4|~horizontal_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_211, c_0_293]), c_0_294])).\n", "fillcolor": "lightskyblue1"}, {"id": 309, "label": "c_0_309", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_309", "role": "negated_conjecture", "formula": "vertical_3_2|horizontal_3_2|vertical_2_2|horizontal_2_1", "inference_rule": "spm", "raw_label": "cnf(c_0_309, negated_conjecture, (vertical_3_2|horizontal_3_2|vertical_2_2|horizontal_2_1),\\ninference(spm,[status(thm)],[c_0_295, c_0_296])).\n", "fillcolor": "lightskyblue1"}, {"id": 310, "label": "c_0_310", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_310", "role": "negated_conjecture", "formula": "~horizontal_2_2|~vertical_2_2", "inference_rule": null, "raw_label": "cnf(c_0_310, negated_conjecture, (~horizontal_2_2|~vertical_2_2),\\nc_0_297).\n", "fillcolor": "lightskyblue1"}, {"id": 311, "label": "c_0_311", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_311", "role": "negated_conjecture", "formula": "vertical_5_5|~vertical_5_2|~vertical_4_3", "inference_rule": "spm", "raw_label": "cnf(c_0_311, negated_conjecture, (vertical_5_5|~vertical_5_2|~vertical_4_3),\\ninference(spm,[status(thm)],[c_0_298, c_0_155])).\n", "fillcolor": "lightskyblue1"}, {"id": 312, "label": "c_0_312", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_312", "role": "negated_conjecture", "formula": "vertical_5_5|~horizontal_5_1|~vertical_4_3", "inference_rule": "spm", "raw_label": "cnf(c_0_312, negated_conjecture, (vertical_5_5|~horizontal_5_1|~vertical_4_3),\\ninference(spm,[status(thm)],[c_0_298, c_0_156])).\n", "fillcolor": "lightskyblue1"}, {"id": 313, "label": "c_0_313", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_313", "role": "negated_conjecture", "formula": "~horizontal_4_2|~horizontal_4_1", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_313, negated_conjecture, (~horizontal_4_2|~horizontal_4_1),\\ninference(fof_simplification,[status(thm)],[unique_cover2_4_2])).\n", "fillcolor": "lightskyblue1"}, {"id": 314, "label": "c_0_314", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_314", "role": "negated_conjecture", "formula": "~horizontal_4_2|~vertical_4_2", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_314, negated_conjecture, (~horizontal_4_2|~vertical_4_2),\\ninference(fof_simplification,[status(thm)],[unique_cover1_4_2])).\n", "fillcolor": "lightskyblue1"}, {"id": 315, "label": "c_0_315", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_315", "role": "negated_conjecture", "formula": "horizontal_5_5|horizontal_2_4|~horizontal_3_4", "inference_rule": "csr", "raw_label": "cnf(c_0_315, negated_conjecture, (horizontal_5_5|horizontal_2_4|~horizontal_3_4),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_299, c_0_254]), c_0_300]), c_0_292])).\n", "fillcolor": "lightskyblue1"}, {"id": 316, "label": "c_0_316", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_316", "role": "negated_conjecture", "formula": "horizontal_5_5|vertical_1_4|vertical_1_2|~horizontal_3_4", "inference_rule": "spm", "raw_label": "cnf(c_0_316, negated_conjecture, (horizontal_5_5|vertical_1_4|vertical_1_2|~horizontal_3_4),\\ninference(spm,[status(thm)],[c_0_284, c_0_235])).\n", "fillcolor": "lightskyblue1"}, {"id": 317, "label": "c_0_317", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_317", "role": "negated_conjecture", "formula": "horizontal_2_4|~vertical_5_5|~vertical_3_5", "inference_rule": "spm", "raw_label": "cnf(c_0_317, negated_conjecture, (horizontal_2_4|~vertical_5_5|~vertical_3_5),\\ninference(spm,[status(thm)],[c_0_164, c_0_301])).\n", "fillcolor": "lightskyblue1"}, {"id": 318, "label": "c_0_318", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_318", "role": "negated_conjecture", "formula": "horizontal_5_5|vertical_1_2|~vertical_3_5|~horizontal_2_4", "inference_rule": "spm", "raw_label": "cnf(c_0_318, negated_conjecture, (horizontal_5_5|vertical_1_2|~vertical_3_5|~horizontal_2_4),\\ninference(spm,[status(thm)],[c_0_302, c_0_303])).\n", "fillcolor": "lightskyblue1"}, {"id": 319, "label": "c_0_319", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_319", "role": "negated_conjecture", "formula": "horizontal_2_1|vertical_1_2|~horizontal_5_3|~vertical_2_4", "inference_rule": "csr", "raw_label": "cnf(c_0_319, negated_conjecture, (horizontal_2_1|vertical_1_2|~horizontal_5_3|~vertical_2_4),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_278, c_0_196]), c_0_280]), c_0_304])).\n", "fillcolor": "lightskyblue1"}, {"id": 320, "label": "c_0_320", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_320", "role": "negated_conjecture", "formula": "~horizontal_3_3|~horizontal_3_2", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_320, negated_conjecture, (~horizontal_3_3|~horizontal_3_2),\\ninference(fof_simplification,[status(thm)],[unique_cover2_3_3])).\n", "fillcolor": "lightskyblue1"}, {"id": 321, "label": "c_0_321", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_321", "role": "negated_conjecture", "formula": "~vertical_3_4|~vertical_2_5|~horizontal_2_2", "inference_rule": "spm", "raw_label": "cnf(c_0_321, negated_conjecture, (~vertical_3_4|~vertical_2_5|~horizontal_2_2),\\ninference(spm,[status(thm)],[c_0_274, c_0_305])).\n", "fillcolor": "lightskyblue1"}, {"id": 322, "label": "c_0_322", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_322", "role": "negated_conjecture", "formula": "~vertical_3_2|~vertical_2_5|~horizontal_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_322, negated_conjecture, (~vertical_3_2|~vertical_2_5|~horizontal_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_306, c_0_307]), c_0_308])).\n", "fillcolor": "lightskyblue1"}, {"id": 323, "label": "c_0_323", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_323", "role": "negated_conjecture", "formula": "vertical_3_2|horizontal_3_2|~horizontal_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_323, negated_conjecture, (vertical_3_2|horizontal_3_2|~horizontal_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_211, c_0_309]), c_0_310])).\n", "fillcolor": "lightskyblue1"}, {"id": 324, "label": "c_0_324", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_324", "role": "negated_conjecture", "formula": "vertical_5_5|horizontal_4_1|horizontal_2_1|~vertical_4_3", "inference_rule": "spm", "raw_label": "cnf(c_0_324, negated_conjecture, (vertical_5_5|horizontal_4_1|horizontal_2_1|~vertical_4_3),\\ninference(spm,[status(thm)],[c_0_298, c_0_163])).\n", "fillcolor": "lightskyblue1"}, {"id": 325, "label": "c_0_325", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_325", "role": "negated_conjecture", "formula": "vertical_5_5|vertical_4_2|~vertical_4_3", "inference_rule": "csr", "raw_label": "cnf(c_0_325, negated_conjecture, (vertical_5_5|vertical_4_2|~vertical_4_3),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_311, c_0_168]), c_0_312]), c_0_228])).\n", "fillcolor": "lightskyblue1"}, {"id": 326, "label": "c_0_326", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_326", "role": "negated_conjecture", "formula": "~horizontal_4_2|~horizontal_4_1", "inference_rule": null, "raw_label": "cnf(c_0_326, negated_conjecture, (~horizontal_4_2|~horizontal_4_1),\\nc_0_313).\n", "fillcolor": "lightskyblue1"}, {"id": 327, "label": "c_0_327", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_327", "role": "negated_conjecture", "formula": "~horizontal_4_2|~vertical_4_2", "inference_rule": null, "raw_label": "cnf(c_0_327, negated_conjecture, (~horizontal_4_2|~vertical_4_2),\\nc_0_314).\n", "fillcolor": "lightskyblue1"}, {"id": 328, "label": "c_0_328", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_328", "role": "negated_conjecture", "formula": "~horizontal_2_3|~vertical_1_4", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_328, negated_conjecture, (~horizontal_2_3|~vertical_1_4),\\ninference(fof_simplification,[status(thm)],[unique_cover6_2_4])).\n", "fillcolor": "lightskyblue1"}, {"id": 329, "label": "c_0_329", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_329", "role": "negated_conjecture", "formula": "horizontal_2_4|~vertical_5_5|~horizontal_3_4", "inference_rule": "spm", "raw_label": "cnf(c_0_329, negated_conjecture, (horizontal_2_4|~vertical_5_5|~horizontal_3_4),\\ninference(spm,[status(thm)],[c_0_164, c_0_315])).\n", "fillcolor": "lightskyblue1"}, {"id": 330, "label": "c_0_330", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_330", "role": "negated_conjecture", "formula": "horizontal_5_5|vertical_1_2|~horizontal_3_4|~horizontal_2_4", "inference_rule": "spm", "raw_label": "cnf(c_0_330, negated_conjecture, (horizontal_5_5|vertical_1_2|~horizontal_3_4|~horizontal_2_4),\\ninference(spm,[status(thm)],[c_0_302, c_0_316])).\n", "fillcolor": "lightskyblue1"}, {"id": 331, "label": "c_0_331", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_331", "role": "negated_conjecture", "formula": "vertical_4_2|horizontal_2_4|~horizontal_5_3|~vertical_3_5", "inference_rule": "spm", "raw_label": "cnf(c_0_331, negated_conjecture, (vertical_4_2|horizontal_2_4|~horizontal_5_3|~vertical_3_5),\\ninference(spm,[status(thm)],[c_0_317, c_0_181])).\n", "fillcolor": "lightskyblue1"}, {"id": 332, "label": "c_0_332", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_332", "role": "negated_conjecture", "formula": "vertical_1_2|~vertical_5_5|~vertical_3_5|~horizontal_2_4", "inference_rule": "spm", "raw_label": "cnf(c_0_332, negated_conjecture, (vertical_1_2|~vertical_5_5|~vertical_3_5|~horizontal_2_4),\\ninference(spm,[status(thm)],[c_0_164, c_0_318])).\n", "fillcolor": "lightskyblue1"}, {"id": 333, "label": "c_0_333", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_333", "role": "negated_conjecture", "formula": "vertical_1_2|~horizontal_5_4|~vertical_2_5|~vertical_2_4", "inference_rule": "spm", "raw_label": "cnf(c_0_333, negated_conjecture, (vertical_1_2|~horizontal_5_4|~vertical_2_5|~vertical_2_4),\\ninference(spm,[status(thm)],[c_0_200, c_0_264])).\n", "fillcolor": "lightskyblue1"}, {"id": 334, "label": "c_0_334", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_334", "role": "negated_conjecture", "formula": "horizontal_5_4|horizontal_5_3|vertical_4_4|vertical_1_2|~vertical_2_4", "inference_rule": "csr", "raw_label": "cnf(c_0_334, negated_conjecture, (horizontal_5_4|horizontal_5_3|vertical_4_4|vertical_1_2|~vertical_2_4),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_278, c_0_213]), c_0_280])).\n", "fillcolor": "lightskyblue1"}, {"id": 335, "label": "c_0_335", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_335", "role": "negated_conjecture", "formula": "~horizontal_5_3|~vertical_2_4|~horizontal_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_335, negated_conjecture, (~horizontal_5_3|~vertical_2_4|~horizontal_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_211, c_0_319]), c_0_294])).\n", "fillcolor": "lightskyblue1"}, {"id": 336, "label": "c_0_336", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_336", "role": "negated_conjecture", "formula": "~horizontal_3_3|~horizontal_3_2", "inference_rule": null, "raw_label": "cnf(c_0_336, negated_conjecture, (~horizontal_3_3|~horizontal_3_2),\\nc_0_320).\n", "fillcolor": "lightskyblue1"}, {"id": 337, "label": "c_0_337", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_337", "role": "negated_conjecture", "formula": "horizontal_3_3|vertical_2_4|~vertical_2_5|~horizontal_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_337, negated_conjecture, (horizontal_3_3|vertical_2_4|~vertical_2_5|~horizontal_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_321, c_0_291]), c_0_292])).\n", "fillcolor": "lightskyblue1"}, {"id": 338, "label": "c_0_338", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_338", "role": "negated_conjecture", "formula": "horizontal_3_2|~vertical_2_5|~horizontal_2_2", "inference_rule": "spm", "raw_label": "cnf(c_0_338, negated_conjecture, (horizontal_3_2|~vertical_2_5|~horizontal_2_2),\\ninference(spm,[status(thm)],[c_0_322, c_0_323])).\n", "fillcolor": "lightskyblue1"}, {"id": 339, "label": "c_0_339", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_339", "role": "negated_conjecture", "formula": "vertical_5_5|horizontal_2_1|~vertical_4_3|~vertical_4_2", "inference_rule": "spm", "raw_label": "cnf(c_0_339, negated_conjecture, (vertical_5_5|horizontal_2_1|~vertical_4_3|~vertical_4_2),\\ninference(spm,[status(thm)],[c_0_178, c_0_324])).\n", "fillcolor": "lightskyblue1"}, {"id": 340, "label": "c_0_340", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_340", "role": "negated_conjecture", "formula": "vertical_4_2|vertical_1_2|~vertical_4_3|~vertical_2_4", "inference_rule": "csr", "raw_label": "cnf(c_0_340, negated_conjecture, (vertical_4_2|vertical_1_2|~vertical_4_3|~vertical_2_4),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_278, c_0_325]), c_0_280])).\n", "fillcolor": "lightskyblue1"}, {"id": 341, "label": "c_0_341", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_341", "role": "negated_conjecture", "formula": "vertical_5_5|horizontal_2_1|~horizontal_5_2|~horizontal_4_2", "inference_rule": "spm", "raw_label": "cnf(c_0_341, negated_conjecture, (vertical_5_5|horizontal_2_1|~horizontal_5_2|~horizontal_4_2),\\ninference(spm,[status(thm)],[c_0_326, c_0_265])).\n", "fillcolor": "lightskyblue1"}, {"id": 342, "label": "c_0_342", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_342", "role": "negated_conjecture", "formula": "horizontal_5_2|horizontal_2_1|~horizontal_4_2", "inference_rule": "csr", "raw_label": "cnf(c_0_342, negated_conjecture, (horizontal_5_2|horizontal_2_1|~horizontal_4_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_326, c_0_203]), c_0_327])).\n", "fillcolor": "lightskyblue1"}, {"id": 343, "label": "c_0_343", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_343", "role": "negated_conjecture", "formula": "~horizontal_2_3|~vertical_1_4", "inference_rule": null, "raw_label": "cnf(c_0_343, negated_conjecture, (~horizontal_2_3|~vertical_1_4),\\nc_0_328).\n", "fillcolor": "lightskyblue1"}, {"id": 344, "label": "c_0_344", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_344", "role": "negated_conjecture", "formula": "~horizontal_2_4|~horizontal_2_3", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_344, negated_conjecture, (~horizontal_2_4|~horizontal_2_3),\\ninference(fof_simplification,[status(thm)],[unique_cover2_2_4])).\n", "fillcolor": "lightskyblue1"}, {"id": 345, "label": "c_0_345", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_345", "role": "negated_conjecture", "formula": "vertical_4_2|horizontal_2_4|~horizontal_5_3|~horizontal_3_4", "inference_rule": "spm", "raw_label": "cnf(c_0_345, negated_conjecture, (vertical_4_2|horizontal_2_4|~horizontal_5_3|~horizontal_3_4),\\ninference(spm,[status(thm)],[c_0_329, c_0_181])).\n", "fillcolor": "lightskyblue1"}, {"id": 346, "label": "c_0_346", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_346", "role": "negated_conjecture", "formula": "vertical_1_2|~vertical_5_5|~horizontal_3_4|~horizontal_2_4", "inference_rule": "spm", "raw_label": "cnf(c_0_346, negated_conjecture, (vertical_1_2|~vertical_5_5|~horizontal_3_4|~horizontal_2_4),\\ninference(spm,[status(thm)],[c_0_164, c_0_330])).\n", "fillcolor": "lightskyblue1"}, {"id": 347, "label": "c_0_347", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_347", "role": "negated_conjecture", "formula": "horizontal_2_4|~horizontal_5_4|~vertical_3_5", "inference_rule": "spm", "raw_label": "cnf(c_0_347, negated_conjecture, (horizontal_2_4|~horizontal_5_4|~vertical_3_5),\\ninference(spm,[status(thm)],[c_0_200, c_0_301])).\n", "fillcolor": "lightskyblue1"}, {"id": 348, "label": "c_0_348", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_348", "role": "negated_conjecture", "formula": "horizontal_2_4|horizontal_2_1|~horizontal_5_3|~vertical_3_5", "inference_rule": "csr", "raw_label": "cnf(c_0_348, negated_conjecture, (horizontal_2_4|horizontal_2_1|~horizontal_5_3|~vertical_3_5),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_317, c_0_196]), c_0_331])).\n", "fillcolor": "lightskyblue1"}, {"id": 349, "label": "c_0_349", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_349", "role": "negated_conjecture", "formula": "vertical_4_2|vertical_1_2|~horizontal_5_3|~vertical_3_5", "inference_rule": "csr", "raw_label": "cnf(c_0_349, negated_conjecture, (vertical_4_2|vertical_1_2|~horizontal_5_3|~vertical_3_5),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_332, c_0_181]), c_0_331])).\n", "fillcolor": "lightskyblue1"}, {"id": 350, "label": "c_0_350", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_350", "role": "negated_conjecture", "formula": "~horizontal_5_4|~vertical_2_5|~vertical_2_4|~horizontal_2_2", "inference_rule": "spm", "raw_label": "cnf(c_0_350, negated_conjecture, (~horizontal_5_4|~vertical_2_5|~vertical_2_4|~horizontal_2_2),\\ninference(spm,[status(thm)],[c_0_294, c_0_333])).\n", "fillcolor": "lightskyblue1"}, {"id": 351, "label": "c_0_351", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_351", "role": "negated_conjecture", "formula": "horizontal_5_4|vertical_4_4|~vertical_2_4|~horizontal_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_351, negated_conjecture, (horizontal_5_4|vertical_4_4|~vertical_2_4|~horizontal_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_294, c_0_334]), c_0_335])).\n", "fillcolor": "lightskyblue1"}, {"id": 352, "label": "c_0_352", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_352", "role": "negated_conjecture", "formula": "vertical_2_4|~vertical_2_5|~horizontal_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_352, negated_conjecture, (vertical_2_4|~vertical_2_5|~horizontal_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_336, c_0_337]), c_0_338])).\n", "fillcolor": "lightskyblue1"}, {"id": 353, "label": "c_0_353", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_353", "role": "negated_conjecture", "formula": "horizontal_2_1|vertical_1_2|~vertical_4_3|~vertical_2_4", "inference_rule": "csr", "raw_label": "cnf(c_0_353, negated_conjecture, (horizontal_2_1|vertical_1_2|~vertical_4_3|~vertical_2_4),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_278, c_0_339]), c_0_280]), c_0_340])).\n", "fillcolor": "lightskyblue1"}, {"id": 354, "label": "c_0_354", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_354", "role": "negated_conjecture", "formula": "horizontal_2_1|vertical_1_2|~horizontal_4_2|~vertical_2_4", "inference_rule": "csr", "raw_label": "cnf(c_0_354, negated_conjecture, (horizontal_2_1|vertical_1_2|~horizontal_4_2|~vertical_2_4),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_278, c_0_341]), c_0_280]), c_0_342])).\n", "fillcolor": "lightskyblue1"}, {"id": 355, "label": "c_0_355", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_355", "role": "negated_conjecture", "formula": "horizontal_5_5|vertical_1_2|~vertical_2_5|~horizontal_2_3", "inference_rule": "spm", "raw_label": "cnf(c_0_355, negated_conjecture, (horizontal_5_5|vertical_1_2|~vertical_2_5|~horizontal_2_3),\\ninference(spm,[status(thm)],[c_0_343, c_0_251])).\n", "fillcolor": "lightskyblue1"}, {"id": 356, "label": "c_0_356", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_356", "role": "negated_conjecture", "formula": "~horizontal_2_4|~horizontal_2_3", "inference_rule": null, "raw_label": "cnf(c_0_356, negated_conjecture, (~horizontal_2_4|~horizontal_2_3),\\nc_0_344).\n", "fillcolor": "lightskyblue1"}, {"id": 357, "label": "c_0_357", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_357", "role": "negated_conjecture", "formula": "horizontal_2_4|~horizontal_5_4|~horizontal_3_4", "inference_rule": "spm", "raw_label": "cnf(c_0_357, negated_conjecture, (horizontal_2_4|~horizontal_5_4|~horizontal_3_4),\\ninference(spm,[status(thm)],[c_0_200, c_0_315])).\n", "fillcolor": "lightskyblue1"}, {"id": 358, "label": "c_0_358", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_358", "role": "negated_conjecture", "formula": "horizontal_2_4|horizontal_2_1|~horizontal_5_3|~horizontal_3_4", "inference_rule": "csr", "raw_label": "cnf(c_0_358, negated_conjecture, (horizontal_2_4|horizontal_2_1|~horizontal_5_3|~horizontal_3_4),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_329, c_0_196]), c_0_345])).\n", "fillcolor": "lightskyblue1"}, {"id": 359, "label": "c_0_359", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_359", "role": "negated_conjecture", "formula": "vertical_4_2|vertical_1_2|~horizontal_5_3|~horizontal_3_4", "inference_rule": "csr", "raw_label": "cnf(c_0_359, negated_conjecture, (vertical_4_2|vertical_1_2|~horizontal_5_3|~horizontal_3_4),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_346, c_0_181]), c_0_345])).\n", "fillcolor": "lightskyblue1"}, {"id": 360, "label": "c_0_360", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_360", "role": "negated_conjecture", "formula": "horizontal_5_3|vertical_4_4|horizontal_2_4|~vertical_3_5", "inference_rule": "csr", "raw_label": "cnf(c_0_360, negated_conjecture, (horizontal_5_3|vertical_4_4|horizontal_2_4|~vertical_3_5),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_317, c_0_213]), c_0_347])).\n", "fillcolor": "lightskyblue1"}, {"id": 361, "label": "c_0_361", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_361", "role": "negated_conjecture", "formula": "horizontal_2_1|vertical_1_2|~horizontal_5_3|~vertical_3_5", "inference_rule": "csr", "raw_label": "cnf(c_0_361, negated_conjecture, (horizontal_2_1|vertical_1_2|~horizontal_5_3|~vertical_3_5),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_332, c_0_196]), c_0_348]), c_0_349])).\n", "fillcolor": "lightskyblue1"}, {"id": 362, "label": "c_0_362", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_362", "role": "negated_conjecture", "formula": "vertical_4_4|~vertical_2_5|~horizontal_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_362, negated_conjecture, (vertical_4_4|~vertical_2_5|~horizontal_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_350, c_0_351]), c_0_352])).\n", "fillcolor": "lightskyblue1"}, {"id": 363, "label": "c_0_363", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_363", "role": "negated_conjecture", "formula": "~vertical_4_3|~vertical_2_4|~horizontal_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_363, negated_conjecture, (~vertical_4_3|~vertical_2_4|~horizontal_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_211, c_0_353]), c_0_294])).\n", "fillcolor": "lightskyblue1"}, {"id": 364, "label": "c_0_364", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_364", "role": "negated_conjecture", "formula": "~horizontal_4_2|~vertical_2_4|~horizontal_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_364, negated_conjecture, (~horizontal_4_2|~vertical_2_4|~horizontal_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_211, c_0_354]), c_0_294])).\n", "fillcolor": "lightskyblue1"}, {"id": 365, "label": "c_0_365", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_365", "role": "negated_conjecture", "formula": "~horizontal_2_1|~vertical_1_2", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_365, negated_conjecture, (~horizontal_2_1|~vertical_1_2),\\ninference(fof_simplification,[status(thm)],[unique_cover6_2_2])).\n", "fillcolor": "lightskyblue1"}, {"id": 366, "label": "c_0_366", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_366", "role": "negated_conjecture", "formula": "vertical_1_2|~vertical_5_5|~vertical_2_5|~horizontal_2_3", "inference_rule": "spm", "raw_label": "cnf(c_0_366, negated_conjecture, (vertical_1_2|~vertical_5_5|~vertical_2_5|~horizontal_2_3),\\ninference(spm,[status(thm)],[c_0_164, c_0_355])).\n", "fillcolor": "lightskyblue1"}, {"id": 367, "label": "c_0_367", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_367", "role": "negated_conjecture", "formula": "vertical_2_5|vertical_1_2|~horizontal_2_3", "inference_rule": "csr", "raw_label": "cnf(c_0_367, negated_conjecture, (vertical_2_5|vertical_1_2|~horizontal_2_3),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_343, c_0_266]), c_0_356])).\n", "fillcolor": "lightskyblue1"}, {"id": 368, "label": "c_0_368", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_368", "role": "negated_conjecture", "formula": "horizontal_5_3|vertical_4_4|horizontal_2_4|~horizontal_3_4", "inference_rule": "csr", "raw_label": "cnf(c_0_368, negated_conjecture, (horizontal_5_3|vertical_4_4|horizontal_2_4|~horizontal_3_4),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_329, c_0_213]), c_0_357])).\n", "fillcolor": "lightskyblue1"}, {"id": 369, "label": "c_0_369", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_369", "role": "negated_conjecture", "formula": "horizontal_2_1|vertical_1_2|~horizontal_5_3|~horizontal_3_4", "inference_rule": "csr", "raw_label": "cnf(c_0_369, negated_conjecture, (horizontal_2_1|vertical_1_2|~horizontal_5_3|~horizontal_3_4),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_346, c_0_196]), c_0_358]), c_0_359])).\n", "fillcolor": "lightskyblue1"}, {"id": 370, "label": "c_0_370", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_370", "role": "negated_conjecture", "formula": "vertical_1_2|~horizontal_5_4|~vertical_3_5|~horizontal_2_4", "inference_rule": "spm", "raw_label": "cnf(c_0_370, negated_conjecture, (vertical_1_2|~horizontal_5_4|~vertical_3_5|~horizontal_2_4),\\ninference(spm,[status(thm)],[c_0_200, c_0_318])).\n", "fillcolor": "lightskyblue1"}, {"id": 371, "label": "c_0_371", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_371", "role": "negated_conjecture", "formula": "horizontal_5_4|horizontal_5_3|vertical_4_4|vertical_1_2|~vertical_3_5", "inference_rule": "csr", "raw_label": "cnf(c_0_371, negated_conjecture, (horizontal_5_4|horizontal_5_3|vertical_4_4|vertical_1_2|~vertical_3_5),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_332, c_0_213]), c_0_360])).\n", "fillcolor": "lightskyblue1"}, {"id": 372, "label": "c_0_372", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_372", "role": "negated_conjecture", "formula": "~horizontal_5_3|~vertical_3_5|~horizontal_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_372, negated_conjecture, (~horizontal_5_3|~vertical_3_5|~horizontal_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_211, c_0_361]), c_0_294])).\n", "fillcolor": "lightskyblue1"}, {"id": 373, "label": "c_0_373", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_373", "role": "negated_conjecture", "formula": "horizontal_2_4|~horizontal_5_3|~vertical_3_5|~horizontal_2_2", "inference_rule": "spm", "raw_label": "cnf(c_0_373, negated_conjecture, (horizontal_2_4|~horizontal_5_3|~vertical_3_5|~horizontal_2_2),\\ninference(spm,[status(thm)],[c_0_211, c_0_348])).\n", "fillcolor": "lightskyblue1"}, {"id": 374, "label": "c_0_374", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_374", "role": "negated_conjecture", "formula": "~vertical_3_3|~horizontal_3_2", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_374, negated_conjecture, (~vertical_3_3|~horizontal_3_2),\\ninference(fof_simplification,[status(thm)],[unique_cover4_3_3])).\n", "fillcolor": "lightskyblue1"}, {"id": 375, "label": "c_0_375", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_375", "role": "negated_conjecture", "formula": "~horizontal_4_3|~vertical_2_5|~horizontal_2_2", "inference_rule": "spm", "raw_label": "cnf(c_0_375, negated_conjecture, (~horizontal_4_3|~vertical_2_5|~horizontal_2_2),\\ninference(spm,[status(thm)],[c_0_241, c_0_362])).\n", "fillcolor": "lightskyblue1"}, {"id": 376, "label": "c_0_376", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_376", "role": "negated_conjecture", "formula": "horizontal_4_3|vertical_3_3|~vertical_2_4|~horizontal_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_376, negated_conjecture, (horizontal_4_3|vertical_3_3|~vertical_2_4|~horizontal_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_363, c_0_245]), c_0_364])).\n", "fillcolor": "lightskyblue1"}, {"id": 377, "label": "c_0_377", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_377", "role": "negated_conjecture", "formula": "horizontal_2_4|horizontal_2_1|~vertical_3_5|~vertical_3_2", "inference_rule": "csr", "raw_label": "cnf(c_0_377, negated_conjecture, (horizontal_2_4|horizontal_2_1|~vertical_3_5|~vertical_3_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_317, c_0_279]), c_0_216])).\n", "fillcolor": "lightskyblue1"}, {"id": 378, "label": "c_0_378", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_378", "role": "negated_conjecture", "formula": "~horizontal_2_1|~vertical_1_2", "inference_rule": null, "raw_label": "cnf(c_0_378, negated_conjecture, (~horizontal_2_1|~vertical_1_2),\\nc_0_365).\n", "fillcolor": "lightskyblue1"}, {"id": 379, "label": "c_0_379", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_379", "role": "negated_conjecture", "formula": "~vertical_2_2|~horizontal_2_1", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_379, negated_conjecture, (~vertical_2_2|~horizontal_2_1),\\ninference(fof_simplification,[status(thm)],[unique_cover4_2_2])).\n", "fillcolor": "lightskyblue1"}, {"id": 380, "label": "c_0_380", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_380", "role": "negated_conjecture", "formula": "~vertical_2_2|~vertical_1_2", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_380, negated_conjecture, (~vertical_2_2|~vertical_1_2),\\ninference(fof_simplification,[status(thm)],[unique_cover5_2_2])).\n", "fillcolor": "lightskyblue1"}, {"id": 381, "label": "c_0_381", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_381", "role": "negated_conjecture", "formula": "vertical_4_2|vertical_1_2|~horizontal_5_3|~horizontal_2_3", "inference_rule": "csr", "raw_label": "cnf(c_0_381, negated_conjecture, (vertical_4_2|vertical_1_2|~horizontal_5_3|~horizontal_2_3),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_366, c_0_181]), c_0_367])).\n", "fillcolor": "lightskyblue1"}, {"id": 382, "label": "c_0_382", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_382", "role": "negated_conjecture", "formula": "vertical_4_2|horizontal_2_4|~vertical_4_3|~horizontal_3_4", "inference_rule": "spm", "raw_label": "cnf(c_0_382, negated_conjecture, (vertical_4_2|horizontal_2_4|~vertical_4_3|~horizontal_3_4),\\ninference(spm,[status(thm)],[c_0_329, c_0_325])).\n", "fillcolor": "lightskyblue1"}, {"id": 383, "label": "c_0_383", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_383", "role": "negated_conjecture", "formula": "vertical_1_2|~horizontal_5_4|~horizontal_3_4|~horizontal_2_4", "inference_rule": "spm", "raw_label": "cnf(c_0_383, negated_conjecture, (vertical_1_2|~horizontal_5_4|~horizontal_3_4|~horizontal_2_4),\\ninference(spm,[status(thm)],[c_0_200, c_0_330])).\n", "fillcolor": "lightskyblue1"}, {"id": 384, "label": "c_0_384", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_384", "role": "negated_conjecture", "formula": "horizontal_5_4|horizontal_5_3|vertical_4_4|vertical_1_2|~horizontal_3_4", "inference_rule": "csr", "raw_label": "cnf(c_0_384, negated_conjecture, (horizontal_5_4|horizontal_5_3|vertical_4_4|vertical_1_2|~horizontal_3_4),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_346, c_0_213]), c_0_368])).\n", "fillcolor": "lightskyblue1"}, {"id": 385, "label": "c_0_385", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_385", "role": "negated_conjecture", "formula": "~horizontal_5_3|~horizontal_3_4|~horizontal_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_385, negated_conjecture, (~horizontal_5_3|~horizontal_3_4|~horizontal_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_211, c_0_369]), c_0_294])).\n", "fillcolor": "lightskyblue1"}, {"id": 386, "label": "c_0_386", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_386", "role": "negated_conjecture", "formula": "horizontal_2_4|~horizontal_5_3|~horizontal_3_4|~horizontal_2_2", "inference_rule": "spm", "raw_label": "cnf(c_0_386, negated_conjecture, (horizontal_2_4|~horizontal_5_3|~horizontal_3_4|~horizontal_2_2),\\ninference(spm,[status(thm)],[c_0_211, c_0_358])).\n", "fillcolor": "lightskyblue1"}, {"id": 387, "label": "c_0_387", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_387", "role": "negated_conjecture", "formula": "~horizontal_5_4|~vertical_3_5|~horizontal_2_4|~horizontal_2_2", "inference_rule": "spm", "raw_label": "cnf(c_0_387, negated_conjecture, (~horizontal_5_4|~vertical_3_5|~horizontal_2_4|~horizontal_2_2),\\ninference(spm,[status(thm)],[c_0_294, c_0_370])).\n", "fillcolor": "lightskyblue1"}, {"id": 388, "label": "c_0_388", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_388", "role": "negated_conjecture", "formula": "horizontal_5_4|vertical_4_4|~vertical_3_5|~horizontal_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_388, negated_conjecture, (horizontal_5_4|vertical_4_4|~vertical_3_5|~horizontal_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_294, c_0_371]), c_0_372])).\n", "fillcolor": "lightskyblue1"}, {"id": 389, "label": "c_0_389", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_389", "role": "negated_conjecture", "formula": "vertical_4_4|horizontal_2_4|~vertical_3_5|~horizontal_2_2", "inference_rule": "spm", "raw_label": "cnf(c_0_389, negated_conjecture, (vertical_4_4|horizontal_2_4|~vertical_3_5|~horizontal_2_2),\\ninference(spm,[status(thm)],[c_0_373, c_0_360])).\n", "fillcolor": "lightskyblue1"}, {"id": 390, "label": "c_0_390", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_390", "role": "negated_conjecture", "formula": "~vertical_3_3|~horizontal_3_2", "inference_rule": null, "raw_label": "cnf(c_0_390, negated_conjecture, (~vertical_3_3|~horizontal_3_2),\\nc_0_374).\n", "fillcolor": "lightskyblue1"}, {"id": 391, "label": "c_0_391", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_391", "role": "negated_conjecture", "formula": "vertical_3_3|~vertical_2_5|~horizontal_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_391, negated_conjecture, (vertical_3_3|~vertical_2_5|~horizontal_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_375, c_0_376]), c_0_352])).\n", "fillcolor": "lightskyblue1"}, {"id": 392, "label": "c_0_392", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_392", "role": "negated_conjecture", "formula": "horizontal_2_1|vertical_1_2|~vertical_3_5|~vertical_3_2", "inference_rule": "csr", "raw_label": "cnf(c_0_392, negated_conjecture, (horizontal_2_1|vertical_1_2|~vertical_3_5|~vertical_3_2),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_332, c_0_279]), c_0_377]), c_0_216])).\n", "fillcolor": "lightskyblue1"}, {"id": 393, "label": "c_0_393", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_393", "role": "negated_conjecture", "formula": "horizontal_2_4|~horizontal_5_3|~vertical_3_5|~vertical_1_2", "inference_rule": "spm", "raw_label": "cnf(c_0_393, negated_conjecture, (horizontal_2_4|~horizontal_5_3|~vertical_3_5|~vertical_1_2),\\ninference(spm,[status(thm)],[c_0_378, c_0_348])).\n", "fillcolor": "lightskyblue1"}, {"id": 394, "label": "c_0_394", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_394", "role": "negated_conjecture", "formula": "~horizontal_1_3|~vertical_1_3", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_394, negated_conjecture, (~horizontal_1_3|~vertical_1_3),\\ninference(fof_simplification,[status(thm)],[unique_cover1_1_3])).\n", "fillcolor": "lightskyblue1"}, {"id": 395, "label": "c_0_395", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_395", "role": "negated_conjecture", "formula": "~vertical_2_2|~horizontal_2_1", "inference_rule": null, "raw_label": "cnf(c_0_395, negated_conjecture, (~vertical_2_2|~horizontal_2_1),\\nc_0_379).\n", "fillcolor": "lightskyblue1"}, {"id": 396, "label": "c_0_396", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_396", "role": "negated_conjecture", "formula": "~vertical_2_2|~vertical_1_2", "inference_rule": null, "raw_label": "cnf(c_0_396, negated_conjecture, (~vertical_2_2|~vertical_1_2),\\nc_0_380).\n", "fillcolor": "lightskyblue1"}, {"id": 397, "label": "c_0_397", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_397", "role": "negated_conjecture", "formula": "horizontal_2_1|vertical_1_2|~horizontal_5_3|~horizontal_2_3", "inference_rule": "csr", "raw_label": "cnf(c_0_397, negated_conjecture, (horizontal_2_1|vertical_1_2|~horizontal_5_3|~horizontal_2_3),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_366, c_0_196]), c_0_367]), c_0_381])).\n", "fillcolor": "lightskyblue1"}, {"id": 398, "label": "c_0_398", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_398", "role": "negated_conjecture", "formula": "horizontal_2_4|horizontal_2_1|~vertical_4_3|~horizontal_3_4", "inference_rule": "csr", "raw_label": "cnf(c_0_398, negated_conjecture, (horizontal_2_4|horizontal_2_1|~vertical_4_3|~horizontal_3_4),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_329, c_0_339]), c_0_382])).\n", "fillcolor": "lightskyblue1"}, {"id": 399, "label": "c_0_399", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_399", "role": "negated_conjecture", "formula": "vertical_4_2|vertical_1_2|~vertical_4_3|~horizontal_3_4", "inference_rule": "csr", "raw_label": "cnf(c_0_399, negated_conjecture, (vertical_4_2|vertical_1_2|~vertical_4_3|~horizontal_3_4),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_346, c_0_325]), c_0_382])).\n", "fillcolor": "lightskyblue1"}, {"id": 400, "label": "c_0_400", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_400", "role": "negated_conjecture", "formula": "horizontal_2_4|horizontal_2_1|~horizontal_4_2|~horizontal_3_4", "inference_rule": "csr", "raw_label": "cnf(c_0_400, negated_conjecture, (horizontal_2_4|horizontal_2_1|~horizontal_4_2|~horizontal_3_4),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_329, c_0_341]), c_0_342])).\n", "fillcolor": "lightskyblue1"}, {"id": 401, "label": "c_0_401", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_401", "role": "negated_conjecture", "formula": "~horizontal_5_4|~horizontal_3_4|~horizontal_2_4|~horizontal_2_2", "inference_rule": "spm", "raw_label": "cnf(c_0_401, negated_conjecture, (~horizontal_5_4|~horizontal_3_4|~horizontal_2_4|~horizontal_2_2),\\ninference(spm,[status(thm)],[c_0_294, c_0_383])).\n", "fillcolor": "lightskyblue1"}, {"id": 402, "label": "c_0_402", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_402", "role": "negated_conjecture", "formula": "horizontal_5_4|vertical_4_4|~horizontal_3_4|~horizontal_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_402, negated_conjecture, (horizontal_5_4|vertical_4_4|~horizontal_3_4|~horizontal_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_294, c_0_384]), c_0_385])).\n", "fillcolor": "lightskyblue1"}, {"id": 403, "label": "c_0_403", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_403", "role": "negated_conjecture", "formula": "vertical_4_4|horizontal_2_4|~horizontal_3_4|~horizontal_2_2", "inference_rule": "spm", "raw_label": "cnf(c_0_403, negated_conjecture, (vertical_4_4|horizontal_2_4|~horizontal_3_4|~horizontal_2_2),\\ninference(spm,[status(thm)],[c_0_386, c_0_368])).\n", "fillcolor": "lightskyblue1"}, {"id": 404, "label": "c_0_404", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_404", "role": "negated_conjecture", "formula": "horizontal_2_4|horizontal_2_1|~horizontal_3_4|~vertical_3_2", "inference_rule": "csr", "raw_label": "cnf(c_0_404, negated_conjecture, (horizontal_2_4|horizontal_2_1|~horizontal_3_4|~vertical_3_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_329, c_0_279]), c_0_216])).\n", "fillcolor": "lightskyblue1"}, {"id": 405, "label": "c_0_405", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_405", "role": "negated_conjecture", "formula": "vertical_4_4|~vertical_3_5|~horizontal_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_405, negated_conjecture, (vertical_4_4|~vertical_3_5|~horizontal_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_387, c_0_388]), c_0_389])).\n", "fillcolor": "lightskyblue1"}, {"id": 406, "label": "c_0_406", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_406", "role": "negated_conjecture", "formula": "~vertical_2_5|~horizontal_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_406, negated_conjecture, (~vertical_2_5|~horizontal_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_390, c_0_391]), c_0_338])).\n", "fillcolor": "lightskyblue1"}, {"id": 407, "label": "c_0_407", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_407", "role": "negated_conjecture", "formula": "vertical_2_5|~vertical_2_4|~horizontal_2_2", "inference_rule": "spm", "raw_label": "cnf(c_0_407, negated_conjecture, (vertical_2_5|~vertical_2_4|~horizontal_2_2),\\ninference(spm,[status(thm)],[c_0_294, c_0_280])).\n", "fillcolor": "lightskyblue1"}, {"id": 408, "label": "c_0_408", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_408", "role": "negated_conjecture", "formula": "~vertical_3_5|~vertical_3_2|~horizontal_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_408, negated_conjecture, (~vertical_3_5|~vertical_3_2|~horizontal_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_211, c_0_392]), c_0_294])).\n", "fillcolor": "lightskyblue1"}, {"id": 409, "label": "c_0_409", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_409", "role": "negated_conjecture", "formula": "horizontal_4_4|vertical_3_5|~horizontal_5_3|~vertical_1_2", "inference_rule": "spm", "raw_label": "cnf(c_0_409, negated_conjecture, (horizontal_4_4|vertical_3_5|~horizontal_5_3|~vertical_1_2),\\ninference(spm,[status(thm)],[c_0_378, c_0_212])).\n", "fillcolor": "lightskyblue1"}, {"id": 410, "label": "c_0_410", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_410", "role": "negated_conjecture", "formula": "~horizontal_5_3|~vertical_3_5|~vertical_2_4", "inference_rule": "csr", "raw_label": "cnf(c_0_410, negated_conjecture, (~horizontal_5_3|~vertical_3_5|~vertical_2_4),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_280]), c_0_267]), c_0_274])).\n", "fillcolor": "lightskyblue1"}, {"id": 411, "label": "c_0_411", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_411", "role": "negated_conjecture", "formula": "~horizontal_1_3|~vertical_1_3", "inference_rule": null, "raw_label": "cnf(c_0_411, negated_conjecture, (~horizontal_1_3|~vertical_1_3),\\nc_0_394).\n", "fillcolor": "lightskyblue1"}, {"id": 412, "label": "c_0_412", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_412", "role": "negated_conjecture", "formula": "~horizontal_5_3|~horizontal_3_4|~vertical_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_412, negated_conjecture, (~horizontal_5_3|~horizontal_3_4|~vertical_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_395, c_0_369]), c_0_396])).\n", "fillcolor": "lightskyblue1"}, {"id": 413, "label": "c_0_413", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_413", "role": "negated_conjecture", "formula": "horizontal_2_4|~horizontal_5_3|~horizontal_3_4|~vertical_2_2", "inference_rule": "spm", "raw_label": "cnf(c_0_413, negated_conjecture, (horizontal_2_4|~horizontal_5_3|~horizontal_3_4|~vertical_2_2),\\ninference(spm,[status(thm)],[c_0_395, c_0_358])).\n", "fillcolor": "lightskyblue1"}, {"id": 414, "label": "c_0_414", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_414", "role": "negated_conjecture", "formula": "~horizontal_5_3|~vertical_3_5|~vertical_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_414, negated_conjecture, (~horizontal_5_3|~vertical_3_5|~vertical_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_395, c_0_361]), c_0_396])).\n", "fillcolor": "lightskyblue1"}, {"id": 415, "label": "c_0_415", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_415", "role": "negated_conjecture", "formula": "horizontal_2_4|~horizontal_5_3|~vertical_3_5|~vertical_2_2", "inference_rule": "spm", "raw_label": "cnf(c_0_415, negated_conjecture, (horizontal_2_4|~horizontal_5_3|~vertical_3_5|~vertical_2_2),\\ninference(spm,[status(thm)],[c_0_395, c_0_348])).\n", "fillcolor": "lightskyblue1"}, {"id": 416, "label": "c_0_416", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_416", "role": "negated_conjecture", "formula": "horizontal_4_4|vertical_3_5|~horizontal_5_3|~vertical_2_2", "inference_rule": "spm", "raw_label": "cnf(c_0_416, negated_conjecture, (horizontal_4_4|vertical_3_5|~horizontal_5_3|~vertical_2_2),\\ninference(spm,[status(thm)],[c_0_395, c_0_212])).\n", "fillcolor": "lightskyblue1"}, {"id": 417, "label": "c_0_417", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_417", "role": "negated_conjecture", "formula": "~horizontal_5_3|~vertical_2_4|~vertical_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_417, negated_conjecture, (~horizontal_5_3|~vertical_2_4|~vertical_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_395, c_0_319]), c_0_396])).\n", "fillcolor": "lightskyblue1"}, {"id": 418, "label": "c_0_418", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_418", "role": "negated_conjecture", "formula": "vertical_1_2|~horizontal_5_4|~vertical_2_5|~horizontal_2_3", "inference_rule": "spm", "raw_label": "cnf(c_0_418, negated_conjecture, (vertical_1_2|~horizontal_5_4|~vertical_2_5|~horizontal_2_3),\\ninference(spm,[status(thm)],[c_0_200, c_0_355])).\n", "fillcolor": "lightskyblue1"}, {"id": 419, "label": "c_0_419", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_419", "role": "negated_conjecture", "formula": "horizontal_5_4|horizontal_5_3|vertical_4_4|vertical_1_2|~horizontal_2_3", "inference_rule": "csr", "raw_label": "cnf(c_0_419, negated_conjecture, (horizontal_5_4|horizontal_5_3|vertical_4_4|vertical_1_2|~horizontal_2_3),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_366, c_0_213]), c_0_367])).\n", "fillcolor": "lightskyblue1"}, {"id": 420, "label": "c_0_420", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_420", "role": "negated_conjecture", "formula": "~horizontal_5_3|~horizontal_2_3|~vertical_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_420, negated_conjecture, (~horizontal_5_3|~horizontal_2_3|~vertical_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_395, c_0_397]), c_0_396])).\n", "fillcolor": "lightskyblue1"}, {"id": 421, "label": "c_0_421", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_421", "role": "negated_conjecture", "formula": "horizontal_2_1|vertical_1_2|~vertical_4_3|~horizontal_3_4", "inference_rule": "csr", "raw_label": "cnf(c_0_421, negated_conjecture, (horizontal_2_1|vertical_1_2|~vertical_4_3|~horizontal_3_4),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_346, c_0_339]), c_0_398]), c_0_399])).\n", "fillcolor": "lightskyblue1"}, {"id": 422, "label": "c_0_422", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_422", "role": "negated_conjecture", "formula": "horizontal_2_1|vertical_1_2|~horizontal_4_2|~horizontal_3_4", "inference_rule": "csr", "raw_label": "cnf(c_0_422, negated_conjecture, (horizontal_2_1|vertical_1_2|~horizontal_4_2|~horizontal_3_4),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_346, c_0_341]), c_0_400]), c_0_342])).\n", "fillcolor": "lightskyblue1"}, {"id": 423, "label": "c_0_423", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_423", "role": "negated_conjecture", "formula": "vertical_4_4|~horizontal_3_4|~horizontal_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_423, negated_conjecture, (vertical_4_4|~horizontal_3_4|~horizontal_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_401, c_0_402]), c_0_403])).\n", "fillcolor": "lightskyblue1"}, {"id": 424, "label": "c_0_424", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_424", "role": "negated_conjecture", "formula": "horizontal_2_1|vertical_1_2|~horizontal_3_4|~vertical_3_2", "inference_rule": "csr", "raw_label": "cnf(c_0_424, negated_conjecture, (horizontal_2_1|vertical_1_2|~horizontal_3_4|~vertical_3_2),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_346, c_0_279]), c_0_404]), c_0_216])).\n", "fillcolor": "lightskyblue1"}, {"id": 425, "label": "c_0_425", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_425", "role": "negated_conjecture", "formula": "~vertical_3_4|~horizontal_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_425, negated_conjecture, (~vertical_3_4|~horizontal_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_247, c_0_405]), c_0_305])).\n", "fillcolor": "lightskyblue1"}, {"id": 426, "label": "c_0_426", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_426", "role": "negated_conjecture", "formula": "~vertical_2_4|~horizontal_2_2", "inference_rule": "spm", "raw_label": "cnf(c_0_426, negated_conjecture, (~vertical_2_4|~horizontal_2_2),\\ninference(spm,[status(thm)],[c_0_406, c_0_407])).\n", "fillcolor": "lightskyblue1"}, {"id": 427, "label": "c_0_427", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_427", "role": "negated_conjecture", "formula": "vertical_3_3|~vertical_3_2|~horizontal_2_2", "inference_rule": "spm", "raw_label": "cnf(c_0_427, negated_conjecture, (vertical_3_3|~vertical_3_2|~horizontal_2_2),\\ninference(spm,[status(thm)],[c_0_408, c_0_275])).\n", "fillcolor": "lightskyblue1"}, {"id": 428, "label": "c_0_428", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_428", "role": "negated_conjecture", "formula": "horizontal_4_4|vertical_3_5|vertical_2_5|~horizontal_5_3|~vertical_2_4", "inference_rule": "spm", "raw_label": "cnf(c_0_428, negated_conjecture, (horizontal_4_4|vertical_3_5|vertical_2_5|~horizontal_5_3|~vertical_2_4),\\ninference(spm,[status(thm)],[c_0_409, c_0_280])).\n", "fillcolor": "lightskyblue1"}, {"id": 429, "label": "c_0_429", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_429", "role": "negated_conjecture", "formula": "vertical_4_4|~vertical_3_5|~vertical_2_4", "inference_rule": "csr", "raw_label": "cnf(c_0_429, negated_conjecture, (vertical_4_4|~vertical_3_5|~vertical_2_4),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_410, c_0_360]), c_0_267])).\n", "fillcolor": "lightskyblue1"}, {"id": 430, "label": "c_0_430", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_430", "role": "negated_conjecture", "formula": "horizontal_5_2|~vertical_3_2|~vertical_1_2", "inference_rule": "spm", "raw_label": "cnf(c_0_430, negated_conjecture, (horizontal_5_2|~vertical_3_2|~vertical_1_2),\\ninference(spm,[status(thm)],[c_0_378, c_0_216])).\n", "fillcolor": "lightskyblue1"}, {"id": 431, "label": "c_0_431", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_431", "role": "negated_conjecture", "formula": "vertical_1_4|~horizontal_2_5|~vertical_1_3", "inference_rule": "spm", "raw_label": "cnf(c_0_431, negated_conjecture, (vertical_1_4|~horizontal_2_5|~vertical_1_3),\\ninference(spm,[status(thm)],[c_0_411, c_0_220])).\n", "fillcolor": "lightskyblue1"}, {"id": 432, "label": "c_0_432", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_432", "role": "negated_conjecture", "formula": "vertical_1_4|~vertical_1_5|~vertical_1_3", "inference_rule": "spm", "raw_label": "cnf(c_0_432, negated_conjecture, (vertical_1_4|~vertical_1_5|~vertical_1_3),\\ninference(spm,[status(thm)],[c_0_411, c_0_221])).\n", "fillcolor": "lightskyblue1"}, {"id": 433, "label": "c_0_433", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_433", "role": "negated_conjecture", "formula": "~horizontal_5_4|~horizontal_3_4|~horizontal_2_4|~vertical_2_2", "inference_rule": "spm", "raw_label": "cnf(c_0_433, negated_conjecture, (~horizontal_5_4|~horizontal_3_4|~horizontal_2_4|~vertical_2_2),\\ninference(spm,[status(thm)],[c_0_396, c_0_383])).\n", "fillcolor": "lightskyblue1"}, {"id": 434, "label": "c_0_434", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_434", "role": "negated_conjecture", "formula": "horizontal_5_4|vertical_4_4|~horizontal_3_4|~vertical_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_434, negated_conjecture, (horizontal_5_4|vertical_4_4|~horizontal_3_4|~vertical_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_396, c_0_384]), c_0_412])).\n", "fillcolor": "lightskyblue1"}, {"id": 435, "label": "c_0_435", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_435", "role": "negated_conjecture", "formula": "vertical_4_4|horizontal_2_4|~horizontal_3_4|~vertical_2_2", "inference_rule": "spm", "raw_label": "cnf(c_0_435, negated_conjecture, (vertical_4_4|horizontal_2_4|~horizontal_3_4|~vertical_2_2),\\ninference(spm,[status(thm)],[c_0_413, c_0_368])).\n", "fillcolor": "lightskyblue1"}, {"id": 436, "label": "c_0_436", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_436", "role": "negated_conjecture", "formula": "~horizontal_5_4|~vertical_3_5|~horizontal_2_4|~vertical_2_2", "inference_rule": "spm", "raw_label": "cnf(c_0_436, negated_conjecture, (~horizontal_5_4|~vertical_3_5|~horizontal_2_4|~vertical_2_2),\\ninference(spm,[status(thm)],[c_0_396, c_0_370])).\n", "fillcolor": "lightskyblue1"}, {"id": 437, "label": "c_0_437", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_437", "role": "negated_conjecture", "formula": "horizontal_5_4|vertical_4_4|~vertical_3_5|~vertical_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_437, negated_conjecture, (horizontal_5_4|vertical_4_4|~vertical_3_5|~vertical_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_396, c_0_371]), c_0_414])).\n", "fillcolor": "lightskyblue1"}, {"id": 438, "label": "c_0_438", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_438", "role": "negated_conjecture", "formula": "vertical_4_4|horizontal_2_4|~vertical_3_5|~vertical_2_2", "inference_rule": "spm", "raw_label": "cnf(c_0_438, negated_conjecture, (vertical_4_4|horizontal_2_4|~vertical_3_5|~vertical_2_2),\\ninference(spm,[status(thm)],[c_0_415, c_0_360])).\n", "fillcolor": "lightskyblue1"}, {"id": 439, "label": "c_0_439", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_439", "role": "negated_conjecture", "formula": "vertical_4_4|horizontal_4_4|vertical_3_5|~vertical_2_2", "inference_rule": "spm", "raw_label": "cnf(c_0_439, negated_conjecture, (vertical_4_4|horizontal_4_4|vertical_3_5|~vertical_2_2),\\ninference(spm,[status(thm)],[c_0_416, c_0_226])).\n", "fillcolor": "lightskyblue1"}, {"id": 440, "label": "c_0_440", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_440", "role": "negated_conjecture", "formula": "~horizontal_5_4|~vertical_2_5|~vertical_2_4|~vertical_2_2", "inference_rule": "spm", "raw_label": "cnf(c_0_440, negated_conjecture, (~horizontal_5_4|~vertical_2_5|~vertical_2_4|~vertical_2_2),\\ninference(spm,[status(thm)],[c_0_396, c_0_333])).\n", "fillcolor": "lightskyblue1"}, {"id": 441, "label": "c_0_441", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_441", "role": "negated_conjecture", "formula": "horizontal_5_4|vertical_4_4|~vertical_2_4|~vertical_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_441, negated_conjecture, (horizontal_5_4|vertical_4_4|~vertical_2_4|~vertical_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_396, c_0_334]), c_0_417])).\n", "fillcolor": "lightskyblue1"}, {"id": 442, "label": "c_0_442", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_442", "role": "negated_conjecture", "formula": "vertical_2_5|~vertical_2_4|~vertical_2_2", "inference_rule": "spm", "raw_label": "cnf(c_0_442, negated_conjecture, (vertical_2_5|~vertical_2_4|~vertical_2_2),\\ninference(spm,[status(thm)],[c_0_396, c_0_280])).\n", "fillcolor": "lightskyblue1"}, {"id": 443, "label": "c_0_443", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_443", "role": "negated_conjecture", "formula": "vertical_4_2|vertical_1_2|~vertical_4_3|~horizontal_2_3", "inference_rule": "csr", "raw_label": "cnf(c_0_443, negated_conjecture, (vertical_4_2|vertical_1_2|~vertical_4_3|~horizontal_2_3),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_366, c_0_325]), c_0_367])).\n", "fillcolor": "lightskyblue1"}, {"id": 444, "label": "c_0_444", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_444", "role": "negated_conjecture", "formula": "~horizontal_5_4|~vertical_2_5|~horizontal_2_3|~vertical_2_2", "inference_rule": "spm", "raw_label": "cnf(c_0_444, negated_conjecture, (~horizontal_5_4|~vertical_2_5|~horizontal_2_3|~vertical_2_2),\\ninference(spm,[status(thm)],[c_0_396, c_0_418])).\n", "fillcolor": "lightskyblue1"}, {"id": 445, "label": "c_0_445", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_445", "role": "negated_conjecture", "formula": "horizontal_5_4|vertical_4_4|~horizontal_2_3|~vertical_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_445, negated_conjecture, (horizontal_5_4|vertical_4_4|~horizontal_2_3|~vertical_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_396, c_0_419]), c_0_420])).\n", "fillcolor": "lightskyblue1"}, {"id": 446, "label": "c_0_446", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_446", "role": "negated_conjecture", "formula": "vertical_2_5|~horizontal_2_3|~vertical_2_2", "inference_rule": "spm", "raw_label": "cnf(c_0_446, negated_conjecture, (vertical_2_5|~horizontal_2_3|~vertical_2_2),\\ninference(spm,[status(thm)],[c_0_396, c_0_367])).\n", "fillcolor": "lightskyblue1"}, {"id": 447, "label": "c_0_447", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_447", "role": "negated_conjecture", "formula": "~horizontal_5_3|~vertical_4_3", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_447, negated_conjecture, (~horizontal_5_3|~vertical_4_3),\\ninference(fof_simplification,[status(thm)],[unique_cover3_5_3])).\n", "fillcolor": "lightskyblue1"}, {"id": 448, "label": "c_0_448", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_448", "role": "negated_conjecture", "formula": "~vertical_4_3|~horizontal_3_4|~horizontal_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_448, negated_conjecture, (~vertical_4_3|~horizontal_3_4|~horizontal_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_211, c_0_421]), c_0_294])).\n", "fillcolor": "lightskyblue1"}, {"id": 449, "label": "c_0_449", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_449", "role": "negated_conjecture", "formula": "~horizontal_4_2|~horizontal_3_4|~horizontal_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_449, negated_conjecture, (~horizontal_4_2|~horizontal_3_4|~horizontal_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_211, c_0_422]), c_0_294])).\n", "fillcolor": "lightskyblue1"}, {"id": 450, "label": "c_0_450", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_450", "role": "negated_conjecture", "formula": "~horizontal_4_3|~horizontal_3_4|~horizontal_2_2", "inference_rule": "spm", "raw_label": "cnf(c_0_450, negated_conjecture, (~horizontal_4_3|~horizontal_3_4|~horizontal_2_2),\\ninference(spm,[status(thm)],[c_0_241, c_0_423])).\n", "fillcolor": "lightskyblue1"}, {"id": 451, "label": "c_0_451", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_451", "role": "negated_conjecture", "formula": "~horizontal_3_4|~vertical_3_2|~horizontal_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_451, negated_conjecture, (~horizontal_3_4|~vertical_3_2|~horizontal_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_211, c_0_424]), c_0_294])).\n", "fillcolor": "lightskyblue1"}, {"id": 452, "label": "c_0_452", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_452", "role": "negated_conjecture", "formula": "horizontal_3_4|horizontal_3_3|~horizontal_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_452, negated_conjecture, (horizontal_3_4|horizontal_3_3|~horizontal_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_425, c_0_291]), c_0_426])).\n", "fillcolor": "lightskyblue1"}, {"id": 453, "label": "c_0_453", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_453", "role": "negated_conjecture", "formula": "~horizontal_3_3|~vertical_3_2|~horizontal_2_2", "inference_rule": "spm", "raw_label": "cnf(c_0_453, negated_conjecture, (~horizontal_3_3|~vertical_3_2|~horizontal_2_2),\\ninference(spm,[status(thm)],[c_0_288, c_0_427])).\n", "fillcolor": "lightskyblue1"}, {"id": 454, "label": "c_0_454", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_454", "role": "negated_conjecture", "formula": "vertical_4_4|horizontal_4_4|vertical_2_5|~vertical_2_4", "inference_rule": "csr", "raw_label": "cnf(c_0_454, negated_conjecture, (vertical_4_4|horizontal_4_4|vertical_2_5|~vertical_2_4),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_428, c_0_226]), c_0_429])).\n", "fillcolor": "lightskyblue1"}, {"id": 455, "label": "c_0_455", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_455", "role": "negated_conjecture", "formula": "~vertical_4_3|~vertical_3_2|~vertical_2_4", "inference_rule": "csr", "raw_label": "cnf(c_0_455, negated_conjecture, (~vertical_4_3|~vertical_3_2|~vertical_2_4),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_430, c_0_340]), c_0_204]), c_0_228])).\n", "fillcolor": "lightskyblue1"}, {"id": 456, "label": "c_0_456", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_456", "role": "negated_conjecture", "formula": "vertical_2_5|horizontal_2_4|vertical_1_4|~vertical_1_3", "inference_rule": "csr", "raw_label": "cnf(c_0_456, negated_conjecture, (vertical_2_5|horizontal_2_4|vertical_1_4|~vertical_1_3),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_431, c_0_254]), c_0_432])).\n", "fillcolor": "lightskyblue1"}, {"id": 457, "label": "c_0_457", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_457", "role": "negated_conjecture", "formula": "~vertical_2_4|~horizontal_2_3", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_457, negated_conjecture, (~vertical_2_4|~horizontal_2_3),\\ninference(fof_simplification,[status(thm)],[unique_cover4_2_4])).\n", "fillcolor": "lightskyblue1"}, {"id": 458, "label": "c_0_458", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_458", "role": "negated_conjecture", "formula": "vertical_4_4|~horizontal_3_4|~vertical_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_458, negated_conjecture, (vertical_4_4|~horizontal_3_4|~vertical_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_433, c_0_434]), c_0_435])).\n", "fillcolor": "lightskyblue1"}, {"id": 459, "label": "c_0_459", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_459", "role": "negated_conjecture", "formula": "vertical_4_4|~vertical_3_5|~vertical_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_459, negated_conjecture, (vertical_4_4|~vertical_3_5|~vertical_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_436, c_0_437]), c_0_438])).\n", "fillcolor": "lightskyblue1"}, {"id": 460, "label": "c_0_460", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_460", "role": "negated_conjecture", "formula": "vertical_3_5|~vertical_3_4|~vertical_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_460, negated_conjecture, (vertical_3_5|~vertical_3_4|~vertical_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_247, c_0_439]), c_0_249])).\n", "fillcolor": "lightskyblue1"}, {"id": 461, "label": "c_0_461", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_461", "role": "negated_conjecture", "formula": "vertical_4_4|~vertical_2_4|~vertical_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_461, negated_conjecture, (vertical_4_4|~vertical_2_4|~vertical_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_440, c_0_441]), c_0_442])).\n", "fillcolor": "lightskyblue1"}, {"id": 462, "label": "c_0_462", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_462", "role": "negated_conjecture", "formula": "~vertical_1_3|~horizontal_1_2", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_462, negated_conjecture, (~vertical_1_3|~horizontal_1_2),\\ninference(fof_simplification,[status(thm)],[unique_cover4_1_3])).\n", "fillcolor": "lightskyblue1"}, {"id": 463, "label": "c_0_463", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_463", "role": "negated_conjecture", "formula": "horizontal_2_1|vertical_1_2|~vertical_4_3|~horizontal_2_3", "inference_rule": "csr", "raw_label": "cnf(c_0_463, negated_conjecture, (horizontal_2_1|vertical_1_2|~vertical_4_3|~horizontal_2_3),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_366, c_0_339]), c_0_367]), c_0_443])).\n", "fillcolor": "lightskyblue1"}, {"id": 464, "label": "c_0_464", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_464", "role": "negated_conjecture", "formula": "horizontal_2_1|vertical_1_2|~horizontal_4_2|~horizontal_2_3", "inference_rule": "csr", "raw_label": "cnf(c_0_464, negated_conjecture, (horizontal_2_1|vertical_1_2|~horizontal_4_2|~horizontal_2_3),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_366, c_0_341]), c_0_367]), c_0_342])).\n", "fillcolor": "lightskyblue1"}, {"id": 465, "label": "c_0_465", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_465", "role": "negated_conjecture", "formula": "vertical_4_4|~horizontal_2_3|~vertical_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_465, negated_conjecture, (vertical_4_4|~horizontal_2_3|~vertical_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_444, c_0_445]), c_0_446])).\n", "fillcolor": "lightskyblue1"}, {"id": 466, "label": "c_0_466", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_466", "role": "negated_conjecture", "formula": "vertical_4_2|horizontal_2_4|~vertical_4_3|~vertical_3_5", "inference_rule": "spm", "raw_label": "cnf(c_0_466, negated_conjecture, (vertical_4_2|horizontal_2_4|~vertical_4_3|~vertical_3_5),\\ninference(spm,[status(thm)],[c_0_317, c_0_325])).\n", "fillcolor": "lightskyblue1"}, {"id": 467, "label": "c_0_467", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_467", "role": "negated_conjecture", "formula": "~horizontal_5_3|~vertical_4_3", "inference_rule": null, "raw_label": "cnf(c_0_467, negated_conjecture, (~horizontal_5_3|~vertical_4_3),\\nc_0_447).\n", "fillcolor": "lightskyblue1"}, {"id": 468, "label": "c_0_468", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_468", "role": "negated_conjecture", "formula": "~vertical_3_5|~horizontal_3_4", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_468, negated_conjecture, (~vertical_3_5|~horizontal_3_4),\\ninference(fof_simplification,[status(thm)],[unique_cover4_3_5])).\n", "fillcolor": "lightskyblue1"}, {"id": 469, "label": "c_0_469", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_469", "role": "negated_conjecture", "formula": "vertical_3_3|~horizontal_3_4|~horizontal_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_469, negated_conjecture, (vertical_3_3|~horizontal_3_4|~horizontal_2_2),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_448, c_0_245]), c_0_449]), c_0_450])).\n", "fillcolor": "lightskyblue1"}, {"id": 470, "label": "c_0_470", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_470", "role": "negated_conjecture", "formula": "~vertical_3_2|~horizontal_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_470, negated_conjecture, (~vertical_3_2|~horizontal_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_451, c_0_452]), c_0_453])).\n", "fillcolor": "lightskyblue1"}, {"id": 471, "label": "c_0_471", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_471", "role": "negated_conjecture", "formula": "~vertical_3_3|~vertical_2_3", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_471, negated_conjecture, (~vertical_3_3|~vertical_2_3),\\ninference(fof_simplification,[status(thm)],[unique_cover5_3_3])).\n", "fillcolor": "lightskyblue1"}, {"id": 472, "label": "c_0_472", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_472", "role": "negated_conjecture", "formula": "vertical_2_5|~horizontal_4_3|~vertical_2_4", "inference_rule": "csr", "raw_label": "cnf(c_0_472, negated_conjecture, (vertical_2_5|~horizontal_4_3|~vertical_2_4),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_241, c_0_454]), c_0_243])).\n", "fillcolor": "lightskyblue1"}, {"id": 473, "label": "c_0_473", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_473", "role": "negated_conjecture", "formula": "horizontal_4_3|vertical_3_3|~vertical_3_2|~vertical_2_4", "inference_rule": "csr", "raw_label": "cnf(c_0_473, negated_conjecture, (horizontal_4_3|vertical_3_3|~vertical_3_2|~vertical_2_4),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_455, c_0_245]), c_0_246])).\n", "fillcolor": "lightskyblue1"}, {"id": 474, "label": "c_0_474", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_474", "role": "negated_conjecture", "formula": "vertical_2_5|~vertical_2_4|~vertical_1_3", "inference_rule": "csr", "raw_label": "cnf(c_0_474, negated_conjecture, (vertical_2_5|~vertical_2_4|~vertical_1_3),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_250, c_0_456]), c_0_267])).\n", "fillcolor": "lightskyblue1"}, {"id": 475, "label": "c_0_475", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_475", "role": "negated_conjecture", "formula": "horizontal_2_3|vertical_2_3|horizontal_2_2|vertical_1_3", "inference_rule": null, "raw_label": "cnf(c_0_475, negated_conjecture, (horizontal_2_3|vertical_2_3|horizontal_2_2|vertical_1_3),\\ncovered_2_3).\n", "fillcolor": "lightskyblue1"}, {"id": 476, "label": "c_0_476", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_476", "role": "negated_conjecture", "formula": "~vertical_2_4|~horizontal_2_3", "inference_rule": null, "raw_label": "cnf(c_0_476, negated_conjecture, (~vertical_2_4|~horizontal_2_3),\\nc_0_457).\n", "fillcolor": "lightskyblue1"}, {"id": 477, "label": "c_0_477", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_477", "role": "negated_conjecture", "formula": "horizontal_3_5|vertical_1_4|~vertical_4_6|~vertical_1_3", "inference_rule": "spm", "raw_label": "cnf(c_0_477, negated_conjecture, (horizontal_3_5|vertical_1_4|~vertical_4_6|~vertical_1_3),\\ninference(spm,[status(thm)],[c_0_411, c_0_190])).\n", "fillcolor": "lightskyblue1"}, {"id": 478, "label": "c_0_478", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_478", "role": "negated_conjecture", "formula": "~horizontal_3_1|~vertical_3_1", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_478, negated_conjecture, (~horizontal_3_1|~vertical_3_1),\\ninference(fof_simplification,[status(thm)],[unique_cover1_3_1])).\n", "fillcolor": "lightskyblue1"}, {"id": 479, "label": "c_0_479", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_479", "role": "negated_conjecture", "formula": "~vertical_4_3|~horizontal_3_4|~vertical_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_479, negated_conjecture, (~vertical_4_3|~horizontal_3_4|~vertical_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_395, c_0_421]), c_0_396])).\n", "fillcolor": "lightskyblue1"}, {"id": 480, "label": "c_0_480", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_480", "role": "negated_conjecture", "formula": "~horizontal_4_2|~horizontal_3_4|~vertical_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_480, negated_conjecture, (~horizontal_4_2|~horizontal_3_4|~vertical_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_395, c_0_422]), c_0_396])).\n", "fillcolor": "lightskyblue1"}, {"id": 481, "label": "c_0_481", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_481", "role": "negated_conjecture", "formula": "~horizontal_4_3|~horizontal_3_4|~vertical_2_2", "inference_rule": "spm", "raw_label": "cnf(c_0_481, negated_conjecture, (~horizontal_4_3|~horizontal_3_4|~vertical_2_2),\\ninference(spm,[status(thm)],[c_0_241, c_0_458])).\n", "fillcolor": "lightskyblue1"}, {"id": 482, "label": "c_0_482", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_482", "role": "negated_conjecture", "formula": "~vertical_3_4|~vertical_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_482, negated_conjecture, (~vertical_3_4|~vertical_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_247, c_0_459]), c_0_460])).\n", "fillcolor": "lightskyblue1"}, {"id": 483, "label": "c_0_483", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_483", "role": "negated_conjecture", "formula": "~vertical_4_3|~vertical_2_4|~vertical_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_483, negated_conjecture, (~vertical_4_3|~vertical_2_4|~vertical_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_395, c_0_353]), c_0_396])).\n", "fillcolor": "lightskyblue1"}, {"id": 484, "label": "c_0_484", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_484", "role": "negated_conjecture", "formula": "~horizontal_4_2|~vertical_2_4|~vertical_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_484, negated_conjecture, (~horizontal_4_2|~vertical_2_4|~vertical_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_395, c_0_354]), c_0_396])).\n", "fillcolor": "lightskyblue1"}, {"id": 485, "label": "c_0_485", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_485", "role": "negated_conjecture", "formula": "~horizontal_4_3|~vertical_2_4|~vertical_2_2", "inference_rule": "spm", "raw_label": "cnf(c_0_485, negated_conjecture, (~horizontal_4_3|~vertical_2_4|~vertical_2_2),\\ninference(spm,[status(thm)],[c_0_241, c_0_461])).\n", "fillcolor": "lightskyblue1"}, {"id": 486, "label": "c_0_486", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_486", "role": "negated_conjecture", "formula": "~vertical_1_3|~horizontal_1_2", "inference_rule": null, "raw_label": "cnf(c_0_486, negated_conjecture, (~vertical_1_3|~horizontal_1_2),\\nc_0_462).\n", "fillcolor": "lightskyblue1"}, {"id": 487, "label": "c_0_487", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_487", "role": "negated_conjecture", "formula": "~vertical_4_3|~horizontal_2_3|~vertical_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_487, negated_conjecture, (~vertical_4_3|~horizontal_2_3|~vertical_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_395, c_0_463]), c_0_396])).\n", "fillcolor": "lightskyblue1"}, {"id": 488, "label": "c_0_488", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_488", "role": "negated_conjecture", "formula": "~horizontal_4_2|~horizontal_2_3|~vertical_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_488, negated_conjecture, (~horizontal_4_2|~horizontal_2_3|~vertical_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_395, c_0_464]), c_0_396])).\n", "fillcolor": "lightskyblue1"}, {"id": 489, "label": "c_0_489", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_489", "role": "negated_conjecture", "formula": "~horizontal_4_3|~horizontal_2_3|~vertical_2_2", "inference_rule": "spm", "raw_label": "cnf(c_0_489, negated_conjecture, (~horizontal_4_3|~horizontal_2_3|~vertical_2_2),\\ninference(spm,[status(thm)],[c_0_241, c_0_465])).\n", "fillcolor": "lightskyblue1"}, {"id": 490, "label": "c_0_490", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_490", "role": "negated_conjecture", "formula": "~horizontal_5_3|~vertical_3_5|~horizontal_2_3", "inference_rule": "csr", "raw_label": "cnf(c_0_490, negated_conjecture, (~horizontal_5_3|~vertical_3_5|~horizontal_2_3),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_367]), c_0_356]), c_0_274])).\n", "fillcolor": "lightskyblue1"}, {"id": 491, "label": "c_0_491", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_491", "role": "negated_conjecture", "formula": "vertical_4_4|vertical_4_2|horizontal_2_4|~vertical_3_5", "inference_rule": "spm", "raw_label": "cnf(c_0_491, negated_conjecture, (vertical_4_4|vertical_4_2|horizontal_2_4|~vertical_3_5),\\ninference(spm,[status(thm)],[c_0_331, c_0_360])).\n", "fillcolor": "lightskyblue1"}, {"id": 492, "label": "c_0_492", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_492", "role": "negated_conjecture", "formula": "vertical_4_2|vertical_1_2|~vertical_4_3|~vertical_3_5", "inference_rule": "csr", "raw_label": "cnf(c_0_492, negated_conjecture, (vertical_4_2|vertical_1_2|~vertical_4_3|~vertical_3_5),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_332, c_0_325]), c_0_466])).\n", "fillcolor": "lightskyblue1"}, {"id": 493, "label": "c_0_493", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_493", "role": "negated_conjecture", "formula": "vertical_4_4|horizontal_4_4|vertical_3_5|~vertical_4_3", "inference_rule": "spm", "raw_label": "cnf(c_0_493, negated_conjecture, (vertical_4_4|horizontal_4_4|vertical_3_5|~vertical_4_3),\\ninference(spm,[status(thm)],[c_0_467, c_0_226])).\n", "fillcolor": "lightskyblue1"}, {"id": 494, "label": "c_0_494", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_494", "role": "negated_conjecture", "formula": "~horizontal_4_3|~vertical_3_4", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_494, negated_conjecture, (~horizontal_4_3|~vertical_3_4),\\ninference(fof_simplification,[status(thm)],[unique_cover6_4_4])).\n", "fillcolor": "lightskyblue1"}, {"id": 495, "label": "c_0_495", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_495", "role": "negated_conjecture", "formula": "~horizontal_5_3|~vertical_3_2|~vertical_2_4", "inference_rule": "csr", "raw_label": "cnf(c_0_495, negated_conjecture, (~horizontal_5_3|~vertical_3_2|~vertical_2_4),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_430, c_0_304]), c_0_204]), c_0_170])).\n", "fillcolor": "lightskyblue1"}, {"id": 496, "label": "c_0_496", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_496", "role": "negated_conjecture", "formula": "~horizontal_5_3|~horizontal_3_4|~vertical_3_2", "inference_rule": "csr", "raw_label": "cnf(c_0_496, negated_conjecture, (~horizontal_5_3|~horizontal_3_4|~vertical_3_2),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_430, c_0_359]), c_0_204]), c_0_170])).\n", "fillcolor": "lightskyblue1"}, {"id": 497, "label": "c_0_497", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_497", "role": "negated_conjecture", "formula": "~vertical_3_5|~horizontal_3_4", "inference_rule": null, "raw_label": "cnf(c_0_497, negated_conjecture, (~vertical_3_5|~horizontal_3_4),\\nc_0_468).\n", "fillcolor": "lightskyblue1"}, {"id": 498, "label": "c_0_498", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_498", "role": "negated_conjecture", "formula": "vertical_3_3|horizontal_3_3|~horizontal_2_2", "inference_rule": "spm", "raw_label": "cnf(c_0_498, negated_conjecture, (vertical_3_3|horizontal_3_3|~horizontal_2_2),\\ninference(spm,[status(thm)],[c_0_469, c_0_452])).\n", "fillcolor": "lightskyblue1"}, {"id": 499, "label": "c_0_499", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_499", "role": "negated_conjecture", "formula": "horizontal_3_2|~horizontal_2_2", "inference_rule": "spm", "raw_label": "cnf(c_0_499, negated_conjecture, (horizontal_3_2|~horizontal_2_2),\\ninference(spm,[status(thm)],[c_0_470, c_0_323])).\n", "fillcolor": "lightskyblue1"}, {"id": 500, "label": "c_0_500", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_500", "role": "negated_conjecture", "formula": "~vertical_3_3|~vertical_2_3", "inference_rule": null, "raw_label": "cnf(c_0_500, negated_conjecture, (~vertical_3_3|~vertical_2_3),\\nc_0_471).\n", "fillcolor": "lightskyblue1"}, {"id": 501, "label": "c_0_501", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_501", "role": "negated_conjecture", "formula": "vertical_3_3|vertical_2_5|~vertical_3_2|~vertical_2_4", "inference_rule": "spm", "raw_label": "cnf(c_0_501, negated_conjecture, (vertical_3_3|vertical_2_5|~vertical_3_2|~vertical_2_4),\\ninference(spm,[status(thm)],[c_0_472, c_0_473])).\n", "fillcolor": "lightskyblue1"}, {"id": 502, "label": "c_0_502", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_502", "role": "negated_conjecture", "formula": "vertical_2_5|vertical_2_3|~vertical_2_4", "inference_rule": "csr", "raw_label": "cnf(c_0_502, negated_conjecture, (vertical_2_5|vertical_2_3|~vertical_2_4),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_474, c_0_475]), c_0_407]), c_0_476])).\n", "fillcolor": "lightskyblue1"}, {"id": 503, "label": "c_0_503", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_503", "role": "negated_conjecture", "formula": "vertical_3_2|horizontal_3_2|~vertical_1_2", "inference_rule": "csr", "raw_label": "cnf(c_0_503, negated_conjecture, (vertical_3_2|horizontal_3_2|~vertical_1_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_378, c_0_309]), c_0_396])).\n", "fillcolor": "lightskyblue1"}, {"id": 504, "label": "c_0_504", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_504", "role": "negated_conjecture", "formula": "horizontal_5_5|horizontal_3_5|vertical_1_4|~vertical_1_3", "inference_rule": "spm", "raw_label": "cnf(c_0_504, negated_conjecture, (horizontal_5_5|horizontal_3_5|vertical_1_4|~vertical_1_3),\\ninference(spm,[status(thm)],[c_0_477, c_0_138])).\n", "fillcolor": "lightskyblue1"}, {"id": 505, "label": "c_0_505", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_505", "role": "negated_conjecture", "formula": "~horizontal_3_1|~vertical_3_1", "inference_rule": null, "raw_label": "cnf(c_0_505, negated_conjecture, (~horizontal_3_1|~vertical_3_1),\\nc_0_478).\n", "fillcolor": "lightskyblue1"}, {"id": 506, "label": "c_0_506", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_506", "role": "negated_conjecture", "formula": "vertical_3_3|~horizontal_3_4|~vertical_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_506, negated_conjecture, (vertical_3_3|~horizontal_3_4|~vertical_2_2),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_479, c_0_245]), c_0_480]), c_0_481])).\n", "fillcolor": "lightskyblue1"}, {"id": 507, "label": "c_0_507", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_507", "role": "negated_conjecture", "formula": "horizontal_3_4|horizontal_3_3|vertical_2_4|~vertical_2_2", "inference_rule": "spm", "raw_label": "cnf(c_0_507, negated_conjecture, (horizontal_3_4|horizontal_3_3|vertical_2_4|~vertical_2_2),\\ninference(spm,[status(thm)],[c_0_482, c_0_291])).\n", "fillcolor": "lightskyblue1"}, {"id": 508, "label": "c_0_508", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_508", "role": "negated_conjecture", "formula": "vertical_3_3|~vertical_2_4|~vertical_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_508, negated_conjecture, (vertical_3_3|~vertical_2_4|~vertical_2_2),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_483, c_0_245]), c_0_484]), c_0_485])).\n", "fillcolor": "lightskyblue1"}, {"id": 509, "label": "c_0_509", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_509", "role": "negated_conjecture", "formula": "~horizontal_3_3|~vertical_2_3", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_509, negated_conjecture, (~horizontal_3_3|~vertical_2_3),\\ninference(fof_simplification,[status(thm)],[unique_cover3_3_3])).\n", "fillcolor": "lightskyblue1"}, {"id": 510, "label": "c_0_510", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_510", "role": "negated_conjecture", "formula": "vertical_1_2|~vertical_1_3", "inference_rule": "spm", "raw_label": "cnf(c_0_510, negated_conjecture, (vertical_1_2|~vertical_1_3),\\ninference(spm,[status(thm)],[c_0_486, c_0_219])).\n", "fillcolor": "lightskyblue1"}, {"id": 511, "label": "c_0_511", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_511", "role": "negated_conjecture", "formula": "vertical_3_3|~horizontal_2_3|~vertical_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_511, negated_conjecture, (vertical_3_3|~horizontal_2_3|~vertical_2_2),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_487, c_0_245]), c_0_488]), c_0_489])).\n", "fillcolor": "lightskyblue1"}, {"id": 512, "label": "c_0_512", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_512", "role": "negated_conjecture", "formula": "~vertical_3_4|~vertical_2_5|~vertical_2_2", "inference_rule": "spm", "raw_label": "cnf(c_0_512, negated_conjecture, (~vertical_3_4|~vertical_2_5|~vertical_2_2),\\ninference(spm,[status(thm)],[c_0_274, c_0_460])).\n", "fillcolor": "lightskyblue1"}, {"id": 513, "label": "c_0_513", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_513", "role": "negated_conjecture", "formula": "~horizontal_5_3|~vertical_3_2|~horizontal_2_3", "inference_rule": "csr", "raw_label": "cnf(c_0_513, negated_conjecture, (~horizontal_5_3|~vertical_3_2|~horizontal_2_3),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_430, c_0_381]), c_0_204]), c_0_170])).\n", "fillcolor": "lightskyblue1"}, {"id": 514, "label": "c_0_514", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_514", "role": "negated_conjecture", "formula": "vertical_4_4|~vertical_3_5|~horizontal_2_3", "inference_rule": "csr", "raw_label": "cnf(c_0_514, negated_conjecture, (vertical_4_4|~vertical_3_5|~horizontal_2_3),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_490, c_0_360]), c_0_356])).\n", "fillcolor": "lightskyblue1"}, {"id": 515, "label": "c_0_515", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_515", "role": "negated_conjecture", "formula": "vertical_4_2|horizontal_2_4|~vertical_3_4", "inference_rule": "csr", "raw_label": "cnf(c_0_515, negated_conjecture, (vertical_4_2|horizontal_2_4|~vertical_3_4),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_247, c_0_491]), c_0_263])).\n", "fillcolor": "lightskyblue1"}, {"id": 516, "label": "c_0_516", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_516", "role": "negated_conjecture", "formula": "~vertical_4_3|~vertical_3_5|~vertical_3_2", "inference_rule": "csr", "raw_label": "cnf(c_0_516, negated_conjecture, (~vertical_4_3|~vertical_3_5|~vertical_3_2),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_430, c_0_492]), c_0_204]), c_0_228])).\n", "fillcolor": "lightskyblue1"}, {"id": 517, "label": "c_0_517", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_517", "role": "negated_conjecture", "formula": "vertical_3_5|~vertical_4_3|~vertical_3_4", "inference_rule": "csr", "raw_label": "cnf(c_0_517, negated_conjecture, (vertical_3_5|~vertical_4_3|~vertical_3_4),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_247, c_0_493]), c_0_249])).\n", "fillcolor": "lightskyblue1"}, {"id": 518, "label": "c_0_518", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_518", "role": "negated_conjecture", "formula": "vertical_3_5|~horizontal_4_2|~vertical_3_4", "inference_rule": "spm", "raw_label": "cnf(c_0_518, negated_conjecture, (vertical_3_5|~horizontal_4_2|~vertical_3_4),\\ninference(spm,[status(thm)],[c_0_327, c_0_263])).\n", "fillcolor": "lightskyblue1"}, {"id": 519, "label": "c_0_519", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_519", "role": "negated_conjecture", "formula": "~horizontal_4_3|~vertical_3_4", "inference_rule": null, "raw_label": "cnf(c_0_519, negated_conjecture, (~horizontal_4_3|~vertical_3_4),\\nc_0_494).\n", "fillcolor": "lightskyblue1"}, {"id": 520, "label": "c_0_520", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_520", "role": "negated_conjecture", "formula": "vertical_4_4|horizontal_4_4|~vertical_3_2|~vertical_2_4", "inference_rule": "csr", "raw_label": "cnf(c_0_520, negated_conjecture, (vertical_4_4|horizontal_4_4|~vertical_3_2|~vertical_2_4),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_495, c_0_226]), c_0_429])).\n", "fillcolor": "lightskyblue1"}, {"id": 521, "label": "c_0_521", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_521", "role": "negated_conjecture", "formula": "vertical_4_4|horizontal_4_4|~horizontal_3_4|~vertical_3_2", "inference_rule": "csr", "raw_label": "cnf(c_0_521, negated_conjecture, (vertical_4_4|horizontal_4_4|~horizontal_3_4|~vertical_3_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_496, c_0_226]), c_0_497])).\n", "fillcolor": "lightskyblue1"}, {"id": 522, "label": "c_0_522", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_522", "role": "negated_conjecture", "formula": "~vertical_4_3|~horizontal_3_4|~vertical_3_2", "inference_rule": "csr", "raw_label": "cnf(c_0_522, negated_conjecture, (~vertical_4_3|~horizontal_3_4|~vertical_3_2),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_430, c_0_399]), c_0_204]), c_0_228])).\n", "fillcolor": "lightskyblue1"}, {"id": 523, "label": "c_0_523", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_523", "role": "negated_conjecture", "formula": "horizontal_3_3|~horizontal_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_523, negated_conjecture, (horizontal_3_3|~horizontal_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_390, c_0_498]), c_0_499])).\n", "fillcolor": "lightskyblue1"}, {"id": 524, "label": "c_0_524", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_524", "role": "negated_conjecture", "formula": "~horizontal_3_2|~vertical_2_3", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_524, negated_conjecture, (~horizontal_3_2|~vertical_2_3),\\ninference(fof_simplification,[status(thm)],[unique_cover6_3_3])).\n", "fillcolor": "lightskyblue1"}, {"id": 525, "label": "c_0_525", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_525", "role": "negated_conjecture", "formula": "vertical_2_5|~vertical_3_2|~vertical_2_4", "inference_rule": "csr", "raw_label": "cnf(c_0_525, negated_conjecture, (vertical_2_5|~vertical_3_2|~vertical_2_4),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_500, c_0_501]), c_0_502])).\n", "fillcolor": "lightskyblue1"}, {"id": 526, "label": "c_0_526", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_526", "role": "negated_conjecture", "formula": "vertical_3_2|horizontal_3_2|vertical_2_5|~vertical_2_4", "inference_rule": "spm", "raw_label": "cnf(c_0_526, negated_conjecture, (vertical_3_2|horizontal_3_2|vertical_2_5|~vertical_2_4),\\ninference(spm,[status(thm)],[c_0_503, c_0_280])).\n", "fillcolor": "lightskyblue1"}, {"id": 527, "label": "c_0_527", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_527", "role": "negated_conjecture", "formula": "horizontal_5_5|vertical_1_4|~vertical_2_5|~vertical_1_3", "inference_rule": "spm", "raw_label": "cnf(c_0_527, negated_conjecture, (horizontal_5_5|vertical_1_4|~vertical_2_5|~vertical_1_3),\\ninference(spm,[status(thm)],[c_0_234, c_0_504])).\n", "fillcolor": "lightskyblue1"}, {"id": 528, "label": "c_0_528", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_528", "role": "negated_conjecture", "formula": "~horizontal_4_5|~vertical_3_6", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_528, negated_conjecture, (~horizontal_4_5|~vertical_3_6),\\ninference(fof_simplification,[status(thm)],[unique_cover6_4_6])).\n", "fillcolor": "lightskyblue1"}, {"id": 529, "label": "c_0_529", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_529", "role": "negated_conjecture", "formula": "vertical_5_3|horizontal_4_1|~horizontal_6_4|~horizontal_3_1", "inference_rule": "spm", "raw_label": "cnf(c_0_529, negated_conjecture, (vertical_5_3|horizontal_4_1|~horizontal_6_4|~horizontal_3_1),\\ninference(spm,[status(thm)],[c_0_505, c_0_118])).\n", "fillcolor": "lightskyblue1"}, {"id": 530, "label": "c_0_530", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_530", "role": "negated_conjecture", "formula": "vertical_3_3|horizontal_3_3|~vertical_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_530, negated_conjecture, (vertical_3_3|horizontal_3_3|~vertical_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_506, c_0_507]), c_0_508])).\n", "fillcolor": "lightskyblue1"}, {"id": 531, "label": "c_0_531", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_531", "role": "negated_conjecture", "formula": "~horizontal_3_3|~vertical_2_3", "inference_rule": null, "raw_label": "cnf(c_0_531, negated_conjecture, (~horizontal_3_3|~vertical_2_3),\\nc_0_509).\n", "fillcolor": "lightskyblue1"}, {"id": 532, "label": "c_0_532", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_532", "role": "negated_conjecture", "formula": "vertical_2_3|horizontal_2_3|horizontal_2_2|vertical_1_2", "inference_rule": "spm", "raw_label": "cnf(c_0_532, negated_conjecture, (vertical_2_3|horizontal_2_3|horizontal_2_2|vertical_1_2),\\ninference(spm,[status(thm)],[c_0_510, c_0_475])).\n", "fillcolor": "lightskyblue1"}, {"id": 533, "label": "c_0_533", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_533", "role": "negated_conjecture", "formula": "~horizontal_3_3|~horizontal_2_3|~vertical_2_2", "inference_rule": "spm", "raw_label": "cnf(c_0_533, negated_conjecture, (~horizontal_3_3|~horizontal_2_3|~vertical_2_2),\\ninference(spm,[status(thm)],[c_0_288, c_0_511])).\n", "fillcolor": "lightskyblue1"}, {"id": 534, "label": "c_0_534", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_534", "role": "negated_conjecture", "formula": "horizontal_3_3|vertical_2_4|~vertical_2_5|~vertical_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_534, negated_conjecture, (horizontal_3_3|vertical_2_4|~vertical_2_5|~vertical_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_512, c_0_291]), c_0_292])).\n", "fillcolor": "lightskyblue1"}, {"id": 535, "label": "c_0_535", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_535", "role": "negated_conjecture", "formula": "vertical_4_4|horizontal_4_4|~vertical_3_2|~horizontal_2_3", "inference_rule": "csr", "raw_label": "cnf(c_0_535, negated_conjecture, (vertical_4_4|horizontal_4_4|~vertical_3_2|~horizontal_2_3),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_513, c_0_226]), c_0_514])).\n", "fillcolor": "lightskyblue1"}, {"id": 536, "label": "c_0_536", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_536", "role": "negated_conjecture", "formula": "~vertical_4_3|~vertical_3_2|~horizontal_2_3", "inference_rule": "csr", "raw_label": "cnf(c_0_536, negated_conjecture, (~vertical_4_3|~vertical_3_2|~horizontal_2_3),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_430, c_0_443]), c_0_204]), c_0_228])).\n", "fillcolor": "lightskyblue1"}, {"id": 537, "label": "c_0_537", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_537", "role": "negated_conjecture", "formula": "horizontal_2_4|~horizontal_3_4|~vertical_3_2|~vertical_1_2", "inference_rule": "spm", "raw_label": "cnf(c_0_537, negated_conjecture, (horizontal_2_4|~horizontal_3_4|~vertical_3_2|~vertical_1_2),\\ninference(spm,[status(thm)],[c_0_378, c_0_404])).\n", "fillcolor": "lightskyblue1"}, {"id": 538, "label": "c_0_538", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_538", "role": "negated_conjecture", "formula": "horizontal_2_4|~vertical_3_4|~vertical_3_2", "inference_rule": "spm", "raw_label": "cnf(c_0_538, negated_conjecture, (horizontal_2_4|~vertical_3_4|~vertical_3_2),\\ninference(spm,[status(thm)],[c_0_204, c_0_515])).\n", "fillcolor": "lightskyblue1"}, {"id": 539, "label": "c_0_539", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_539", "role": "negated_conjecture", "formula": "horizontal_4_3|vertical_3_3|~vertical_3_5|~vertical_3_2", "inference_rule": "csr", "raw_label": "cnf(c_0_539, negated_conjecture, (horizontal_4_3|vertical_3_3|~vertical_3_5|~vertical_3_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_516, c_0_245]), c_0_246])).\n", "fillcolor": "lightskyblue1"}, {"id": 540, "label": "c_0_540", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_540", "role": "negated_conjecture", "formula": "vertical_3_5|vertical_3_3|~vertical_3_4", "inference_rule": "csr", "raw_label": "cnf(c_0_540, negated_conjecture, (vertical_3_5|vertical_3_3|~vertical_3_4),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_517, c_0_245]), c_0_518]), c_0_519])).\n", "fillcolor": "lightskyblue1"}, {"id": 541, "label": "c_0_541", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_541", "role": "negated_conjecture", "formula": "~horizontal_4_3|~vertical_3_2|~vertical_2_4", "inference_rule": "csr", "raw_label": "cnf(c_0_541, negated_conjecture, (~horizontal_4_3|~vertical_3_2|~vertical_2_4),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_241, c_0_520]), c_0_243])).\n", "fillcolor": "lightskyblue1"}, {"id": 542, "label": "c_0_542", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_542", "role": "negated_conjecture", "formula": "~horizontal_4_3|~horizontal_3_4|~vertical_3_2", "inference_rule": "csr", "raw_label": "cnf(c_0_542, negated_conjecture, (~horizontal_4_3|~horizontal_3_4|~vertical_3_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_241, c_0_521]), c_0_243])).\n", "fillcolor": "lightskyblue1"}, {"id": 543, "label": "c_0_543", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_543", "role": "negated_conjecture", "formula": "horizontal_4_3|vertical_3_3|~horizontal_3_4|~vertical_3_2", "inference_rule": "csr", "raw_label": "cnf(c_0_543, negated_conjecture, (horizontal_4_3|vertical_3_3|~horizontal_3_4|~vertical_3_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_522, c_0_245]), c_0_246])).\n", "fillcolor": "lightskyblue1"}, {"id": 544, "label": "c_0_544", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_544", "role": "negated_conjecture", "formula": "~horizontal_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_544, negated_conjecture, (~horizontal_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_336, c_0_523]), c_0_499])).\n", "fillcolor": "lightskyblue1"}, {"id": 545, "label": "c_0_545", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_545", "role": "negated_conjecture", "formula": "~horizontal_3_2|~vertical_2_3", "inference_rule": null, "raw_label": "cnf(c_0_545, negated_conjecture, (~horizontal_3_2|~vertical_2_3),\\nc_0_524).\n", "fillcolor": "lightskyblue1"}, {"id": 546, "label": "c_0_546", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_546", "role": "negated_conjecture", "formula": "horizontal_3_2|vertical_2_5|~vertical_2_4", "inference_rule": "spm", "raw_label": "cnf(c_0_546, negated_conjecture, (horizontal_3_2|vertical_2_5|~vertical_2_4),\\ninference(spm,[status(thm)],[c_0_525, c_0_526])).\n", "fillcolor": "lightskyblue1"}, {"id": 547, "label": "c_0_547", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_547", "role": "negated_conjecture", "formula": "horizontal_5_5|~vertical_2_5|~vertical_2_4|~vertical_1_3", "inference_rule": "spm", "raw_label": "cnf(c_0_547, negated_conjecture, (horizontal_5_5|~vertical_2_5|~vertical_2_4|~vertical_1_3),\\ninference(spm,[status(thm)],[c_0_250, c_0_527])).\n", "fillcolor": "lightskyblue1"}, {"id": 548, "label": "c_0_548", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_548", "role": "negated_conjecture", "formula": "~horizontal_6_3|~vertical_5_4", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_548, negated_conjecture, (~horizontal_6_3|~vertical_5_4),\\ninference(fof_simplification,[status(thm)],[unique_cover6_6_4])).\n", "fillcolor": "lightskyblue1"}, {"id": 549, "label": "c_0_549", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_549", "role": "negated_conjecture", "formula": "~horizontal_4_5|~vertical_3_6", "inference_rule": null, "raw_label": "cnf(c_0_549, negated_conjecture, (~horizontal_4_5|~vertical_3_6),\\nc_0_528).\n", "fillcolor": "lightskyblue1"}, {"id": 550, "label": "c_0_550", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_550", "role": "negated_conjecture", "formula": "horizontal_2_4|~horizontal_5_3|~horizontal_3_4|~vertical_1_2", "inference_rule": "spm", "raw_label": "cnf(c_0_550, negated_conjecture, (horizontal_2_4|~horizontal_5_3|~horizontal_3_4|~vertical_1_2),\\ninference(spm,[status(thm)],[c_0_378, c_0_358])).\n", "fillcolor": "lightskyblue1"}, {"id": 551, "label": "c_0_551", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_551", "role": "negated_conjecture", "formula": "vertical_5_5|vertical_5_3|horizontal_4_1|~horizontal_3_1", "inference_rule": "spm", "raw_label": "cnf(c_0_551, negated_conjecture, (vertical_5_5|vertical_5_3|horizontal_4_1|~horizontal_3_1),\\ninference(spm,[status(thm)],[c_0_529, c_0_134])).\n", "fillcolor": "lightskyblue1"}, {"id": 552, "label": "c_0_552", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_552", "role": "negated_conjecture", "formula": "~vertical_2_3|~vertical_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_552, negated_conjecture, (~vertical_2_3|~vertical_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_500, c_0_530]), c_0_531])).\n", "fillcolor": "lightskyblue1"}, {"id": 553, "label": "c_0_553", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_553", "role": "negated_conjecture", "formula": "vertical_2_3|horizontal_2_3|~vertical_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_553, negated_conjecture, (vertical_2_3|horizontal_2_3|~vertical_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_396, c_0_532]), c_0_310])).\n", "fillcolor": "lightskyblue1"}, {"id": 554, "label": "c_0_554", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_554", "role": "negated_conjecture", "formula": "~horizontal_2_3|~vertical_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_554, negated_conjecture, (~horizontal_2_3|~vertical_2_2),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_533, c_0_534]), c_0_446]), c_0_476])).\n", "fillcolor": "lightskyblue1"}, {"id": 555, "label": "c_0_555", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_555", "role": "negated_conjecture", "formula": "~horizontal_4_3|~vertical_3_2|~horizontal_2_3", "inference_rule": "csr", "raw_label": "cnf(c_0_555, negated_conjecture, (~horizontal_4_3|~vertical_3_2|~horizontal_2_3),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_241, c_0_535]), c_0_243])).\n", "fillcolor": "lightskyblue1"}, {"id": 556, "label": "c_0_556", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_556", "role": "negated_conjecture", "formula": "horizontal_4_3|vertical_3_3|~vertical_3_2|~horizontal_2_3", "inference_rule": "csr", "raw_label": "cnf(c_0_556, negated_conjecture, (horizontal_4_3|vertical_3_3|~vertical_3_2|~horizontal_2_3),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_536, c_0_245]), c_0_246])).\n", "fillcolor": "lightskyblue1"}, {"id": 557, "label": "c_0_557", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_557", "role": "negated_conjecture", "formula": "~horizontal_3_4|~vertical_3_2|~horizontal_2_3", "inference_rule": "csr", "raw_label": "cnf(c_0_557, negated_conjecture, (~horizontal_3_4|~vertical_3_2|~horizontal_2_3),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_537, c_0_367]), c_0_356]), c_0_292])).\n", "fillcolor": "lightskyblue1"}, {"id": 558, "label": "c_0_558", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_558", "role": "negated_conjecture", "formula": "horizontal_3_4|horizontal_3_3|vertical_2_4|horizontal_2_4|~vertical_3_2", "inference_rule": "spm", "raw_label": "cnf(c_0_558, negated_conjecture, (horizontal_3_4|horizontal_3_3|vertical_2_4|horizontal_2_4|~vertical_3_2),\\ninference(spm,[status(thm)],[c_0_538, c_0_291])).\n", "fillcolor": "lightskyblue1"}, {"id": 559, "label": "c_0_559", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_559", "role": "negated_conjecture", "formula": "vertical_3_3|~vertical_3_4|~vertical_3_2", "inference_rule": "csr", "raw_label": "cnf(c_0_559, negated_conjecture, (vertical_3_3|~vertical_3_4|~vertical_3_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_519, c_0_539]), c_0_540])).\n", "fillcolor": "lightskyblue1"}, {"id": 560, "label": "c_0_560", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_560", "role": "negated_conjecture", "formula": "vertical_3_3|~vertical_3_2|~vertical_2_4", "inference_rule": "spm", "raw_label": "cnf(c_0_560, negated_conjecture, (vertical_3_3|~vertical_3_2|~vertical_2_4),\\ninference(spm,[status(thm)],[c_0_541, c_0_473])).\n", "fillcolor": "lightskyblue1"}, {"id": 561, "label": "c_0_561", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_561", "role": "negated_conjecture", "formula": "vertical_3_3|~horizontal_3_4|~vertical_3_2", "inference_rule": "spm", "raw_label": "cnf(c_0_561, negated_conjecture, (vertical_3_3|~horizontal_3_4|~vertical_3_2),\\ninference(spm,[status(thm)],[c_0_542, c_0_543])).\n", "fillcolor": "lightskyblue1"}, {"id": 562, "label": "c_0_562", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_562", "role": "negated_conjecture", "formula": "vertical_2_3|horizontal_2_3|vertical_1_3", "inference_rule": "sr", "raw_label": "cnf(c_0_562, negated_conjecture, (vertical_2_3|horizontal_2_3|vertical_1_3),\\ninference(sr,[status(thm)],[c_0_475, c_0_544])).\n", "fillcolor": "lightskyblue1"}, {"id": 563, "label": "c_0_563", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_563", "role": "negated_conjecture", "formula": "vertical_2_5|~vertical_2_4", "inference_rule": "csr", "raw_label": "cnf(c_0_563, negated_conjecture, (vertical_2_5|~vertical_2_4),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_545, c_0_546]), c_0_502])).\n", "fillcolor": "lightskyblue1"}, {"id": 564, "label": "c_0_564", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_564", "role": "negated_conjecture", "formula": "horizontal_5_5|vertical_2_3|horizontal_2_2|~vertical_2_4", "inference_rule": "csr", "raw_label": "cnf(c_0_564, negated_conjecture, (horizontal_5_5|vertical_2_3|horizontal_2_2|~vertical_2_4),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_547, c_0_475]), c_0_502]), c_0_476])).\n", "fillcolor": "lightskyblue1"}, {"id": 565, "label": "c_0_565", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_565", "role": "negated_conjecture", "formula": "~horizontal_6_3|~vertical_5_4", "inference_rule": null, "raw_label": "cnf(c_0_565, negated_conjecture, (~horizontal_6_3|~vertical_5_4),\\nc_0_548).\n", "fillcolor": "lightskyblue1"}, {"id": 566, "label": "c_0_566", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_566", "role": "negated_conjecture", "formula": "horizontal_3_5|vertical_2_6|~horizontal_4_5", "inference_rule": "spm", "raw_label": "cnf(c_0_566, negated_conjecture, (horizontal_3_5|vertical_2_6|~horizontal_4_5),\\ninference(spm,[status(thm)],[c_0_549, c_0_147])).\n", "fillcolor": "lightskyblue1"}, {"id": 567, "label": "c_0_567", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_567", "role": "negated_conjecture", "formula": "~horizontal_5_3|~horizontal_3_4|~horizontal_2_3", "inference_rule": "csr", "raw_label": "cnf(c_0_567, negated_conjecture, (~horizontal_5_3|~horizontal_3_4|~horizontal_2_3),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_550, c_0_367]), c_0_356]), c_0_292])).\n", "fillcolor": "lightskyblue1"}, {"id": 568, "label": "c_0_568", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_568", "role": "negated_conjecture", "formula": "vertical_5_5|horizontal_4_1|~horizontal_5_3|~horizontal_3_1", "inference_rule": "spm", "raw_label": "cnf(c_0_568, negated_conjecture, (vertical_5_5|horizontal_4_1|~horizontal_5_3|~horizontal_3_1),\\ninference(spm,[status(thm)],[c_0_154, c_0_551])).\n", "fillcolor": "lightskyblue1"}, {"id": 569, "label": "c_0_569", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_569", "role": "negated_conjecture", "formula": "~vertical_2_2", "inference_rule": "csr", "raw_label": "cnf(c_0_569, negated_conjecture, (~vertical_2_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_552, c_0_553]), c_0_554])).\n", "fillcolor": "lightskyblue1"}, {"id": 570, "label": "c_0_570", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_570", "role": "negated_conjecture", "formula": "horizontal_4_4|vertical_3_5|horizontal_2_1|~vertical_3_2", "inference_rule": "csr", "raw_label": "cnf(c_0_570, negated_conjecture, (horizontal_4_4|vertical_3_5|horizontal_2_1|~vertical_3_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_180, c_0_279]), c_0_216])).\n", "fillcolor": "lightskyblue1"}, {"id": 571, "label": "c_0_571", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_571", "role": "negated_conjecture", "formula": "vertical_3_3|~vertical_3_2|~horizontal_2_3", "inference_rule": "spm", "raw_label": "cnf(c_0_571, negated_conjecture, (vertical_3_3|~vertical_3_2|~horizontal_2_3),\\ninference(spm,[status(thm)],[c_0_555, c_0_556])).\n", "fillcolor": "lightskyblue1"}, {"id": 572, "label": "c_0_572", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_572", "role": "negated_conjecture", "formula": "horizontal_3_3|~vertical_3_2|~horizontal_2_3", "inference_rule": "csr", "raw_label": "cnf(c_0_572, negated_conjecture, (horizontal_3_3|~vertical_3_2|~horizontal_2_3),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_557, c_0_558]), c_0_356]), c_0_476])).\n", "fillcolor": "lightskyblue1"}, {"id": 573, "label": "c_0_573", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_573", "role": "negated_conjecture", "formula": "vertical_3_3|horizontal_3_3|~vertical_3_2", "inference_rule": "csr", "raw_label": "cnf(c_0_573, negated_conjecture, (vertical_3_3|horizontal_3_3|~vertical_3_2),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_559, c_0_291]), c_0_560]), c_0_561])).\n", "fillcolor": "lightskyblue1"}, {"id": 574, "label": "c_0_574", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_574", "role": "negated_conjecture", "formula": "horizontal_5_5|vertical_2_3|~vertical_2_4", "inference_rule": "csr", "raw_label": "cnf(c_0_574, negated_conjecture, (horizontal_5_5|vertical_2_3|~vertical_2_4),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_547, c_0_562]), c_0_563]), c_0_476])).\n", "fillcolor": "lightskyblue1"}, {"id": 575, "label": "c_0_575", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_575", "role": "negated_conjecture", "formula": "vertical_2_3|horizontal_2_2|~vertical_5_5|~vertical_2_4", "inference_rule": "spm", "raw_label": "cnf(c_0_575, negated_conjecture, (vertical_2_3|horizontal_2_2|~vertical_5_5|~vertical_2_4),\\ninference(spm,[status(thm)],[c_0_164, c_0_564])).\n", "fillcolor": "lightskyblue1"}, {"id": 576, "label": "c_0_576", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_576", "role": "negated_conjecture", "formula": "horizontal_6_2|vertical_5_3|~vertical_5_4", "inference_rule": "spm", "raw_label": "cnf(c_0_576, negated_conjecture, (horizontal_6_2|vertical_5_3|~vertical_5_4),\\ninference(spm,[status(thm)],[c_0_565, c_0_103])).\n", "fillcolor": "lightskyblue1"}, {"id": 577, "label": "c_0_577", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_577", "role": "negated_conjecture", "formula": "horizontal_3_5|~horizontal_4_5|~horizontal_1_4", "inference_rule": "spm", "raw_label": "cnf(c_0_577, negated_conjecture, (horizontal_3_5|~horizontal_4_5|~horizontal_1_4),\\ninference(spm,[status(thm)],[c_0_160, c_0_566])).\n", "fillcolor": "lightskyblue1"}, {"id": 578, "label": "c_0_578", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_578", "role": "negated_conjecture", "formula": "horizontal_4_4|vertical_3_5|vertical_2_5|~horizontal_5_3|~horizontal_2_3", "inference_rule": "spm", "raw_label": "cnf(c_0_578, negated_conjecture, (horizontal_4_4|vertical_3_5|vertical_2_5|~horizontal_5_3|~horizontal_2_3),\\ninference(spm,[status(thm)],[c_0_409, c_0_367])).\n", "fillcolor": "lightskyblue1"}, {"id": 579, "label": "c_0_579", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_579", "role": "negated_conjecture", "formula": "horizontal_2_4|~vertical_4_3|~horizontal_3_4|~vertical_1_2", "inference_rule": "spm", "raw_label": "cnf(c_0_579, negated_conjecture, (horizontal_2_4|~vertical_4_3|~horizontal_3_4|~vertical_1_2),\\ninference(spm,[status(thm)],[c_0_378, c_0_398])).\n", "fillcolor": "lightskyblue1"}, {"id": 580, "label": "c_0_580", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_580", "role": "negated_conjecture", "formula": "horizontal_2_4|~horizontal_4_2|~horizontal_3_4|~vertical_1_2", "inference_rule": "spm", "raw_label": "cnf(c_0_580, negated_conjecture, (horizontal_2_4|~horizontal_4_2|~horizontal_3_4|~vertical_1_2),\\ninference(spm,[status(thm)],[c_0_378, c_0_400])).\n", "fillcolor": "lightskyblue1"}, {"id": 581, "label": "c_0_581", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_581", "role": "negated_conjecture", "formula": "vertical_4_4|~horizontal_3_4|~horizontal_2_3", "inference_rule": "csr", "raw_label": "cnf(c_0_581, negated_conjecture, (vertical_4_4|~horizontal_3_4|~horizontal_2_3),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_567, c_0_368]), c_0_356])).\n", "fillcolor": "lightskyblue1"}, {"id": 582, "label": "c_0_582", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_582", "role": "negated_conjecture", "formula": "vertical_5_5|~horizontal_5_3|~vertical_4_2|~horizontal_3_1", "inference_rule": "spm", "raw_label": "cnf(c_0_582, negated_conjecture, (vertical_5_5|~horizontal_5_3|~vertical_4_2|~horizontal_3_1),\\ninference(spm,[status(thm)],[c_0_178, c_0_568])).\n", "fillcolor": "lightskyblue1"}, {"id": 583, "label": "c_0_583", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_583", "role": "negated_conjecture", "formula": "vertical_3_2|horizontal_3_2|horizontal_3_1", "inference_rule": "sr", "raw_label": "cnf(c_0_583, negated_conjecture, (vertical_3_2|horizontal_3_2|horizontal_3_1),\\ninference(sr,[status(thm)],[c_0_296, c_0_569])).\n", "fillcolor": "lightskyblue1"}, {"id": 584, "label": "c_0_584", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_584", "role": "negated_conjecture", "formula": "horizontal_4_4|vertical_3_5|~vertical_3_2|~vertical_1_2", "inference_rule": "spm", "raw_label": "cnf(c_0_584, negated_conjecture, (horizontal_4_4|vertical_3_5|~vertical_3_2|~vertical_1_2),\\ninference(spm,[status(thm)],[c_0_378, c_0_570])).\n", "fillcolor": "lightskyblue1"}, {"id": 585, "label": "c_0_585", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_585", "role": "negated_conjecture", "formula": "vertical_2_3|horizontal_2_3|vertical_1_2", "inference_rule": "sr", "raw_label": "cnf(c_0_585, negated_conjecture, (vertical_2_3|horizontal_2_3|vertical_1_2),\\ninference(sr,[status(thm)],[c_0_532, c_0_544])).\n", "fillcolor": "lightskyblue1"}, {"id": 586, "label": "c_0_586", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_586", "role": "negated_conjecture", "formula": "~vertical_3_2|~horizontal_2_3", "inference_rule": "csr", "raw_label": "cnf(c_0_586, negated_conjecture, (~vertical_3_2|~horizontal_2_3),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_288, c_0_571]), c_0_572])).\n", "fillcolor": "lightskyblue1"}, {"id": 587, "label": "c_0_587", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_587", "role": "negated_conjecture", "formula": "~vertical_3_2|~vertical_2_3", "inference_rule": "csr", "raw_label": "cnf(c_0_587, negated_conjecture, (~vertical_3_2|~vertical_2_3),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_500, c_0_573]), c_0_531])).\n", "fillcolor": "lightskyblue1"}, {"id": 588, "label": "c_0_588", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_588", "role": "negated_conjecture", "formula": "vertical_2_3|~vertical_5_5|~vertical_2_4", "inference_rule": "spm", "raw_label": "cnf(c_0_588, negated_conjecture, (vertical_2_3|~vertical_5_5|~vertical_2_4),\\ninference(spm,[status(thm)],[c_0_164, c_0_574])).\n", "fillcolor": "lightskyblue1"}, {"id": 589, "label": "c_0_589", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_589", "role": "negated_conjecture", "formula": "vertical_4_2|vertical_2_3|~horizontal_5_3|~vertical_2_4", "inference_rule": "csr", "raw_label": "cnf(c_0_589, negated_conjecture, (vertical_4_2|vertical_2_3|~horizontal_5_3|~vertical_2_4),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_575, c_0_181]), c_0_335])).\n", "fillcolor": "lightskyblue1"}, {"id": 590, "label": "c_0_590", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_590", "role": "negated_conjecture", "formula": "vertical_5_3|~vertical_5_4|~vertical_4_1", "inference_rule": "spm", "raw_label": "cnf(c_0_590, negated_conjecture, (vertical_5_3|~vertical_5_4|~vertical_4_1),\\ninference(spm,[status(thm)],[c_0_104, c_0_576])).\n", "fillcolor": "lightskyblue1"}, {"id": 591, "label": "c_0_591", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_591", "role": "negated_conjecture", "formula": "~vertical_5_5|~horizontal_5_4", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_591, negated_conjecture, (~vertical_5_5|~horizontal_5_4),\\ninference(fof_simplification,[status(thm)],[unique_cover4_5_5])).\n", "fillcolor": "lightskyblue1"}, {"id": 592, "label": "c_0_592", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_592", "role": "negated_conjecture", "formula": "horizontal_3_5|vertical_1_4|horizontal_1_3|~horizontal_4_5", "inference_rule": "spm", "raw_label": "cnf(c_0_592, negated_conjecture, (horizontal_3_5|vertical_1_4|horizontal_1_3|~horizontal_4_5),\\ninference(spm,[status(thm)],[c_0_577, c_0_176])).\n", "fillcolor": "lightskyblue1"}, {"id": 593, "label": "c_0_593", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_593", "role": "negated_conjecture", "formula": "vertical_4_4|horizontal_4_4|vertical_2_5|~horizontal_2_3", "inference_rule": "csr", "raw_label": "cnf(c_0_593, negated_conjecture, (vertical_4_4|horizontal_4_4|vertical_2_5|~horizontal_2_3),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_578, c_0_226]), c_0_514])).\n", "fillcolor": "lightskyblue1"}, {"id": 594, "label": "c_0_594", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_594", "role": "negated_conjecture", "formula": "~vertical_3_5|~vertical_3_4|~horizontal_2_3", "inference_rule": "spm", "raw_label": "cnf(c_0_594, negated_conjecture, (~vertical_3_5|~vertical_3_4|~horizontal_2_3),\\ninference(spm,[status(thm)],[c_0_247, c_0_514])).\n", "fillcolor": "lightskyblue1"}, {"id": 595, "label": "c_0_595", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_595", "role": "negated_conjecture", "formula": "~vertical_4_3|~horizontal_3_4|~horizontal_2_3", "inference_rule": "csr", "raw_label": "cnf(c_0_595, negated_conjecture, (~vertical_4_3|~horizontal_3_4|~horizontal_2_3),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_579, c_0_367]), c_0_356]), c_0_292])).\n", "fillcolor": "lightskyblue1"}, {"id": 596, "label": "c_0_596", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_596", "role": "negated_conjecture", "formula": "~horizontal_4_2|~horizontal_3_4|~horizontal_2_3", "inference_rule": "csr", "raw_label": "cnf(c_0_596, negated_conjecture, (~horizontal_4_2|~horizontal_3_4|~horizontal_2_3),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_580, c_0_367]), c_0_356]), c_0_292])).\n", "fillcolor": "lightskyblue1"}, {"id": 597, "label": "c_0_597", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_597", "role": "negated_conjecture", "formula": "~horizontal_4_3|~horizontal_3_4|~horizontal_2_3", "inference_rule": "spm", "raw_label": "cnf(c_0_597, negated_conjecture, (~horizontal_4_3|~horizontal_3_4|~horizontal_2_3),\\ninference(spm,[status(thm)],[c_0_241, c_0_581])).\n", "fillcolor": "lightskyblue1"}, {"id": 598, "label": "c_0_598", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_598", "role": "negated_conjecture", "formula": "vertical_5_3|~vertical_5_4|~vertical_5_2", "inference_rule": "spm", "raw_label": "cnf(c_0_598, negated_conjecture, (vertical_5_3|~vertical_5_4|~vertical_5_2),\\ninference(spm,[status(thm)],[c_0_125, c_0_576])).\n", "fillcolor": "lightskyblue1"}, {"id": 599, "label": "c_0_599", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_599", "role": "negated_conjecture", "formula": "vertical_5_3|~vertical_5_4|~horizontal_5_1", "inference_rule": "spm", "raw_label": "cnf(c_0_599, negated_conjecture, (vertical_5_3|~vertical_5_4|~horizontal_5_1),\\ninference(spm,[status(thm)],[c_0_126, c_0_576])).\n", "fillcolor": "lightskyblue1"}, {"id": 600, "label": "c_0_600", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_600", "role": "negated_conjecture", "formula": "vertical_5_5|vertical_3_2|horizontal_3_2|~horizontal_5_3", "inference_rule": "csr", "raw_label": "cnf(c_0_600, negated_conjecture, (vertical_5_5|vertical_3_2|horizontal_3_2|~horizontal_5_3),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_582, c_0_583]), c_0_181])).\n", "fillcolor": "lightskyblue1"}, {"id": 601, "label": "c_0_601", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_601", "role": "negated_conjecture", "formula": "horizontal_4_4|vertical_3_5|~vertical_3_2", "inference_rule": "csr", "raw_label": "cnf(c_0_601, negated_conjecture, (horizontal_4_4|vertical_3_5|~vertical_3_2),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_584, c_0_585]), c_0_586]), c_0_587])).\n", "fillcolor": "lightskyblue1"}, {"id": 602, "label": "c_0_602", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_602", "role": "negated_conjecture", "formula": "vertical_2_3|horizontal_2_1|~horizontal_5_3|~vertical_2_4", "inference_rule": "csr", "raw_label": "cnf(c_0_602, negated_conjecture, (vertical_2_3|horizontal_2_1|~horizontal_5_3|~vertical_2_4),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_588, c_0_196]), c_0_589])).\n", "fillcolor": "lightskyblue1"}, {"id": 603, "label": "c_0_603", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_603", "role": "negated_conjecture", "formula": "vertical_2_3|~horizontal_5_4|~vertical_2_4", "inference_rule": "spm", "raw_label": "cnf(c_0_603, negated_conjecture, (vertical_2_3|~horizontal_5_4|~vertical_2_4),\\ninference(spm,[status(thm)],[c_0_200, c_0_574])).\n", "fillcolor": "lightskyblue1"}, {"id": 604, "label": "c_0_604", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_604", "role": "negated_conjecture", "formula": "vertical_5_3|horizontal_4_1|vertical_3_1|~vertical_5_4", "inference_rule": "spm", "raw_label": "cnf(c_0_604, negated_conjecture, (vertical_5_3|horizontal_4_1|vertical_3_1|~vertical_5_4),\\ninference(spm,[status(thm)],[c_0_590, c_0_110])).\n", "fillcolor": "lightskyblue1"}, {"id": 605, "label": "c_0_605", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_605", "role": "negated_conjecture", "formula": "~vertical_5_5|~horizontal_5_4", "inference_rule": null, "raw_label": "cnf(c_0_605, negated_conjecture, (~vertical_5_5|~horizontal_5_4),\\nc_0_591).\n", "fillcolor": "lightskyblue1"}, {"id": 606, "label": "c_0_606", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_606", "role": "negated_conjecture", "formula": "horizontal_3_5|vertical_1_4|~horizontal_4_5|~horizontal_1_2", "inference_rule": "spm", "raw_label": "cnf(c_0_606, negated_conjecture, (horizontal_3_5|vertical_1_4|~horizontal_4_5|~horizontal_1_2),\\ninference(spm,[status(thm)],[c_0_189, c_0_592])).\n", "fillcolor": "lightskyblue1"}, {"id": 607, "label": "c_0_607", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_607", "role": "negated_conjecture", "formula": "vertical_2_5|~horizontal_4_3|~horizontal_2_3", "inference_rule": "csr", "raw_label": "cnf(c_0_607, negated_conjecture, (vertical_2_5|~horizontal_4_3|~horizontal_2_3),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_241, c_0_593]), c_0_243])).\n", "fillcolor": "lightskyblue1"}, {"id": 608, "label": "c_0_608", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_608", "role": "negated_conjecture", "formula": "vertical_3_3|~vertical_3_4|~horizontal_2_3", "inference_rule": "spm", "raw_label": "cnf(c_0_608, negated_conjecture, (vertical_3_3|~vertical_3_4|~horizontal_2_3),\\ninference(spm,[status(thm)],[c_0_594, c_0_540])).\n", "fillcolor": "lightskyblue1"}, {"id": 609, "label": "c_0_609", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_609", "role": "negated_conjecture", "formula": "vertical_3_3|~horizontal_3_4|~horizontal_2_3", "inference_rule": "csr", "raw_label": "cnf(c_0_609, negated_conjecture, (vertical_3_3|~horizontal_3_4|~horizontal_2_3),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_595, c_0_245]), c_0_596]), c_0_597])).\n", "fillcolor": "lightskyblue1"}, {"id": 610, "label": "c_0_610", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_610", "role": "negated_conjecture", "formula": "~horizontal_5_5|~vertical_4_5", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_610, negated_conjecture, (~horizontal_5_5|~vertical_4_5),\\ninference(fof_simplification,[status(thm)],[unique_cover3_5_5])).\n", "fillcolor": "lightskyblue1"}, {"id": 611, "label": "c_0_611", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_611", "role": "negated_conjecture", "formula": "horizontal_5_4|vertical_5_3|horizontal_5_3|vertical_4_4|~vertical_5_2", "inference_rule": "spm", "raw_label": "cnf(c_0_611, negated_conjecture, (horizontal_5_4|vertical_5_3|horizontal_5_3|vertical_4_4|~vertical_5_2),\\ninference(spm,[status(thm)],[c_0_598, c_0_199])).\n", "fillcolor": "lightskyblue1"}, {"id": 612, "label": "c_0_612", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_612", "role": "negated_conjecture", "formula": "horizontal_5_4|vertical_5_3|horizontal_5_3|vertical_4_4|~horizontal_5_1", "inference_rule": "spm", "raw_label": "cnf(c_0_612, negated_conjecture, (horizontal_5_4|vertical_5_3|horizontal_5_3|vertical_4_4|~horizontal_5_1),\\ninference(spm,[status(thm)],[c_0_599, c_0_199])).\n", "fillcolor": "lightskyblue1"}, {"id": 613, "label": "c_0_613", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_613", "role": "negated_conjecture", "formula": "horizontal_4_4|vertical_3_5|horizontal_3_2|~horizontal_5_3", "inference_rule": "csr", "raw_label": "cnf(c_0_613, negated_conjecture, (horizontal_4_4|vertical_3_5|horizontal_3_2|~horizontal_5_3),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_180, c_0_600]), c_0_601])).\n", "fillcolor": "lightskyblue1"}, {"id": 614, "label": "c_0_614", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_614", "role": "negated_conjecture", "formula": "vertical_4_2|vertical_2_3|~vertical_4_3|~vertical_2_4", "inference_rule": "csr", "raw_label": "cnf(c_0_614, negated_conjecture, (vertical_4_2|vertical_2_3|~vertical_4_3|~vertical_2_4),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_575, c_0_325]), c_0_363])).\n", "fillcolor": "lightskyblue1"}, {"id": 615, "label": "c_0_615", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_615", "role": "negated_conjecture", "formula": "vertical_2_3|~horizontal_5_3|~vertical_2_4|~vertical_1_2", "inference_rule": "spm", "raw_label": "cnf(c_0_615, negated_conjecture, (vertical_2_3|~horizontal_5_3|~vertical_2_4|~vertical_1_2),\\ninference(spm,[status(thm)],[c_0_378, c_0_602])).\n", "fillcolor": "lightskyblue1"}, {"id": 616, "label": "c_0_616", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_616", "role": "negated_conjecture", "formula": "~horizontal_4_4|~vertical_4_4", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_616, negated_conjecture, (~horizontal_4_4|~vertical_4_4),\\ninference(fof_simplification,[status(thm)],[unique_cover1_4_4])).\n", "fillcolor": "lightskyblue1"}, {"id": 617, "label": "c_0_617", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_617", "role": "negated_conjecture", "formula": "horizontal_5_3|vertical_4_4|vertical_2_3|~vertical_2_4", "inference_rule": "csr", "raw_label": "cnf(c_0_617, negated_conjecture, (horizontal_5_3|vertical_4_4|vertical_2_3|~vertical_2_4),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_588, c_0_213]), c_0_603])).\n", "fillcolor": "lightskyblue1"}, {"id": 618, "label": "c_0_618", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_618", "role": "negated_conjecture", "formula": "vertical_5_3|horizontal_4_1|~vertical_5_4|~horizontal_3_1", "inference_rule": "spm", "raw_label": "cnf(c_0_618, negated_conjecture, (vertical_5_3|horizontal_4_1|~vertical_5_4|~horizontal_3_1),\\ninference(spm,[status(thm)],[c_0_505, c_0_604])).\n", "fillcolor": "lightskyblue1"}, {"id": 619, "label": "c_0_619", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_619", "role": "negated_conjecture", "formula": "vertical_5_5|horizontal_4_1|~vertical_4_3|~horizontal_3_1", "inference_rule": "spm", "raw_label": "cnf(c_0_619, negated_conjecture, (vertical_5_5|horizontal_4_1|~vertical_4_3|~horizontal_3_1),\\ninference(spm,[status(thm)],[c_0_298, c_0_551])).\n", "fillcolor": "lightskyblue1"}, {"id": 620, "label": "c_0_620", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_620", "role": "negated_conjecture", "formula": "horizontal_2_1|~horizontal_5_4|~horizontal_5_2|~vertical_3_2", "inference_rule": "spm", "raw_label": "cnf(c_0_620, negated_conjecture, (horizontal_2_1|~horizontal_5_4|~horizontal_5_2|~vertical_3_2),\\ninference(spm,[status(thm)],[c_0_605, c_0_279])).\n", "fillcolor": "lightskyblue1"}, {"id": 621, "label": "c_0_621", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_621", "role": "negated_conjecture", "formula": "horizontal_3_5|vertical_1_4|vertical_1_2|~horizontal_4_5", "inference_rule": "spm", "raw_label": "cnf(c_0_621, negated_conjecture, (horizontal_3_5|vertical_1_4|vertical_1_2|~horizontal_4_5),\\ninference(spm,[status(thm)],[c_0_606, c_0_219])).\n", "fillcolor": "lightskyblue1"}, {"id": 622, "label": "c_0_622", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_622", "role": "negated_conjecture", "formula": "vertical_3_3|vertical_2_5|~vertical_3_2|~horizontal_2_3", "inference_rule": "spm", "raw_label": "cnf(c_0_622, negated_conjecture, (vertical_3_3|vertical_2_5|~vertical_3_2|~horizontal_2_3),\\ninference(spm,[status(thm)],[c_0_607, c_0_556])).\n", "fillcolor": "lightskyblue1"}, {"id": 623, "label": "c_0_623", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_623", "role": "negated_conjecture", "formula": "vertical_3_3|horizontal_3_3|~horizontal_2_3", "inference_rule": "csr", "raw_label": "cnf(c_0_623, negated_conjecture, (vertical_3_3|horizontal_3_3|~horizontal_2_3),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_608, c_0_291]), c_0_476]), c_0_609])).\n", "fillcolor": "lightskyblue1"}, {"id": 624, "label": "c_0_624", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_624", "role": "negated_conjecture", "formula": "horizontal_4_1|~vertical_5_2|~horizontal_3_1", "inference_rule": "spm", "raw_label": "cnf(c_0_624, negated_conjecture, (horizontal_4_1|~vertical_5_2|~horizontal_3_1),\\ninference(spm,[status(thm)],[c_0_505, c_0_158])).\n", "fillcolor": "lightskyblue1"}, {"id": 625, "label": "c_0_625", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_625", "role": "negated_conjecture", "formula": "horizontal_4_1|~horizontal_5_1|~horizontal_3_1", "inference_rule": "spm", "raw_label": "cnf(c_0_625, negated_conjecture, (horizontal_4_1|~horizontal_5_1|~horizontal_3_1),\\ninference(spm,[status(thm)],[c_0_505, c_0_159])).\n", "fillcolor": "lightskyblue1"}, {"id": 626, "label": "c_0_626", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_626", "role": "negated_conjecture", "formula": "horizontal_2_1|~horizontal_5_3|~vertical_4_5|~vertical_4_2", "inference_rule": "spm", "raw_label": "cnf(c_0_626, negated_conjecture, (horizontal_2_1|~horizontal_5_3|~vertical_4_5|~vertical_4_2),\\ninference(spm,[status(thm)],[c_0_166, c_0_196])).\n", "fillcolor": "lightskyblue1"}, {"id": 627, "label": "c_0_627", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_627", "role": "negated_conjecture", "formula": "vertical_4_2|~horizontal_5_3|~vertical_4_5", "inference_rule": "spm", "raw_label": "cnf(c_0_627, negated_conjecture, (vertical_4_2|~horizontal_5_3|~vertical_4_5),\\ninference(spm,[status(thm)],[c_0_166, c_0_181])).\n", "fillcolor": "lightskyblue1"}, {"id": 628, "label": "c_0_628", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_628", "role": "negated_conjecture", "formula": "~horizontal_5_5|~vertical_4_5", "inference_rule": null, "raw_label": "cnf(c_0_628, negated_conjecture, (~horizontal_5_5|~vertical_4_5),\\nc_0_610).\n", "fillcolor": "lightskyblue1"}, {"id": 629, "label": "c_0_629", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_629", "role": "negated_conjecture", "formula": "horizontal_5_4|vertical_4_4|~vertical_5_2|~vertical_4_3", "inference_rule": "csr", "raw_label": "cnf(c_0_629, negated_conjecture, (horizontal_5_4|vertical_4_4|~vertical_5_2|~vertical_4_3),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_298, c_0_611]), c_0_467])).\n", "fillcolor": "lightskyblue1"}, {"id": 630, "label": "c_0_630", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_630", "role": "negated_conjecture", "formula": "horizontal_5_4|vertical_4_4|~horizontal_5_1|~vertical_4_3", "inference_rule": "csr", "raw_label": "cnf(c_0_630, negated_conjecture, (horizontal_5_4|vertical_4_4|~horizontal_5_1|~vertical_4_3),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_298, c_0_612]), c_0_467])).\n", "fillcolor": "lightskyblue1"}, {"id": 631, "label": "c_0_631", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_631", "role": "negated_conjecture", "formula": "vertical_4_2|~horizontal_5_4|~vertical_4_3", "inference_rule": "spm", "raw_label": "cnf(c_0_631, negated_conjecture, (vertical_4_2|~horizontal_5_4|~vertical_4_3),\\ninference(spm,[status(thm)],[c_0_605, c_0_325])).\n", "fillcolor": "lightskyblue1"}, {"id": 632, "label": "c_0_632", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_632", "role": "negated_conjecture", "formula": "vertical_4_4|horizontal_4_4|vertical_3_5|horizontal_3_2", "inference_rule": "spm", "raw_label": "cnf(c_0_632, negated_conjecture, (vertical_4_4|horizontal_4_4|vertical_3_5|horizontal_3_2),\\ninference(spm,[status(thm)],[c_0_613, c_0_226])).\n", "fillcolor": "lightskyblue1"}, {"id": 633, "label": "c_0_633", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_633", "role": "negated_conjecture", "formula": "vertical_2_3|horizontal_2_1|~vertical_4_3|~vertical_2_4", "inference_rule": "csr", "raw_label": "cnf(c_0_633, negated_conjecture, (vertical_2_3|horizontal_2_1|~vertical_4_3|~vertical_2_4),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_588, c_0_339]), c_0_614])).\n", "fillcolor": "lightskyblue1"}, {"id": 634, "label": "c_0_634", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_634", "role": "negated_conjecture", "formula": "vertical_2_3|horizontal_2_1|~horizontal_4_2|~vertical_2_4", "inference_rule": "csr", "raw_label": "cnf(c_0_634, negated_conjecture, (vertical_2_3|horizontal_2_1|~horizontal_4_2|~vertical_2_4),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_588, c_0_341]), c_0_342])).\n", "fillcolor": "lightskyblue1"}, {"id": 635, "label": "c_0_635", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_635", "role": "negated_conjecture", "formula": "vertical_2_3|~horizontal_5_3|~vertical_2_4", "inference_rule": "csr", "raw_label": "cnf(c_0_635, negated_conjecture, (vertical_2_3|~horizontal_5_3|~vertical_2_4),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_615, c_0_585]), c_0_476])).\n", "fillcolor": "lightskyblue1"}, {"id": 636, "label": "c_0_636", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_636", "role": "negated_conjecture", "formula": "~horizontal_4_4|~vertical_4_4", "inference_rule": null, "raw_label": "cnf(c_0_636, negated_conjecture, (~horizontal_4_4|~vertical_4_4),\\nc_0_616).\n", "fillcolor": "lightskyblue1"}, {"id": 637, "label": "c_0_637", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_637", "role": "negated_conjecture", "formula": "vertical_4_4|~vertical_3_2|~vertical_2_4", "inference_rule": "csr", "raw_label": "cnf(c_0_637, negated_conjecture, (vertical_4_4|~vertical_3_2|~vertical_2_4),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_495, c_0_617]), c_0_587])).\n", "fillcolor": "lightskyblue1"}, {"id": 638, "label": "c_0_638", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_638", "role": "negated_conjecture", "formula": "horizontal_2_4|~vertical_3_5|~vertical_3_2|~vertical_1_2", "inference_rule": "spm", "raw_label": "cnf(c_0_638, negated_conjecture, (horizontal_2_4|~vertical_3_5|~vertical_3_2|~vertical_1_2),\\ninference(spm,[status(thm)],[c_0_378, c_0_377])).\n", "fillcolor": "lightskyblue1"}, {"id": 639, "label": "c_0_639", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_639", "role": "negated_conjecture", "formula": "horizontal_5_4|vertical_5_3|horizontal_5_3|vertical_4_4|horizontal_4_1|~horizontal_3_1", "inference_rule": "spm", "raw_label": "cnf(c_0_639, negated_conjecture, (horizontal_5_4|vertical_5_3|horizontal_5_3|vertical_4_4|horizontal_4_1|~horizontal_3_1),\\ninference(spm,[status(thm)],[c_0_618, c_0_199])).\n", "fillcolor": "lightskyblue1"}, {"id": 640, "label": "c_0_640", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_640", "role": "negated_conjecture", "formula": "vertical_5_5|~vertical_4_3|~vertical_4_2|~horizontal_3_1", "inference_rule": "spm", "raw_label": "cnf(c_0_640, negated_conjecture, (vertical_5_5|~vertical_4_3|~vertical_4_2|~horizontal_3_1),\\ninference(spm,[status(thm)],[c_0_178, c_0_619])).\n", "fillcolor": "lightskyblue1"}, {"id": 641, "label": "c_0_641", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_641", "role": "negated_conjecture", "formula": "~horizontal_5_4|~horizontal_5_2|~vertical_3_2|~vertical_1_2", "inference_rule": "spm", "raw_label": "cnf(c_0_641, negated_conjecture, (~horizontal_5_4|~horizontal_5_2|~vertical_3_2|~vertical_1_2),\\ninference(spm,[status(thm)],[c_0_378, c_0_620])).\n", "fillcolor": "lightskyblue1"}, {"id": 642, "label": "c_0_642", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_642", "role": "negated_conjecture", "formula": "horizontal_5_2|vertical_2_3|horizontal_2_3|~vertical_3_2", "inference_rule": "csr", "raw_label": "cnf(c_0_642, negated_conjecture, (horizontal_5_2|vertical_2_3|horizontal_2_3|~vertical_3_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_430, c_0_532]), c_0_229])).\n", "fillcolor": "lightskyblue1"}, {"id": 643, "label": "c_0_643", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_643", "role": "negated_conjecture", "formula": "vertical_4_5|horizontal_4_4|vertical_3_5|horizontal_3_5|vertical_1_4|vertical_1_2", "inference_rule": "spm", "raw_label": "cnf(c_0_643, negated_conjecture, (vertical_4_5|horizontal_4_4|vertical_3_5|horizontal_3_5|vertical_1_4|vertical_1_2),\\ninference(spm,[status(thm)],[c_0_621, c_0_152])).\n", "fillcolor": "lightskyblue1"}, {"id": 644, "label": "c_0_644", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_644", "role": "negated_conjecture", "formula": "vertical_2_5|~vertical_3_2|~horizontal_2_3", "inference_rule": "csr", "raw_label": "cnf(c_0_644, negated_conjecture, (vertical_2_5|~vertical_3_2|~horizontal_2_3),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_288, c_0_622]), c_0_572])).\n", "fillcolor": "lightskyblue1"}, {"id": 645, "label": "c_0_645", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_645", "role": "negated_conjecture", "formula": "vertical_3_2|horizontal_3_2|vertical_2_5|~horizontal_2_3", "inference_rule": "spm", "raw_label": "cnf(c_0_645, negated_conjecture, (vertical_3_2|horizontal_3_2|vertical_2_5|~horizontal_2_3),\\ninference(spm,[status(thm)],[c_0_503, c_0_367])).\n", "fillcolor": "lightskyblue1"}, {"id": 646, "label": "c_0_646", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_646", "role": "negated_conjecture", "formula": "~horizontal_3_2|~horizontal_2_3", "inference_rule": "csr", "raw_label": "cnf(c_0_646, negated_conjecture, (~horizontal_3_2|~horizontal_2_3),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_390, c_0_623]), c_0_336])).\n", "fillcolor": "lightskyblue1"}, {"id": 647, "label": "c_0_647", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_647", "role": "negated_conjecture", "formula": "horizontal_5_2|vertical_4_2|horizontal_4_1|~horizontal_3_1", "inference_rule": "csr", "raw_label": "cnf(c_0_647, negated_conjecture, (horizontal_5_2|vertical_4_2|horizontal_4_1|~horizontal_3_1),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_624, c_0_168]), c_0_625])).\n", "fillcolor": "lightskyblue1"}, {"id": 648, "label": "c_0_648", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_648", "role": "negated_conjecture", "formula": "~horizontal_5_3|~vertical_4_5|~vertical_1_2", "inference_rule": "csr", "raw_label": "cnf(c_0_648, negated_conjecture, (~horizontal_5_3|~vertical_4_5|~vertical_1_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_378, c_0_626]), c_0_627])).\n", "fillcolor": "lightskyblue1"}, {"id": 649, "label": "c_0_649", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_649", "role": "negated_conjecture", "formula": "vertical_1_2|~vertical_4_5|~vertical_2_5|~horizontal_2_3", "inference_rule": "spm", "raw_label": "cnf(c_0_649, negated_conjecture, (vertical_1_2|~vertical_4_5|~vertical_2_5|~horizontal_2_3),\\ninference(spm,[status(thm)],[c_0_628, c_0_355])).\n", "fillcolor": "lightskyblue1"}, {"id": 650, "label": "c_0_650", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_650", "role": "negated_conjecture", "formula": "vertical_4_4|vertical_4_2|~vertical_4_3", "inference_rule": "csr", "raw_label": "cnf(c_0_650, negated_conjecture, (vertical_4_4|vertical_4_2|~vertical_4_3),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_629, c_0_168]), c_0_630]), c_0_228]), c_0_631])).\n", "fillcolor": "lightskyblue1"}, {"id": 651, "label": "c_0_651", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_651", "role": "negated_conjecture", "formula": "horizontal_2_4|vertical_2_3|~horizontal_5_3|~vertical_3_5", "inference_rule": "csr", "raw_label": "cnf(c_0_651, negated_conjecture, (horizontal_2_4|vertical_2_3|~horizontal_5_3|~vertical_3_5),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_585]), c_0_490])).\n", "fillcolor": "lightskyblue1"}, {"id": 652, "label": "c_0_652", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_652", "role": "negated_conjecture", "formula": "horizontal_4_4|vertical_3_5|vertical_2_3|horizontal_2_3|~horizontal_5_3", "inference_rule": "spm", "raw_label": "cnf(c_0_652, negated_conjecture, (horizontal_4_4|vertical_3_5|vertical_2_3|horizontal_2_3|~horizontal_5_3),\\ninference(spm,[status(thm)],[c_0_409, c_0_585])).\n", "fillcolor": "lightskyblue1"}, {"id": 653, "label": "c_0_653", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_653", "role": "negated_conjecture", "formula": "vertical_3_5|horizontal_3_2|~vertical_3_4", "inference_rule": "csr", "raw_label": "cnf(c_0_653, negated_conjecture, (vertical_3_5|horizontal_3_2|~vertical_3_4),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_247, c_0_632]), c_0_249])).\n", "fillcolor": "lightskyblue1"}, {"id": 654, "label": "c_0_654", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_654", "role": "negated_conjecture", "formula": "vertical_2_3|~vertical_4_3|~vertical_2_4|~vertical_1_2", "inference_rule": "spm", "raw_label": "cnf(c_0_654, negated_conjecture, (vertical_2_3|~vertical_4_3|~vertical_2_4|~vertical_1_2),\\ninference(spm,[status(thm)],[c_0_378, c_0_633])).\n", "fillcolor": "lightskyblue1"}, {"id": 655, "label": "c_0_655", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_655", "role": "negated_conjecture", "formula": "vertical_2_3|~horizontal_4_2|~vertical_2_4|~vertical_1_2", "inference_rule": "spm", "raw_label": "cnf(c_0_655, negated_conjecture, (vertical_2_3|~horizontal_4_2|~vertical_2_4|~vertical_1_2),\\ninference(spm,[status(thm)],[c_0_378, c_0_634])).\n", "fillcolor": "lightskyblue1"}, {"id": 656, "label": "c_0_656", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_656", "role": "negated_conjecture", "formula": "vertical_4_4|vertical_2_3|~vertical_2_4", "inference_rule": "spm", "raw_label": "cnf(c_0_656, negated_conjecture, (vertical_4_4|vertical_2_3|~vertical_2_4),\\ninference(spm,[status(thm)],[c_0_635, c_0_617])).\n", "fillcolor": "lightskyblue1"}, {"id": 657, "label": "c_0_657", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_657", "role": "negated_conjecture", "formula": "~horizontal_4_4|~vertical_3_2|~vertical_2_4", "inference_rule": "spm", "raw_label": "cnf(c_0_657, negated_conjecture, (~horizontal_4_4|~vertical_3_2|~vertical_2_4),\\ninference(spm,[status(thm)],[c_0_636, c_0_637])).\n", "fillcolor": "lightskyblue1"}, {"id": 658, "label": "c_0_658", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_658", "role": "negated_conjecture", "formula": "~vertical_3_5|~vertical_3_2|~vertical_2_4", "inference_rule": "csr", "raw_label": "cnf(c_0_658, negated_conjecture, (~vertical_3_5|~vertical_3_2|~vertical_2_4),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_638, c_0_280]), c_0_267]), c_0_274])).\n", "fillcolor": "lightskyblue1"}, {"id": 659, "label": "c_0_659", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_659", "role": "negated_conjecture", "formula": "horizontal_3_5|~horizontal_4_5|~horizontal_2_5", "inference_rule": "spm", "raw_label": "cnf(c_0_659, negated_conjecture, (horizontal_3_5|~horizontal_4_5|~horizontal_2_5),\\ninference(spm,[status(thm)],[c_0_239, c_0_566])).\n", "fillcolor": "lightskyblue1"}, {"id": 660, "label": "c_0_660", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_660", "role": "negated_conjecture", "formula": "horizontal_3_5|~horizontal_4_5|~vertical_1_5", "inference_rule": "spm", "raw_label": "cnf(c_0_660, negated_conjecture, (horizontal_3_5|~horizontal_4_5|~vertical_1_5),\\ninference(spm,[status(thm)],[c_0_240, c_0_566])).\n", "fillcolor": "lightskyblue1"}, {"id": 661, "label": "c_0_661", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_661", "role": "negated_conjecture", "formula": "horizontal_5_4|vertical_4_4|horizontal_4_1|~vertical_4_3|~horizontal_3_1", "inference_rule": "csr", "raw_label": "cnf(c_0_661, negated_conjecture, (horizontal_5_4|vertical_4_4|horizontal_4_1|~vertical_4_3|~horizontal_3_1),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_298, c_0_639]), c_0_467])).\n", "fillcolor": "lightskyblue1"}, {"id": 662, "label": "c_0_662", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_662", "role": "negated_conjecture", "formula": "vertical_5_5|vertical_3_2|horizontal_3_2|~vertical_4_3", "inference_rule": "csr", "raw_label": "cnf(c_0_662, negated_conjecture, (vertical_5_5|vertical_3_2|horizontal_3_2|~vertical_4_3),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_640, c_0_583]), c_0_325])).\n", "fillcolor": "lightskyblue1"}, {"id": 663, "label": "c_0_663", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_663", "role": "negated_conjecture", "formula": "~horizontal_5_4|~vertical_3_2", "inference_rule": "csr", "raw_label": "cnf(c_0_663, negated_conjecture, (~horizontal_5_4|~vertical_3_2),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_641, c_0_585]), c_0_642]), c_0_586]), c_0_587])).\n", "fillcolor": "lightskyblue1"}, {"id": 664, "label": "c_0_664", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_664", "role": "negated_conjecture", "formula": "vertical_4_5|horizontal_4_4|vertical_1_4|vertical_1_2|~vertical_2_5", "inference_rule": "csr", "raw_label": "cnf(c_0_664, negated_conjecture, (vertical_4_5|horizontal_4_4|vertical_1_4|vertical_1_2|~vertical_2_5),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_234, c_0_643]), c_0_274])).\n", "fillcolor": "lightskyblue1"}, {"id": 665, "label": "c_0_665", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_665", "role": "negated_conjecture", "formula": "vertical_2_5|~horizontal_2_3", "inference_rule": "csr", "raw_label": "cnf(c_0_665, negated_conjecture, (vertical_2_5|~horizontal_2_3),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_644, c_0_645]), c_0_646])).\n", "fillcolor": "lightskyblue1"}, {"id": 666, "label": "c_0_666", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_666", "role": "negated_conjecture", "formula": "vertical_5_5|horizontal_4_1|~horizontal_5_2|~horizontal_3_1", "inference_rule": "spm", "raw_label": "cnf(c_0_666, negated_conjecture, (vertical_5_5|horizontal_4_1|~horizontal_5_2|~horizontal_3_1),\\ninference(spm,[status(thm)],[c_0_252, c_0_551])).\n", "fillcolor": "lightskyblue1"}, {"id": 667, "label": "c_0_667", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_667", "role": "negated_conjecture", "formula": "horizontal_5_2|~horizontal_4_2|~horizontal_3_1", "inference_rule": "csr", "raw_label": "cnf(c_0_667, negated_conjecture, (horizontal_5_2|~horizontal_4_2|~horizontal_3_1),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_326, c_0_647]), c_0_327])).\n", "fillcolor": "lightskyblue1"}, {"id": 668, "label": "c_0_668", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_668", "role": "negated_conjecture", "formula": "horizontal_5_2|~horizontal_4_2|~vertical_2_2", "inference_rule": "spm", "raw_label": "cnf(c_0_668, negated_conjecture, (horizontal_5_2|~horizontal_4_2|~vertical_2_2),\\ninference(spm,[status(thm)],[c_0_395, c_0_342])).\n", "fillcolor": "lightskyblue1"}, {"id": 669, "label": "c_0_669", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_669", "role": "negated_conjecture", "formula": "~horizontal_5_3|~vertical_4_5|~vertical_2_5|~horizontal_2_3", "inference_rule": "spm", "raw_label": "cnf(c_0_669, negated_conjecture, (~horizontal_5_3|~vertical_4_5|~vertical_2_5|~horizontal_2_3),\\ninference(spm,[status(thm)],[c_0_648, c_0_649])).\n", "fillcolor": "lightskyblue1"}, {"id": 670, "label": "c_0_670", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_670", "role": "negated_conjecture", "formula": "vertical_4_2|~horizontal_4_4|~vertical_4_3", "inference_rule": "spm", "raw_label": "cnf(c_0_670, negated_conjecture, (vertical_4_2|~horizontal_4_4|~vertical_4_3),\\ninference(spm,[status(thm)],[c_0_636, c_0_650])).\n", "fillcolor": "lightskyblue1"}, {"id": 671, "label": "c_0_671", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_671", "role": "negated_conjecture", "formula": "horizontal_4_4|vertical_4_2|vertical_3_5|~vertical_4_3", "inference_rule": "spm", "raw_label": "cnf(c_0_671, negated_conjecture, (horizontal_4_4|vertical_4_2|vertical_3_5|~vertical_4_3),\\ninference(spm,[status(thm)],[c_0_180, c_0_325])).\n", "fillcolor": "lightskyblue1"}, {"id": 672, "label": "c_0_672", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_672", "role": "negated_conjecture", "formula": "vertical_4_4|horizontal_2_4|vertical_2_3|~vertical_3_5", "inference_rule": "spm", "raw_label": "cnf(c_0_672, negated_conjecture, (vertical_4_4|horizontal_2_4|vertical_2_3|~vertical_3_5),\\ninference(spm,[status(thm)],[c_0_651, c_0_360])).\n", "fillcolor": "lightskyblue1"}, {"id": 673, "label": "c_0_673", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_673", "role": "negated_conjecture", "formula": "vertical_4_4|horizontal_4_4|vertical_3_5|vertical_2_3|horizontal_2_3", "inference_rule": "spm", "raw_label": "cnf(c_0_673, negated_conjecture, (vertical_4_4|horizontal_4_4|vertical_3_5|vertical_2_3|horizontal_2_3),\\ninference(spm,[status(thm)],[c_0_652, c_0_226])).\n", "fillcolor": "lightskyblue1"}, {"id": 674, "label": "c_0_674", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_674", "role": "negated_conjecture", "formula": "~vertical_3_4|~horizontal_2_3", "inference_rule": "csr", "raw_label": "cnf(c_0_674, negated_conjecture, (~vertical_3_4|~horizontal_2_3),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_594, c_0_653]), c_0_646])).\n", "fillcolor": "lightskyblue1"}, {"id": 675, "label": "c_0_675", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_675", "role": "negated_conjecture", "formula": "vertical_2_3|~vertical_4_3|~vertical_2_4", "inference_rule": "csr", "raw_label": "cnf(c_0_675, negated_conjecture, (vertical_2_3|~vertical_4_3|~vertical_2_4),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_654, c_0_585]), c_0_476])).\n", "fillcolor": "lightskyblue1"}, {"id": 676, "label": "c_0_676", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_676", "role": "negated_conjecture", "formula": "vertical_2_3|~horizontal_4_2|~vertical_2_4", "inference_rule": "csr", "raw_label": "cnf(c_0_676, negated_conjecture, (vertical_2_3|~horizontal_4_2|~vertical_2_4),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_655, c_0_585]), c_0_476])).\n", "fillcolor": "lightskyblue1"}, {"id": 677, "label": "c_0_677", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_677", "role": "negated_conjecture", "formula": "vertical_2_3|~horizontal_4_3|~vertical_2_4", "inference_rule": "spm", "raw_label": "cnf(c_0_677, negated_conjecture, (vertical_2_3|~horizontal_4_3|~vertical_2_4),\\ninference(spm,[status(thm)],[c_0_241, c_0_656])).\n", "fillcolor": "lightskyblue1"}, {"id": 678, "label": "c_0_678", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_678", "role": "negated_conjecture", "formula": "~vertical_3_2|~vertical_2_4", "inference_rule": "csr", "raw_label": "cnf(c_0_678, negated_conjecture, (~vertical_3_2|~vertical_2_4),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_657, c_0_601]), c_0_658])).\n", "fillcolor": "lightskyblue1"}, {"id": 679, "label": "c_0_679", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_679", "role": "negated_conjecture", "formula": "vertical_3_2|horizontal_3_2|vertical_2_3|horizontal_2_3", "inference_rule": "csr", "raw_label": "cnf(c_0_679, negated_conjecture, (vertical_3_2|horizontal_3_2|vertical_2_3|horizontal_2_3),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_503, c_0_532]), c_0_323])).\n", "fillcolor": "lightskyblue1"}, {"id": 680, "label": "c_0_680", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_680", "role": "negated_conjecture", "formula": "horizontal_2_4|vertical_2_3|~horizontal_5_3|~horizontal_3_4", "inference_rule": "csr", "raw_label": "cnf(c_0_680, negated_conjecture, (horizontal_2_4|vertical_2_3|~horizontal_5_3|~horizontal_3_4),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_550, c_0_585]), c_0_567])).\n", "fillcolor": "lightskyblue1"}, {"id": 681, "label": "c_0_681", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_681", "role": "negated_conjecture", "formula": "vertical_4_5|horizontal_4_4|vertical_3_5|horizontal_3_5|~horizontal_2_5", "inference_rule": "spm", "raw_label": "cnf(c_0_681, negated_conjecture, (vertical_4_5|horizontal_4_4|vertical_3_5|horizontal_3_5|~horizontal_2_5),\\ninference(spm,[status(thm)],[c_0_659, c_0_152])).\n", "fillcolor": "lightskyblue1"}, {"id": 682, "label": "c_0_682", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_682", "role": "negated_conjecture", "formula": "vertical_4_5|horizontal_4_4|vertical_3_5|horizontal_3_5|~vertical_1_5", "inference_rule": "spm", "raw_label": "cnf(c_0_682, negated_conjecture, (vertical_4_5|horizontal_4_4|vertical_3_5|horizontal_3_5|~vertical_1_5),\\ninference(spm,[status(thm)],[c_0_660, c_0_152])).\n", "fillcolor": "lightskyblue1"}, {"id": 683, "label": "c_0_683", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_683", "role": "negated_conjecture", "formula": "horizontal_5_4|vertical_4_4|~vertical_4_3|~horizontal_3_1", "inference_rule": "csr", "raw_label": "cnf(c_0_683, negated_conjecture, (horizontal_5_4|vertical_4_4|~vertical_4_3|~horizontal_3_1),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_178, c_0_661]), c_0_650])).\n", "fillcolor": "lightskyblue1"}, {"id": 684, "label": "c_0_684", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_684", "role": "negated_conjecture", "formula": "horizontal_3_2|~horizontal_5_4|~vertical_4_3", "inference_rule": "csr", "raw_label": "cnf(c_0_684, negated_conjecture, (horizontal_3_2|~horizontal_5_4|~vertical_4_3),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_605, c_0_662]), c_0_663])).\n", "fillcolor": "lightskyblue1"}, {"id": 685, "label": "c_0_685", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_685", "role": "negated_conjecture", "formula": "vertical_4_5|horizontal_4_4|vertical_1_2|~horizontal_2_3", "inference_rule": "csr", "raw_label": "cnf(c_0_685, negated_conjecture, (vertical_4_5|horizontal_4_4|vertical_1_2|~horizontal_2_3),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_343, c_0_664]), c_0_665])).\n", "fillcolor": "lightskyblue1"}, {"id": 686, "label": "c_0_686", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_686", "role": "negated_conjecture", "formula": "vertical_5_5|~horizontal_5_2|~horizontal_4_2|~horizontal_3_1", "inference_rule": "spm", "raw_label": "cnf(c_0_686, negated_conjecture, (vertical_5_5|~horizontal_5_2|~horizontal_4_2|~horizontal_3_1),\\ninference(spm,[status(thm)],[c_0_326, c_0_666])).\n", "fillcolor": "lightskyblue1"}, {"id": 687, "label": "c_0_687", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_687", "role": "negated_conjecture", "formula": "horizontal_5_2|horizontal_3_2|~horizontal_4_2", "inference_rule": "csr", "raw_label": "cnf(c_0_687, negated_conjecture, (horizontal_5_2|horizontal_3_2|~horizontal_4_2),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_667, c_0_296]), c_0_668]), c_0_246])).\n", "fillcolor": "lightskyblue1"}, {"id": 688, "label": "c_0_688", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_688", "role": "negated_conjecture", "formula": "~horizontal_5_3|~vertical_4_5|~horizontal_2_3", "inference_rule": "csr", "raw_label": "cnf(c_0_688, negated_conjecture, (~horizontal_5_3|~vertical_4_5|~horizontal_2_3),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_648, c_0_367]), c_0_669])).\n", "fillcolor": "lightskyblue1"}, {"id": 689, "label": "c_0_689", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_689", "role": "negated_conjecture", "formula": "horizontal_5_3|vertical_4_4|~vertical_4_5", "inference_rule": "csr", "raw_label": "cnf(c_0_689, negated_conjecture, (horizontal_5_3|vertical_4_4|~vertical_4_5),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_166, c_0_213]), c_0_201])).\n", "fillcolor": "lightskyblue1"}, {"id": 690, "label": "c_0_690", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_690", "role": "negated_conjecture", "formula": "vertical_4_2|vertical_3_5|~vertical_4_3", "inference_rule": "spm", "raw_label": "cnf(c_0_690, negated_conjecture, (vertical_4_2|vertical_3_5|~vertical_4_3),\\ninference(spm,[status(thm)],[c_0_670, c_0_671])).\n", "fillcolor": "lightskyblue1"}, {"id": 691, "label": "c_0_691", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_691", "role": "negated_conjecture", "formula": "vertical_4_2|vertical_3_5|~horizontal_4_3", "inference_rule": "csr", "raw_label": "cnf(c_0_691, negated_conjecture, (vertical_4_2|vertical_3_5|~horizontal_4_3),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_241, c_0_248]), c_0_243])).\n", "fillcolor": "lightskyblue1"}, {"id": 692, "label": "c_0_692", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_692", "role": "negated_conjecture", "formula": "vertical_4_4|horizontal_2_4|~horizontal_3_4|~vertical_3_2", "inference_rule": "spm", "raw_label": "cnf(c_0_692, negated_conjecture, (vertical_4_4|horizontal_2_4|~horizontal_3_4|~vertical_3_2),\\ninference(spm,[status(thm)],[c_0_496, c_0_368])).\n", "fillcolor": "lightskyblue1"}, {"id": 693, "label": "c_0_693", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_693", "role": "negated_conjecture", "formula": "horizontal_2_4|vertical_2_3|~vertical_3_5|~vertical_3_4", "inference_rule": "spm", "raw_label": "cnf(c_0_693, negated_conjecture, (horizontal_2_4|vertical_2_3|~vertical_3_5|~vertical_3_4),\\ninference(spm,[status(thm)],[c_0_247, c_0_672])).\n", "fillcolor": "lightskyblue1"}, {"id": 694, "label": "c_0_694", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_694", "role": "negated_conjecture", "formula": "vertical_3_5|vertical_2_3|~vertical_3_4", "inference_rule": "csr", "raw_label": "cnf(c_0_694, negated_conjecture, (vertical_3_5|vertical_2_3|~vertical_3_4),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_247, c_0_673]), c_0_674]), c_0_249])).\n", "fillcolor": "lightskyblue1"}, {"id": 695, "label": "c_0_695", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_695", "role": "negated_conjecture", "formula": "vertical_3_3|vertical_2_3|~vertical_2_4", "inference_rule": "csr", "raw_label": "cnf(c_0_695, negated_conjecture, (vertical_3_3|vertical_2_3|~vertical_2_4),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_675, c_0_245]), c_0_676]), c_0_677])).\n", "fillcolor": "lightskyblue1"}, {"id": 696, "label": "c_0_696", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_696", "role": "negated_conjecture", "formula": "horizontal_3_2|vertical_2_3|~vertical_2_4", "inference_rule": "csr", "raw_label": "cnf(c_0_696, negated_conjecture, (horizontal_3_2|vertical_2_3|~vertical_2_4),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_678, c_0_679]), c_0_476])).\n", "fillcolor": "lightskyblue1"}, {"id": 697, "label": "c_0_697", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_697", "role": "negated_conjecture", "formula": "vertical_4_4|horizontal_2_4|vertical_2_3|~horizontal_3_4", "inference_rule": "spm", "raw_label": "cnf(c_0_697, negated_conjecture, (vertical_4_4|horizontal_2_4|vertical_2_3|~horizontal_3_4),\\ninference(spm,[status(thm)],[c_0_680, c_0_368])).\n", "fillcolor": "lightskyblue1"}, {"id": 698, "label": "c_0_698", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_698", "role": "negated_conjecture", "formula": "vertical_4_5|horizontal_4_4|~horizontal_3_4|~horizontal_2_5", "inference_rule": "csr", "raw_label": "cnf(c_0_698, negated_conjecture, (vertical_4_5|horizontal_4_4|~horizontal_3_4|~horizontal_2_5),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_284, c_0_681]), c_0_497])).\n", "fillcolor": "lightskyblue1"}, {"id": 699, "label": "c_0_699", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_699", "role": "negated_conjecture", "formula": "vertical_4_5|horizontal_4_4|~horizontal_3_4|~vertical_1_5", "inference_rule": "csr", "raw_label": "cnf(c_0_699, negated_conjecture, (vertical_4_5|horizontal_4_4|~horizontal_3_4|~vertical_1_5),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_284, c_0_682]), c_0_497])).\n", "fillcolor": "lightskyblue1"}, {"id": 700, "label": "c_0_700", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_700", "role": "negated_conjecture", "formula": "horizontal_2_4|~vertical_4_5|~horizontal_3_4", "inference_rule": "spm", "raw_label": "cnf(c_0_700, negated_conjecture, (horizontal_2_4|~vertical_4_5|~horizontal_3_4),\\ninference(spm,[status(thm)],[c_0_628, c_0_315])).\n", "fillcolor": "lightskyblue1"}, {"id": 701, "label": "c_0_701", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_701", "role": "negated_conjecture", "formula": "vertical_4_4|vertical_3_2|horizontal_3_2|~vertical_4_3", "inference_rule": "csr", "raw_label": "cnf(c_0_701, negated_conjecture, (vertical_4_4|vertical_3_2|horizontal_3_2|~vertical_4_3),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_683, c_0_583]), c_0_684])).\n", "fillcolor": "lightskyblue1"}, {"id": 702, "label": "c_0_702", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_702", "role": "negated_conjecture", "formula": "vertical_3_2|horizontal_3_2|~vertical_4_5|~horizontal_2_3", "inference_rule": "csr", "raw_label": "cnf(c_0_702, negated_conjecture, (vertical_3_2|horizontal_3_2|~vertical_4_5|~horizontal_2_3),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_503, c_0_649]), c_0_645])).\n", "fillcolor": "lightskyblue1"}, {"id": 703, "label": "c_0_703", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_703", "role": "negated_conjecture", "formula": "vertical_4_5|horizontal_4_4|~horizontal_2_3", "inference_rule": "csr", "raw_label": "cnf(c_0_703, negated_conjecture, (vertical_4_5|horizontal_4_4|~horizontal_2_3),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_503, c_0_685]), c_0_646]), c_0_586])).\n", "fillcolor": "lightskyblue1"}, {"id": 704, "label": "c_0_704", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_704", "role": "negated_conjecture", "formula": "vertical_5_5|horizontal_3_2|~horizontal_4_2", "inference_rule": "csr", "raw_label": "cnf(c_0_704, negated_conjecture, (vertical_5_5|horizontal_3_2|~horizontal_4_2),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_686, c_0_583]), c_0_687]), c_0_246])).\n", "fillcolor": "lightskyblue1"}, {"id": 705, "label": "c_0_705", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_705", "role": "negated_conjecture", "formula": "vertical_4_4|~vertical_4_5|~horizontal_2_3", "inference_rule": "spm", "raw_label": "cnf(c_0_705, negated_conjecture, (vertical_4_4|~vertical_4_5|~horizontal_2_3),\\ninference(spm,[status(thm)],[c_0_688, c_0_689])).\n", "fillcolor": "lightskyblue1"}, {"id": 706, "label": "c_0_706", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_706", "role": "negated_conjecture", "formula": "vertical_4_2|horizontal_4_2|vertical_3_5|vertical_3_3", "inference_rule": "csr", "raw_label": "cnf(c_0_706, negated_conjecture, (vertical_4_2|horizontal_4_2|vertical_3_5|vertical_3_3),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_690, c_0_245]), c_0_691])).\n", "fillcolor": "lightskyblue1"}, {"id": 707, "label": "c_0_707", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_707", "role": "negated_conjecture", "formula": "horizontal_2_4|~horizontal_4_4|~horizontal_3_4|~vertical_3_2", "inference_rule": "spm", "raw_label": "cnf(c_0_707, negated_conjecture, (horizontal_2_4|~horizontal_4_4|~horizontal_3_4|~vertical_3_2),\\ninference(spm,[status(thm)],[c_0_636, c_0_692])).\n", "fillcolor": "lightskyblue1"}, {"id": 708, "label": "c_0_708", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_708", "role": "negated_conjecture", "formula": "horizontal_2_4|vertical_2_3|~vertical_3_4", "inference_rule": "spm", "raw_label": "cnf(c_0_708, negated_conjecture, (horizontal_2_4|vertical_2_3|~vertical_3_4),\\ninference(spm,[status(thm)],[c_0_693, c_0_694])).\n", "fillcolor": "lightskyblue1"}, {"id": 709, "label": "c_0_709", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_709", "role": "negated_conjecture", "formula": "vertical_2_3|~vertical_2_4", "inference_rule": "csr", "raw_label": "cnf(c_0_709, negated_conjecture, (vertical_2_3|~vertical_2_4),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_390, c_0_695]), c_0_696])).\n", "fillcolor": "lightskyblue1"}, {"id": 710, "label": "c_0_710", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_710", "role": "negated_conjecture", "formula": "horizontal_2_4|vertical_2_3|~horizontal_4_4|~horizontal_3_4", "inference_rule": "spm", "raw_label": "cnf(c_0_710, negated_conjecture, (horizontal_2_4|vertical_2_3|~horizontal_4_4|~horizontal_3_4),\\ninference(spm,[status(thm)],[c_0_636, c_0_697])).\n", "fillcolor": "lightskyblue1"}, {"id": 711, "label": "c_0_711", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_711", "role": "negated_conjecture", "formula": "horizontal_4_4|horizontal_2_4|~horizontal_3_4", "inference_rule": "csr", "raw_label": "cnf(c_0_711, negated_conjecture, (horizontal_4_4|horizontal_2_4|~horizontal_3_4),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_698, c_0_254]), c_0_699]), c_0_292]), c_0_700])).\n", "fillcolor": "lightskyblue1"}, {"id": 712, "label": "c_0_712", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_712", "role": "negated_conjecture", "formula": "vertical_3_2|horizontal_3_2|~horizontal_4_4|~vertical_4_3", "inference_rule": "spm", "raw_label": "cnf(c_0_712, negated_conjecture, (vertical_3_2|horizontal_3_2|~horizontal_4_4|~vertical_4_3),\\ninference(spm,[status(thm)],[c_0_636, c_0_701])).\n", "fillcolor": "lightskyblue1"}, {"id": 713, "label": "c_0_713", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_713", "role": "negated_conjecture", "formula": "horizontal_4_4|~horizontal_2_3", "inference_rule": "csr", "raw_label": "cnf(c_0_713, negated_conjecture, (horizontal_4_4|~horizontal_2_3),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_702, c_0_703]), c_0_646]), c_0_586])).\n", "fillcolor": "lightskyblue1"}, {"id": 714, "label": "c_0_714", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_714", "role": "negated_conjecture", "formula": "vertical_1_2|~horizontal_4_2|~horizontal_2_3", "inference_rule": "csr", "raw_label": "cnf(c_0_714, negated_conjecture, (vertical_1_2|~horizontal_4_2|~horizontal_2_3),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_366, c_0_704]), c_0_665]), c_0_646])).\n", "fillcolor": "lightskyblue1"}, {"id": 715, "label": "c_0_715", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_715", "role": "negated_conjecture", "formula": "vertical_4_4|horizontal_4_4|~horizontal_2_3", "inference_rule": "spm", "raw_label": "cnf(c_0_715, negated_conjecture, (vertical_4_4|horizontal_4_4|~horizontal_2_3),\\ninference(spm,[status(thm)],[c_0_705, c_0_703])).\n", "fillcolor": "lightskyblue1"}, {"id": 716, "label": "c_0_716", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_716", "role": "negated_conjecture", "formula": "horizontal_2_4|~vertical_3_5|~vertical_3_2", "inference_rule": "csr", "raw_label": "cnf(c_0_716, negated_conjecture, (horizontal_2_4|~vertical_3_5|~vertical_3_2),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_638, c_0_585]), c_0_586]), c_0_587])).\n", "fillcolor": "lightskyblue1"}, {"id": 717, "label": "c_0_717", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_717", "role": "negated_conjecture", "formula": "vertical_3_5|vertical_3_3|~vertical_3_2", "inference_rule": "csr", "raw_label": "cnf(c_0_717, negated_conjecture, (vertical_3_5|vertical_3_3|~vertical_3_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_204, c_0_706]), c_0_246])).\n", "fillcolor": "lightskyblue1"}, {"id": 718, "label": "c_0_718", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_718", "role": "negated_conjecture", "formula": "horizontal_2_4|~horizontal_3_4|~vertical_3_2", "inference_rule": "csr", "raw_label": "cnf(c_0_718, negated_conjecture, (horizontal_2_4|~horizontal_3_4|~vertical_3_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_707, c_0_601]), c_0_497])).\n", "fillcolor": "lightskyblue1"}, {"id": 719, "label": "c_0_719", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_719", "role": "negated_conjecture", "formula": "horizontal_3_4|horizontal_3_3|horizontal_2_4|vertical_2_3", "inference_rule": "csr", "raw_label": "cnf(c_0_719, negated_conjecture, (horizontal_3_4|horizontal_3_3|horizontal_2_4|vertical_2_3),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_708, c_0_291]), c_0_709])).\n", "fillcolor": "lightskyblue1"}, {"id": 720, "label": "c_0_720", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_720", "role": "negated_conjecture", "formula": "horizontal_2_4|vertical_2_3|~horizontal_3_4", "inference_rule": "spm", "raw_label": "cnf(c_0_720, negated_conjecture, (horizontal_2_4|vertical_2_3|~horizontal_3_4),\\ninference(spm,[status(thm)],[c_0_710, c_0_711])).\n", "fillcolor": "lightskyblue1"}, {"id": 721, "label": "c_0_721", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_721", "role": "negated_conjecture", "formula": "~vertical_4_3|~horizontal_2_3", "inference_rule": "csr", "raw_label": "cnf(c_0_721, negated_conjecture, (~vertical_4_3|~horizontal_2_3),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_712, c_0_713]), c_0_646]), c_0_586])).\n", "fillcolor": "lightskyblue1"}, {"id": 722, "label": "c_0_722", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_722", "role": "negated_conjecture", "formula": "~horizontal_4_2|~horizontal_2_3", "inference_rule": "csr", "raw_label": "cnf(c_0_722, negated_conjecture, (~horizontal_4_2|~horizontal_2_3),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_503, c_0_714]), c_0_646]), c_0_246])).\n", "fillcolor": "lightskyblue1"}, {"id": 723, "label": "c_0_723", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_723", "role": "negated_conjecture", "formula": "~horizontal_4_3|~horizontal_2_3", "inference_rule": "csr", "raw_label": "cnf(c_0_723, negated_conjecture, (~horizontal_4_3|~horizontal_2_3),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_241, c_0_715]), c_0_243])).\n", "fillcolor": "lightskyblue1"}, {"id": 724, "label": "c_0_724", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_724", "role": "negated_conjecture", "formula": "horizontal_3_4|horizontal_3_3|~horizontal_2_3", "inference_rule": "csr", "raw_label": "cnf(c_0_724, negated_conjecture, (horizontal_3_4|horizontal_3_3|~horizontal_2_3),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_674, c_0_291]), c_0_476])).\n", "fillcolor": "lightskyblue1"}, {"id": 725, "label": "c_0_725", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_725", "role": "negated_conjecture", "formula": "vertical_3_3|horizontal_2_4|~vertical_3_2", "inference_rule": "spm", "raw_label": "cnf(c_0_725, negated_conjecture, (vertical_3_3|horizontal_2_4|~vertical_3_2),\\ninference(spm,[status(thm)],[c_0_716, c_0_717])).\n", "fillcolor": "lightskyblue1"}, {"id": 726, "label": "c_0_726", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_726", "role": "negated_conjecture", "formula": "horizontal_3_3|horizontal_2_4|~vertical_3_2", "inference_rule": "csr", "raw_label": "cnf(c_0_726, negated_conjecture, (horizontal_3_3|horizontal_2_4|~vertical_3_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_718, c_0_719]), c_0_587])).\n", "fillcolor": "lightskyblue1"}, {"id": 727, "label": "c_0_727", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_727", "role": "negated_conjecture", "formula": "horizontal_3_3|horizontal_2_4|vertical_2_3", "inference_rule": "spm", "raw_label": "cnf(c_0_727, negated_conjecture, (horizontal_3_3|horizontal_2_4|vertical_2_3),\\ninference(spm,[status(thm)],[c_0_720, c_0_719])).\n", "fillcolor": "lightskyblue1"}, {"id": 728, "label": "c_0_728", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_728", "role": "negated_conjecture", "formula": "vertical_3_3|~horizontal_2_3", "inference_rule": "csr", "raw_label": "cnf(c_0_728, negated_conjecture, (vertical_3_3|~horizontal_2_3),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_721, c_0_245]), c_0_722]), c_0_723])).\n", "fillcolor": "lightskyblue1"}, {"id": 729, "label": "c_0_729", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_729", "role": "negated_conjecture", "formula": "horizontal_3_3|~horizontal_2_3", "inference_rule": "csr", "raw_label": "cnf(c_0_729, negated_conjecture, (horizontal_3_3|~horizontal_2_3),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_292, c_0_724]), c_0_665])).\n", "fillcolor": "lightskyblue1"}, {"id": 730, "label": "c_0_730", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_730", "role": "negated_conjecture", "formula": "horizontal_2_4|~vertical_3_2", "inference_rule": "csr", "raw_label": "cnf(c_0_730, negated_conjecture, (horizontal_2_4|~vertical_3_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_288, c_0_725]), c_0_726])).\n", "fillcolor": "lightskyblue1"}, {"id": 731, "label": "c_0_731", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_731", "role": "negated_conjecture", "formula": "horizontal_2_4|~horizontal_3_2", "inference_rule": "csr", "raw_label": "cnf(c_0_731, negated_conjecture, (horizontal_2_4|~horizontal_3_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_336, c_0_727]), c_0_545])).\n", "fillcolor": "lightskyblue1"}, {"id": 732, "label": "c_0_732", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_732", "role": "negated_conjecture", "formula": "horizontal_5_5|vertical_1_4|~horizontal_3_4|~vertical_1_3", "inference_rule": "spm", "raw_label": "cnf(c_0_732, negated_conjecture, (horizontal_5_5|vertical_1_4|~horizontal_3_4|~vertical_1_3),\\ninference(spm,[status(thm)],[c_0_284, c_0_504])).\n", "fillcolor": "lightskyblue1"}, {"id": 733, "label": "c_0_733", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_733", "role": "negated_conjecture", "formula": "~horizontal_2_3", "inference_rule": "csr", "raw_label": "cnf(c_0_733, negated_conjecture, (~horizontal_2_3),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_288, c_0_728]), c_0_729])).\n", "fillcolor": "lightskyblue1"}, {"id": 734, "label": "c_0_734", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_734", "role": "negated_conjecture", "formula": "horizontal_2_4|vertical_2_3|horizontal_2_3", "inference_rule": "csr", "raw_label": "cnf(c_0_734, negated_conjecture, (horizontal_2_4|vertical_2_3|horizontal_2_3),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_730, c_0_679]), c_0_731])).\n", "fillcolor": "lightskyblue1"}, {"id": 735, "label": "c_0_735", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_735", "role": "negated_conjecture", "formula": "horizontal_3_5|vertical_1_4|~horizontal_4_5|~vertical_1_3", "inference_rule": "spm", "raw_label": "cnf(c_0_735, negated_conjecture, (horizontal_3_5|vertical_1_4|~horizontal_4_5|~vertical_1_3),\\ninference(spm,[status(thm)],[c_0_411, c_0_592])).\n", "fillcolor": "lightskyblue1"}, {"id": 736, "label": "c_0_736", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_736", "role": "negated_conjecture", "formula": "vertical_5_3|horizontal_4_1|~vertical_5_4|~vertical_2_1", "inference_rule": "spm", "raw_label": "cnf(c_0_736, negated_conjecture, (vertical_5_3|horizontal_4_1|~vertical_5_4|~vertical_2_1),\\ninference(spm,[status(thm)],[c_0_117, c_0_604])).\n", "fillcolor": "lightskyblue1"}, {"id": 737, "label": "c_0_737", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_737", "role": "negated_conjecture", "formula": "horizontal_5_5|~horizontal_3_4|~horizontal_2_4|~vertical_1_3", "inference_rule": "spm", "raw_label": "cnf(c_0_737, negated_conjecture, (horizontal_5_5|~horizontal_3_4|~horizontal_2_4|~vertical_1_3),\\ninference(spm,[status(thm)],[c_0_302, c_0_732])).\n", "fillcolor": "lightskyblue1"}, {"id": 738, "label": "c_0_738", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_738", "role": "negated_conjecture", "formula": "vertical_2_3|vertical_1_3", "inference_rule": "sr", "raw_label": "cnf(c_0_738, negated_conjecture, (vertical_2_3|vertical_1_3),\\ninference(sr,[status(thm)],[c_0_562, c_0_733])).\n", "fillcolor": "lightskyblue1"}, {"id": 739, "label": "c_0_739", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_739", "role": "negated_conjecture", "formula": "horizontal_2_4|vertical_2_3", "inference_rule": "sr", "raw_label": "cnf(c_0_739, negated_conjecture, (horizontal_2_4|vertical_2_3),\\ninference(sr,[status(thm)],[c_0_734, c_0_733])).\n", "fillcolor": "lightskyblue1"}, {"id": 740, "label": "c_0_740", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_740", "role": "negated_conjecture", "formula": "vertical_4_5|horizontal_4_4|vertical_3_5|horizontal_3_5|vertical_1_4|~vertical_1_3", "inference_rule": "spm", "raw_label": "cnf(c_0_740, negated_conjecture, (vertical_4_5|horizontal_4_4|vertical_3_5|horizontal_3_5|vertical_1_4|~vertical_1_3),\\ninference(spm,[status(thm)],[c_0_735, c_0_152])).\n", "fillcolor": "lightskyblue1"}, {"id": 741, "label": "c_0_741", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_741", "role": "negated_conjecture", "formula": "vertical_5_3|horizontal_4_1|horizontal_2_1|~vertical_5_4", "inference_rule": "spm", "raw_label": "cnf(c_0_741, negated_conjecture, (vertical_5_3|horizontal_4_1|horizontal_2_1|~vertical_5_4),\\ninference(spm,[status(thm)],[c_0_736, c_0_149])).\n", "fillcolor": "lightskyblue1"}, {"id": 742, "label": "c_0_742", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_742", "role": "negated_conjecture", "formula": "horizontal_5_5|vertical_1_4|~vertical_3_5|~vertical_1_3", "inference_rule": "spm", "raw_label": "cnf(c_0_742, negated_conjecture, (horizontal_5_5|vertical_1_4|~vertical_3_5|~vertical_1_3),\\ninference(spm,[status(thm)],[c_0_270, c_0_504])).\n", "fillcolor": "lightskyblue1"}, {"id": 743, "label": "c_0_743", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_743", "role": "negated_conjecture", "formula": "horizontal_5_5|vertical_2_3|horizontal_2_3|~horizontal_3_4", "inference_rule": "csr", "raw_label": "cnf(c_0_743, negated_conjecture, (horizontal_5_5|vertical_2_3|horizontal_2_3|~horizontal_3_4),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_737, c_0_562]), c_0_315])).\n", "fillcolor": "lightskyblue1"}, {"id": 744, "label": "c_0_744", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_744", "role": "negated_conjecture", "formula": "horizontal_5_5|vertical_2_3|~horizontal_3_4", "inference_rule": "csr", "raw_label": "cnf(c_0_744, negated_conjecture, (horizontal_5_5|vertical_2_3|~horizontal_3_4),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_737, c_0_738]), c_0_739])).\n", "fillcolor": "lightskyblue1"}, {"id": 745, "label": "c_0_745", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_745", "role": "negated_conjecture", "formula": "vertical_4_5|horizontal_4_4|vertical_1_4|~horizontal_3_4|~vertical_1_3", "inference_rule": "csr", "raw_label": "cnf(c_0_745, negated_conjecture, (vertical_4_5|horizontal_4_4|vertical_1_4|~horizontal_3_4|~vertical_1_3),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_284, c_0_740]), c_0_497])).\n", "fillcolor": "lightskyblue1"}, {"id": 746, "label": "c_0_746", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_746", "role": "negated_conjecture", "formula": "~horizontal_4_4|~horizontal_3_4|~horizontal_2_3", "inference_rule": "spm", "raw_label": "cnf(c_0_746, negated_conjecture, (~horizontal_4_4|~horizontal_3_4|~horizontal_2_3),\\ninference(spm,[status(thm)],[c_0_636, c_0_581])).\n", "fillcolor": "lightskyblue1"}, {"id": 747, "label": "c_0_747", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_747", "role": "negated_conjecture", "formula": "horizontal_5_4|vertical_5_3|horizontal_5_3|vertical_4_4|horizontal_4_1|horizontal_2_1", "inference_rule": "spm", "raw_label": "cnf(c_0_747, negated_conjecture, (horizontal_5_4|vertical_5_3|horizontal_5_3|vertical_4_4|horizontal_4_1|horizontal_2_1),\\ninference(spm,[status(thm)],[c_0_741, c_0_199])).\n", "fillcolor": "lightskyblue1"}, {"id": 748, "label": "c_0_748", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_748", "role": "negated_conjecture", "formula": "vertical_2_3|vertical_1_2", "inference_rule": "sr", "raw_label": "cnf(c_0_748, negated_conjecture, (vertical_2_3|vertical_1_2),\\ninference(sr,[status(thm)],[c_0_585, c_0_733])).\n", "fillcolor": "lightskyblue1"}, {"id": 749, "label": "c_0_749", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_749", "role": "negated_conjecture", "formula": "horizontal_5_5|~vertical_3_5|~horizontal_2_4|~vertical_1_3", "inference_rule": "spm", "raw_label": "cnf(c_0_749, negated_conjecture, (horizontal_5_5|~vertical_3_5|~horizontal_2_4|~vertical_1_3),\\ninference(spm,[status(thm)],[c_0_302, c_0_742])).\n", "fillcolor": "lightskyblue1"}, {"id": 750, "label": "c_0_750", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_750", "role": "negated_conjecture", "formula": "vertical_2_3|horizontal_2_3|~vertical_5_5|~horizontal_3_4", "inference_rule": "spm", "raw_label": "cnf(c_0_750, negated_conjecture, (vertical_2_3|horizontal_2_3|~vertical_5_5|~horizontal_3_4),\\ninference(spm,[status(thm)],[c_0_164, c_0_743])).\n", "fillcolor": "lightskyblue1"}, {"id": 751, "label": "c_0_751", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_751", "role": "negated_conjecture", "formula": "vertical_2_3|~vertical_5_5|~horizontal_3_4", "inference_rule": "spm", "raw_label": "cnf(c_0_751, negated_conjecture, (vertical_2_3|~vertical_5_5|~horizontal_3_4),\\ninference(spm,[status(thm)],[c_0_164, c_0_744])).\n", "fillcolor": "lightskyblue1"}, {"id": 752, "label": "c_0_752", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_752", "role": "negated_conjecture", "formula": "vertical_2_3|~horizontal_5_4|~horizontal_3_4", "inference_rule": "spm", "raw_label": "cnf(c_0_752, negated_conjecture, (vertical_2_3|~horizontal_5_4|~horizontal_3_4),\\ninference(spm,[status(thm)],[c_0_200, c_0_744])).\n", "fillcolor": "lightskyblue1"}, {"id": 753, "label": "c_0_753", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_753", "role": "negated_conjecture", "formula": "vertical_4_5|horizontal_4_4|~horizontal_3_4|~vertical_1_3", "inference_rule": "csr", "raw_label": "cnf(c_0_753, negated_conjecture, (vertical_4_5|horizontal_4_4|~horizontal_3_4|~vertical_1_3),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_302, c_0_745]), c_0_711])).\n", "fillcolor": "lightskyblue1"}, {"id": 754, "label": "c_0_754", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_754", "role": "negated_conjecture", "formula": "~horizontal_3_4|~horizontal_2_3", "inference_rule": "csr", "raw_label": "cnf(c_0_754, negated_conjecture, (~horizontal_3_4|~horizontal_2_3),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_746, c_0_711]), c_0_356])).\n", "fillcolor": "lightskyblue1"}, {"id": 755, "label": "c_0_755", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_755", "role": "negated_conjecture", "formula": "horizontal_5_4|vertical_4_4|horizontal_4_1|horizontal_2_1|~horizontal_5_2", "inference_rule": "csr", "raw_label": "cnf(c_0_755, negated_conjecture, (horizontal_5_4|vertical_4_4|horizontal_4_1|horizontal_2_1|~horizontal_5_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_252, c_0_747]), c_0_170])).\n", "fillcolor": "lightskyblue1"}, {"id": 756, "label": "c_0_756", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_756", "role": "negated_conjecture", "formula": "horizontal_5_2|~vertical_3_2", "inference_rule": "csr", "raw_label": "cnf(c_0_756, negated_conjecture, (horizontal_5_2|~vertical_3_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_430, c_0_748]), c_0_587])).\n", "fillcolor": "lightskyblue1"}, {"id": 757, "label": "c_0_757", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_757", "role": "negated_conjecture", "formula": "horizontal_5_5|vertical_2_3|horizontal_2_3|~vertical_3_5", "inference_rule": "csr", "raw_label": "cnf(c_0_757, negated_conjecture, (horizontal_5_5|vertical_2_3|horizontal_2_3|~vertical_3_5),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_749, c_0_562]), c_0_301])).\n", "fillcolor": "lightskyblue1"}, {"id": 758, "label": "c_0_758", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_758", "role": "negated_conjecture", "formula": "vertical_4_2|vertical_2_3|~horizontal_5_3|~horizontal_3_4", "inference_rule": "csr", "raw_label": "cnf(c_0_758, negated_conjecture, (vertical_4_2|vertical_2_3|~horizontal_5_3|~horizontal_3_4),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_750, c_0_181]), c_0_567])).\n", "fillcolor": "lightskyblue1"}, {"id": 759, "label": "c_0_759", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_759", "role": "negated_conjecture", "formula": "horizontal_5_3|vertical_4_4|vertical_2_3|~horizontal_3_4", "inference_rule": "csr", "raw_label": "cnf(c_0_759, negated_conjecture, (horizontal_5_3|vertical_4_4|vertical_2_3|~horizontal_3_4),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_751, c_0_213]), c_0_752])).\n", "fillcolor": "lightskyblue1"}, {"id": 760, "label": "c_0_760", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_760", "role": "negated_conjecture", "formula": "vertical_2_3|horizontal_2_3|~vertical_4_5|~horizontal_3_4", "inference_rule": "spm", "raw_label": "cnf(c_0_760, negated_conjecture, (vertical_2_3|horizontal_2_3|~vertical_4_5|~horizontal_3_4),\\ninference(spm,[status(thm)],[c_0_628, c_0_743])).\n", "fillcolor": "lightskyblue1"}, {"id": 761, "label": "c_0_761", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_761", "role": "negated_conjecture", "formula": "vertical_4_5|horizontal_4_4|vertical_2_3|~horizontal_3_4", "inference_rule": "csr", "raw_label": "cnf(c_0_761, negated_conjecture, (vertical_4_5|horizontal_4_4|vertical_2_3|~horizontal_3_4),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_753, c_0_562]), c_0_754])).\n", "fillcolor": "lightskyblue1"}, {"id": 762, "label": "c_0_762", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_762", "role": "negated_conjecture", "formula": "vertical_4_4|horizontal_2_1|~vertical_3_2", "inference_rule": "csr", "raw_label": "cnf(c_0_762, negated_conjecture, (vertical_4_4|horizontal_2_1|~vertical_3_2),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_202, c_0_755]), c_0_756]), c_0_663])).\n", "fillcolor": "lightskyblue1"}, {"id": 763, "label": "c_0_763", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_763", "role": "negated_conjecture", "formula": "horizontal_5_5|vertical_2_3|~vertical_3_5", "inference_rule": "csr", "raw_label": "cnf(c_0_763, negated_conjecture, (horizontal_5_5|vertical_2_3|~vertical_3_5),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_749, c_0_738]), c_0_739])).\n", "fillcolor": "lightskyblue1"}, {"id": 764, "label": "c_0_764", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_764", "role": "negated_conjecture", "formula": "vertical_2_3|horizontal_2_3|~vertical_5_5|~vertical_3_5", "inference_rule": "spm", "raw_label": "cnf(c_0_764, negated_conjecture, (vertical_2_3|horizontal_2_3|~vertical_5_5|~vertical_3_5),\\ninference(spm,[status(thm)],[c_0_164, c_0_757])).\n", "fillcolor": "lightskyblue1"}, {"id": 765, "label": "c_0_765", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_765", "role": "negated_conjecture", "formula": "vertical_4_4|vertical_4_2|vertical_2_3|~horizontal_3_4", "inference_rule": "spm", "raw_label": "cnf(c_0_765, negated_conjecture, (vertical_4_4|vertical_4_2|vertical_2_3|~horizontal_3_4),\\ninference(spm,[status(thm)],[c_0_758, c_0_759])).\n", "fillcolor": "lightskyblue1"}, {"id": 766, "label": "c_0_766", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_766", "role": "negated_conjecture", "formula": "horizontal_4_4|vertical_2_3|~horizontal_3_4", "inference_rule": "csr", "raw_label": "cnf(c_0_766, negated_conjecture, (horizontal_4_4|vertical_2_3|~horizontal_3_4),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_760, c_0_761]), c_0_754])).\n", "fillcolor": "lightskyblue1"}, {"id": 767, "label": "c_0_767", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_767", "role": "negated_conjecture", "formula": "vertical_4_4|~vertical_3_2|~vertical_1_2", "inference_rule": "spm", "raw_label": "cnf(c_0_767, negated_conjecture, (vertical_4_4|~vertical_3_2|~vertical_1_2),\\ninference(spm,[status(thm)],[c_0_378, c_0_762])).\n", "fillcolor": "lightskyblue1"}, {"id": 768, "label": "c_0_768", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_768", "role": "negated_conjecture", "formula": "horizontal_2_1|~horizontal_5_2|~vertical_4_5|~vertical_3_2", "inference_rule": "spm", "raw_label": "cnf(c_0_768, negated_conjecture, (horizontal_2_1|~horizontal_5_2|~vertical_4_5|~vertical_3_2),\\ninference(spm,[status(thm)],[c_0_166, c_0_279])).\n", "fillcolor": "lightskyblue1"}, {"id": 769, "label": "c_0_769", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_769", "role": "negated_conjecture", "formula": "vertical_2_3|~vertical_5_5|~vertical_3_5", "inference_rule": "spm", "raw_label": "cnf(c_0_769, negated_conjecture, (vertical_2_3|~vertical_5_5|~vertical_3_5),\\ninference(spm,[status(thm)],[c_0_164, c_0_763])).\n", "fillcolor": "lightskyblue1"}, {"id": 770, "label": "c_0_770", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_770", "role": "negated_conjecture", "formula": "vertical_4_2|vertical_2_3|~horizontal_5_3|~vertical_3_5", "inference_rule": "csr", "raw_label": "cnf(c_0_770, negated_conjecture, (vertical_4_2|vertical_2_3|~horizontal_5_3|~vertical_3_5),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_764, c_0_181]), c_0_490])).\n", "fillcolor": "lightskyblue1"}, {"id": 771, "label": "c_0_771", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_771", "role": "negated_conjecture", "formula": "vertical_4_2|vertical_2_3|~horizontal_3_4", "inference_rule": "csr", "raw_label": "cnf(c_0_771, negated_conjecture, (vertical_4_2|vertical_2_3|~horizontal_3_4),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_636, c_0_765]), c_0_766])).\n", "fillcolor": "lightskyblue1"}, {"id": 772, "label": "c_0_772", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_772", "role": "negated_conjecture", "formula": "vertical_4_4|~vertical_3_2", "inference_rule": "csr", "raw_label": "cnf(c_0_772, negated_conjecture, (vertical_4_4|~vertical_3_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_767, c_0_748]), c_0_587])).\n", "fillcolor": "lightskyblue1"}, {"id": 773, "label": "c_0_773", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_773", "role": "negated_conjecture", "formula": "~vertical_4_3|~vertical_3_2", "inference_rule": "spm", "raw_label": "cnf(c_0_773, negated_conjecture, (~vertical_4_3|~vertical_3_2),\\ninference(spm,[status(thm)],[c_0_228, c_0_756])).\n", "fillcolor": "lightskyblue1"}, {"id": 774, "label": "c_0_774", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_774", "role": "negated_conjecture", "formula": "~horizontal_5_2|~vertical_4_5|~vertical_3_2|~vertical_1_2", "inference_rule": "spm", "raw_label": "cnf(c_0_774, negated_conjecture, (~horizontal_5_2|~vertical_4_5|~vertical_3_2|~vertical_1_2),\\ninference(spm,[status(thm)],[c_0_378, c_0_768])).\n", "fillcolor": "lightskyblue1"}, {"id": 775, "label": "c_0_775", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_775", "role": "negated_conjecture", "formula": "vertical_2_3|horizontal_2_1|~horizontal_5_3|~vertical_3_5", "inference_rule": "csr", "raw_label": "cnf(c_0_775, negated_conjecture, (vertical_2_3|horizontal_2_1|~horizontal_5_3|~vertical_3_5),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_769, c_0_196]), c_0_770])).\n", "fillcolor": "lightskyblue1"}, {"id": 776, "label": "c_0_776", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_776", "role": "negated_conjecture", "formula": "vertical_2_3|~horizontal_5_4|~vertical_3_5", "inference_rule": "spm", "raw_label": "cnf(c_0_776, negated_conjecture, (vertical_2_3|~horizontal_5_4|~vertical_3_5),\\ninference(spm,[status(thm)],[c_0_200, c_0_763])).\n", "fillcolor": "lightskyblue1"}, {"id": 777, "label": "c_0_777", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_777", "role": "negated_conjecture", "formula": "vertical_2_3|horizontal_2_1|~horizontal_5_3|~horizontal_3_4", "inference_rule": "csr", "raw_label": "cnf(c_0_777, negated_conjecture, (vertical_2_3|horizontal_2_1|~horizontal_5_3|~horizontal_3_4),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_751, c_0_196]), c_0_771])).\n", "fillcolor": "lightskyblue1"}, {"id": 778, "label": "c_0_778", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_778", "role": "negated_conjecture", "formula": "~horizontal_4_3|~vertical_3_2", "inference_rule": "spm", "raw_label": "cnf(c_0_778, negated_conjecture, (~horizontal_4_3|~vertical_3_2),\\ninference(spm,[status(thm)],[c_0_241, c_0_772])).\n", "fillcolor": "lightskyblue1"}, {"id": 779, "label": "c_0_779", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_779", "role": "negated_conjecture", "formula": "horizontal_4_3|vertical_3_3|~vertical_3_2", "inference_rule": "csr", "raw_label": "cnf(c_0_779, negated_conjecture, (horizontal_4_3|vertical_3_3|~vertical_3_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_773, c_0_245]), c_0_246])).\n", "fillcolor": "lightskyblue1"}, {"id": 780, "label": "c_0_780", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_780", "role": "negated_conjecture", "formula": "~vertical_4_5|~vertical_3_2", "inference_rule": "csr", "raw_label": "cnf(c_0_780, negated_conjecture, (~vertical_4_5|~vertical_3_2),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_774, c_0_585]), c_0_642]), c_0_586]), c_0_587])).\n", "fillcolor": "lightskyblue1"}, {"id": 781, "label": "c_0_781", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_781", "role": "negated_conjecture", "formula": "vertical_2_3|~horizontal_5_3|~vertical_3_5", "inference_rule": "csr", "raw_label": "cnf(c_0_781, negated_conjecture, (vertical_2_3|~horizontal_5_3|~vertical_3_5),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_378, c_0_775]), c_0_748])).\n", "fillcolor": "lightskyblue1"}, {"id": 782, "label": "c_0_782", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_782", "role": "negated_conjecture", "formula": "horizontal_5_3|vertical_4_4|vertical_2_3|~vertical_3_5", "inference_rule": "csr", "raw_label": "cnf(c_0_782, negated_conjecture, (horizontal_5_3|vertical_4_4|vertical_2_3|~vertical_3_5),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_769, c_0_213]), c_0_776])).\n", "fillcolor": "lightskyblue1"}, {"id": 783, "label": "c_0_783", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_783", "role": "negated_conjecture", "formula": "vertical_2_3|~horizontal_5_3|~horizontal_3_4", "inference_rule": "csr", "raw_label": "cnf(c_0_783, negated_conjecture, (vertical_2_3|~horizontal_5_3|~horizontal_3_4),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_378, c_0_777]), c_0_748])).\n", "fillcolor": "lightskyblue1"}, {"id": 784, "label": "c_0_784", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_784", "role": "negated_conjecture", "formula": "vertical_3_3|~vertical_3_2", "inference_rule": "spm", "raw_label": "cnf(c_0_784, negated_conjecture, (vertical_3_3|~vertical_3_2),\\ninference(spm,[status(thm)],[c_0_778, c_0_779])).\n", "fillcolor": "lightskyblue1"}, {"id": 785, "label": "c_0_785", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_785", "role": "negated_conjecture", "formula": "~horizontal_4_4|~vertical_3_2", "inference_rule": "spm", "raw_label": "cnf(c_0_785, negated_conjecture, (~horizontal_4_4|~vertical_3_2),\\ninference(spm,[status(thm)],[c_0_636, c_0_772])).\n", "fillcolor": "lightskyblue1"}, {"id": 786, "label": "c_0_786", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_786", "role": "negated_conjecture", "formula": "horizontal_4_4|~horizontal_3_4|~vertical_3_2", "inference_rule": "csr", "raw_label": "cnf(c_0_786, negated_conjecture, (horizontal_4_4|~horizontal_3_4|~vertical_3_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_780, c_0_761]), c_0_587])).\n", "fillcolor": "lightskyblue1"}, {"id": 787, "label": "c_0_787", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_787", "role": "negated_conjecture", "formula": "vertical_4_4|vertical_2_3|~vertical_3_5", "inference_rule": "spm", "raw_label": "cnf(c_0_787, negated_conjecture, (vertical_4_4|vertical_2_3|~vertical_3_5),\\ninference(spm,[status(thm)],[c_0_781, c_0_782])).\n", "fillcolor": "lightskyblue1"}, {"id": 788, "label": "c_0_788", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_788", "role": "negated_conjecture", "formula": "vertical_4_4|vertical_2_3|~horizontal_3_4", "inference_rule": "spm", "raw_label": "cnf(c_0_788, negated_conjecture, (vertical_4_4|vertical_2_3|~horizontal_3_4),\\ninference(spm,[status(thm)],[c_0_783, c_0_759])).\n", "fillcolor": "lightskyblue1"}, {"id": 789, "label": "c_0_789", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_789", "role": "negated_conjecture", "formula": "~vertical_3_4|~vertical_3_2", "inference_rule": "spm", "raw_label": "cnf(c_0_789, negated_conjecture, (~vertical_3_4|~vertical_3_2),\\ninference(spm,[status(thm)],[c_0_247, c_0_772])).\n", "fillcolor": "lightskyblue1"}, {"id": 790, "label": "c_0_790", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_790", "role": "negated_conjecture", "formula": "~horizontal_3_3|~vertical_3_2", "inference_rule": "spm", "raw_label": "cnf(c_0_790, negated_conjecture, (~horizontal_3_3|~vertical_3_2),\\ninference(spm,[status(thm)],[c_0_288, c_0_784])).\n", "fillcolor": "lightskyblue1"}, {"id": 791, "label": "c_0_791", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_791", "role": "negated_conjecture", "formula": "~horizontal_3_4|~vertical_3_2", "inference_rule": "spm", "raw_label": "cnf(c_0_791, negated_conjecture, (~horizontal_3_4|~vertical_3_2),\\ninference(spm,[status(thm)],[c_0_785, c_0_786])).\n", "fillcolor": "lightskyblue1"}, {"id": 792, "label": "c_0_792", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_792", "role": "negated_conjecture", "formula": "vertical_2_3|~vertical_3_4", "inference_rule": "csr", "raw_label": "cnf(c_0_792, negated_conjecture, (vertical_2_3|~vertical_3_4),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_247, c_0_787]), c_0_694])).\n", "fillcolor": "lightskyblue1"}, {"id": 793, "label": "c_0_793", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_793", "role": "negated_conjecture", "formula": "vertical_2_3|~horizontal_3_4", "inference_rule": "csr", "raw_label": "cnf(c_0_793, negated_conjecture, (vertical_2_3|~horizontal_3_4),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_636, c_0_788]), c_0_766])).\n", "fillcolor": "lightskyblue1"}, {"id": 794, "label": "c_0_794", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_794", "role": "negated_conjecture", "formula": "vertical_3_2|horizontal_3_2|vertical_2_3", "inference_rule": "sr", "raw_label": "cnf(c_0_794, negated_conjecture, (vertical_3_2|horizontal_3_2|vertical_2_3),\\ninference(sr,[status(thm)],[c_0_679, c_0_733])).\n", "fillcolor": "lightskyblue1"}, {"id": 795, "label": "c_0_795", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_795", "role": "negated_conjecture", "formula": "~vertical_3_2", "inference_rule": "csr", "raw_label": "cnf(c_0_795, negated_conjecture, (~vertical_3_2),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_789, c_0_291]), c_0_678]), c_0_790]), c_0_791])).\n", "fillcolor": "lightskyblue1"}, {"id": 796, "label": "c_0_796", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_796", "role": "negated_conjecture", "formula": "horizontal_3_3|vertical_2_3", "inference_rule": "csr", "raw_label": "cnf(c_0_796, negated_conjecture, (horizontal_3_3|vertical_2_3),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_792, c_0_291]), c_0_709]), c_0_793])).\n", "fillcolor": "lightskyblue1"}, {"id": 797, "label": "c_0_797", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_797", "role": "negated_conjecture", "formula": "horizontal_3_2|vertical_2_3", "inference_rule": "sr", "raw_label": "cnf(c_0_797, negated_conjecture, (horizontal_3_2|vertical_2_3),\\ninference(sr,[status(thm)],[c_0_794, c_0_795])).\n", "fillcolor": "lightskyblue1"}, {"id": 798, "label": "c_0_798", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_798", "role": "negated_conjecture", "formula": "vertical_3_2|horizontal_3_2|horizontal_2_1", "inference_rule": "sr", "raw_label": "cnf(c_0_798, negated_conjecture, (vertical_3_2|horizontal_3_2|horizontal_2_1),\\ninference(sr,[status(thm)],[c_0_309, c_0_569])).\n", "fillcolor": "lightskyblue1"}, {"id": 799, "label": "c_0_799", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_799", "role": "negated_conjecture", "formula": "vertical_2_3", "inference_rule": "csr", "raw_label": "cnf(c_0_799, negated_conjecture, (vertical_2_3),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_336, c_0_796]), c_0_797])).\n", "fillcolor": "lightskyblue1"}, {"id": 800, "label": "c_0_800", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_800", "role": "negated_conjecture", "formula": "horizontal_2_4|horizontal_2_1|~horizontal_4_2|~vertical_3_5", "inference_rule": "csr", "raw_label": "cnf(c_0_800, negated_conjecture, (horizontal_2_4|horizontal_2_1|~horizontal_4_2|~vertical_3_5),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_317, c_0_341]), c_0_342])).\n", "fillcolor": "lightskyblue1"}, {"id": 801, "label": "c_0_801", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_801", "role": "negated_conjecture", "formula": "horizontal_3_2|horizontal_3_1", "inference_rule": "sr", "raw_label": "cnf(c_0_801, negated_conjecture, (horizontal_3_2|horizontal_3_1),\\ninference(sr,[status(thm)],[c_0_583, c_0_795])).\n", "fillcolor": "lightskyblue1"}, {"id": 802, "label": "c_0_802", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_802", "role": "negated_conjecture", "formula": "horizontal_3_2|horizontal_2_4|~horizontal_5_3|~horizontal_3_4", "inference_rule": "csr", "raw_label": "cnf(c_0_802, negated_conjecture, (horizontal_3_2|horizontal_2_4|~horizontal_5_3|~horizontal_3_4),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_329, c_0_600]), c_0_718])).\n", "fillcolor": "lightskyblue1"}, {"id": 803, "label": "c_0_803", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_803", "role": "negated_conjecture", "formula": "horizontal_3_2|horizontal_2_4|~horizontal_5_3|~vertical_3_5", "inference_rule": "csr", "raw_label": "cnf(c_0_803, negated_conjecture, (horizontal_3_2|horizontal_2_4|~horizontal_5_3|~vertical_3_5),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_317, c_0_600]), c_0_716])).\n", "fillcolor": "lightskyblue1"}, {"id": 804, "label": "c_0_804", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_804", "role": "negated_conjecture", "formula": "horizontal_3_2|horizontal_2_1", "inference_rule": "sr", "raw_label": "cnf(c_0_804, negated_conjecture, (horizontal_3_2|horizontal_2_1),\\ninference(sr,[status(thm)],[c_0_798, c_0_795])).\n", "fillcolor": "lightskyblue1"}, {"id": 805, "label": "c_0_805", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_805", "role": "negated_conjecture", "formula": "~horizontal_3_2", "inference_rule": "cn", "raw_label": "cnf(c_0_805, negated_conjecture, (~horizontal_3_2),\\ninference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_545, c_0_799])])).\n", "fillcolor": "lightskyblue1"}, {"id": 806, "label": "c_0_806", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_806", "role": "negated_conjecture", "formula": "horizontal_5_4|vertical_4_4|horizontal_4_1|~horizontal_5_2|~horizontal_3_1", "inference_rule": "csr", "raw_label": "cnf(c_0_806, negated_conjecture, (horizontal_5_4|vertical_4_4|horizontal_4_1|~horizontal_5_2|~horizontal_3_1),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_252, c_0_639]), c_0_170])).\n", "fillcolor": "lightskyblue1"}, {"id": 807, "label": "c_0_807", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_807", "role": "negated_conjecture", "formula": "vertical_5_5|horizontal_3_2|vertical_2_2|~horizontal_4_2", "inference_rule": "csr", "raw_label": "cnf(c_0_807, negated_conjecture, (vertical_5_5|horizontal_3_2|vertical_2_2|~horizontal_4_2),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_686, c_0_296]), c_0_687]), c_0_246])).\n", "fillcolor": "lightskyblue1"}, {"id": 808, "label": "c_0_808", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_808", "role": "negated_conjecture", "formula": "horizontal_2_4|~horizontal_4_2|~vertical_3_5|~vertical_2_2", "inference_rule": "spm", "raw_label": "cnf(c_0_808, negated_conjecture, (horizontal_2_4|~horizontal_4_2|~vertical_3_5|~vertical_2_2),\\ninference(spm,[status(thm)],[c_0_395, c_0_800])).\n", "fillcolor": "lightskyblue1"}, {"id": 809, "label": "c_0_809", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_809", "role": "negated_conjecture", "formula": "vertical_4_4|horizontal_3_2|~vertical_4_3", "inference_rule": "csr", "raw_label": "cnf(c_0_809, negated_conjecture, (vertical_4_4|horizontal_3_2|~vertical_4_3),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_683, c_0_801]), c_0_684])).\n", "fillcolor": "lightskyblue1"}, {"id": 810, "label": "c_0_810", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_810", "role": "negated_conjecture", "formula": "vertical_4_4|horizontal_3_2|horizontal_2_4|~horizontal_3_4", "inference_rule": "spm", "raw_label": "cnf(c_0_810, negated_conjecture, (vertical_4_4|horizontal_3_2|horizontal_2_4|~horizontal_3_4),\\ninference(spm,[status(thm)],[c_0_802, c_0_368])).\n", "fillcolor": "lightskyblue1"}, {"id": 811, "label": "c_0_811", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_811", "role": "negated_conjecture", "formula": "vertical_4_4|horizontal_3_2|horizontal_2_4|~vertical_3_5", "inference_rule": "spm", "raw_label": "cnf(c_0_811, negated_conjecture, (vertical_4_4|horizontal_3_2|horizontal_2_4|~vertical_3_5),\\ninference(spm,[status(thm)],[c_0_803, c_0_360])).\n", "fillcolor": "lightskyblue1"}, {"id": 812, "label": "c_0_812", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_812", "role": "negated_conjecture", "formula": "horizontal_2_1", "inference_rule": "sr", "raw_label": "cnf(c_0_812, negated_conjecture, (horizontal_2_1),\\ninference(sr,[status(thm)],[c_0_804, c_0_805])).\n", "fillcolor": "lightskyblue1"}, {"id": 813, "label": "c_0_813", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_813", "role": "negated_conjecture", "formula": "horizontal_5_4|vertical_4_4|~horizontal_4_2|~horizontal_3_1", "inference_rule": "csr", "raw_label": "cnf(c_0_813, negated_conjecture, (horizontal_5_4|vertical_4_4|~horizontal_4_2|~horizontal_3_1),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_326, c_0_806]), c_0_667])).\n", "fillcolor": "lightskyblue1"}, {"id": 814, "label": "c_0_814", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_814", "role": "negated_conjecture", "formula": "horizontal_3_2|~horizontal_5_4|~horizontal_4_2", "inference_rule": "spm", "raw_label": "cnf(c_0_814, negated_conjecture, (horizontal_3_2|~horizontal_5_4|~horizontal_4_2),\\ninference(spm,[status(thm)],[c_0_605, c_0_704])).\n", "fillcolor": "lightskyblue1"}, {"id": 815, "label": "c_0_815", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_815", "role": "negated_conjecture", "formula": "horizontal_4_4|vertical_3_5|horizontal_2_1|~horizontal_4_2", "inference_rule": "csr", "raw_label": "cnf(c_0_815, negated_conjecture, (horizontal_4_4|vertical_3_5|horizontal_2_1|~horizontal_4_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_180, c_0_341]), c_0_342])).\n", "fillcolor": "lightskyblue1"}, {"id": 816, "label": "c_0_816", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_816", "role": "negated_conjecture", "formula": "horizontal_3_2|horizontal_2_4|~horizontal_4_2|~vertical_3_5", "inference_rule": "csr", "raw_label": "cnf(c_0_816, negated_conjecture, (horizontal_3_2|horizontal_2_4|~horizontal_4_2|~vertical_3_5),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_317, c_0_807]), c_0_808])).\n", "fillcolor": "lightskyblue1"}, {"id": 817, "label": "c_0_817", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_817", "role": "negated_conjecture", "formula": "horizontal_3_2|~horizontal_4_4|~vertical_4_3", "inference_rule": "spm", "raw_label": "cnf(c_0_817, negated_conjecture, (horizontal_3_2|~horizontal_4_4|~vertical_4_3),\\ninference(spm,[status(thm)],[c_0_636, c_0_809])).\n", "fillcolor": "lightskyblue1"}, {"id": 818, "label": "c_0_818", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_818", "role": "negated_conjecture", "formula": "horizontal_4_4|vertical_3_5|horizontal_3_2|~vertical_4_3", "inference_rule": "csr", "raw_label": "cnf(c_0_818, negated_conjecture, (horizontal_4_4|vertical_3_5|horizontal_3_2|~vertical_4_3),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_180, c_0_662]), c_0_601])).\n", "fillcolor": "lightskyblue1"}, {"id": 819, "label": "c_0_819", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_819", "role": "negated_conjecture", "formula": "~vertical_3_3", "inference_rule": "cn", "raw_label": "cnf(c_0_819, negated_conjecture, (~vertical_3_3),\\ninference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_500, c_0_799])])).\n", "fillcolor": "lightskyblue1"}, {"id": 820, "label": "c_0_820", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_820", "role": "negated_conjecture", "formula": "horizontal_3_2|horizontal_2_4|~horizontal_4_4|~horizontal_3_4", "inference_rule": "spm", "raw_label": "cnf(c_0_820, negated_conjecture, (horizontal_3_2|horizontal_2_4|~horizontal_4_4|~horizontal_3_4),\\ninference(spm,[status(thm)],[c_0_636, c_0_810])).\n", "fillcolor": "lightskyblue1"}, {"id": 821, "label": "c_0_821", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_821", "role": "negated_conjecture", "formula": "horizontal_3_2|horizontal_2_4|~vertical_3_4", "inference_rule": "csr", "raw_label": "cnf(c_0_821, negated_conjecture, (horizontal_3_2|horizontal_2_4|~vertical_3_4),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_247, c_0_811]), c_0_653])).\n", "fillcolor": "lightskyblue1"}, {"id": 822, "label": "c_0_822", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_822", "role": "negated_conjecture", "formula": "horizontal_3_2|~vertical_1_2", "inference_rule": "spm", "raw_label": "cnf(c_0_822, negated_conjecture, (horizontal_3_2|~vertical_1_2),\\ninference(spm,[status(thm)],[c_0_378, c_0_804])).\n", "fillcolor": "lightskyblue1"}, {"id": 823, "label": "c_0_823", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_823", "role": "negated_conjecture", "formula": "vertical_1_2|~vertical_4_5|~vertical_2_5|~vertical_2_4", "inference_rule": "spm", "raw_label": "cnf(c_0_823, negated_conjecture, (vertical_1_2|~vertical_4_5|~vertical_2_5|~vertical_2_4),\\ninference(spm,[status(thm)],[c_0_628, c_0_264])).\n", "fillcolor": "lightskyblue1"}, {"id": 824, "label": "c_0_824", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_824", "role": "negated_conjecture", "formula": "~vertical_1_2", "inference_rule": "cn", "raw_label": "cnf(c_0_824, negated_conjecture, (~vertical_1_2),\\ninference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_378, c_0_812])])).\n", "fillcolor": "lightskyblue1"}, {"id": 825, "label": "c_0_825", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_825", "role": "negated_conjecture", "formula": "vertical_4_5|horizontal_4_4|vertical_1_2|~vertical_2_4", "inference_rule": "csr", "raw_label": "cnf(c_0_825, negated_conjecture, (vertical_4_5|horizontal_4_4|vertical_1_2|~vertical_2_4),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_250, c_0_664]), c_0_563])).\n", "fillcolor": "lightskyblue1"}, {"id": 826, "label": "c_0_826", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_826", "role": "negated_conjecture", "formula": "vertical_4_4|horizontal_3_2|~horizontal_4_2", "inference_rule": "csr", "raw_label": "cnf(c_0_826, negated_conjecture, (vertical_4_4|horizontal_3_2|~horizontal_4_2),\\ninference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_813, c_0_583]), c_0_246]), c_0_814])).\n", "fillcolor": "lightskyblue1"}, {"id": 827, "label": "c_0_827", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_827", "role": "negated_conjecture", "formula": "horizontal_4_4|vertical_3_5|~horizontal_4_2|~vertical_2_2", "inference_rule": "spm", "raw_label": "cnf(c_0_827, negated_conjecture, (horizontal_4_4|vertical_3_5|~horizontal_4_2|~vertical_2_2),\\ninference(spm,[status(thm)],[c_0_395, c_0_815])).\n", "fillcolor": "lightskyblue1"}, {"id": 828, "label": "c_0_828", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_828", "role": "negated_conjecture", "formula": "horizontal_3_2|vertical_1_2|~horizontal_4_2|~vertical_3_5", "inference_rule": "csr", "raw_label": "cnf(c_0_828, negated_conjecture, (horizontal_3_2|vertical_1_2|~horizontal_4_2|~vertical_3_5),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_332, c_0_704]), c_0_816])).\n", "fillcolor": "lightskyblue1"}, {"id": 829, "label": "c_0_829", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_829", "role": "negated_conjecture", "formula": "vertical_3_5|horizontal_3_2|~vertical_4_3", "inference_rule": "spm", "raw_label": "cnf(c_0_829, negated_conjecture, (vertical_3_5|horizontal_3_2|~vertical_4_3),\\ninference(spm,[status(thm)],[c_0_817, c_0_818])).\n", "fillcolor": "lightskyblue1"}, {"id": 830, "label": "c_0_830", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_830", "role": "negated_conjecture", "formula": "vertical_4_3|horizontal_4_3|horizontal_4_2", "inference_rule": "sr", "raw_label": "cnf(c_0_830, negated_conjecture, (vertical_4_3|horizontal_4_3|horizontal_4_2),\\ninference(sr,[status(thm)],[c_0_245, c_0_819])).\n", "fillcolor": "lightskyblue1"}, {"id": 831, "label": "c_0_831", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_831", "role": "negated_conjecture", "formula": "horizontal_3_2|horizontal_2_4|~horizontal_3_4", "inference_rule": "spm", "raw_label": "cnf(c_0_831, negated_conjecture, (horizontal_3_2|horizontal_2_4|~horizontal_3_4),\\ninference(spm,[status(thm)],[c_0_820, c_0_711])).\n", "fillcolor": "lightskyblue1"}, {"id": 832, "label": "c_0_832", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_832", "role": "negated_conjecture", "formula": "horizontal_3_4|horizontal_3_3|horizontal_3_2|vertical_2_4|horizontal_2_4", "inference_rule": "spm", "raw_label": "cnf(c_0_832, negated_conjecture, (horizontal_3_4|horizontal_3_3|horizontal_3_2|vertical_2_4|horizontal_2_4),\\ninference(spm,[status(thm)],[c_0_821, c_0_291])).\n", "fillcolor": "lightskyblue1"}, {"id": 833, "label": "c_0_833", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_833", "role": "negated_conjecture", "formula": "~horizontal_4_4|~vertical_3_5", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_833, negated_conjecture, (~horizontal_4_4|~vertical_3_5),\\ninference(fof_simplification,[status(thm)],[unique_cover6_4_5])).\n", "fillcolor": "lightskyblue1"}, {"id": 834, "label": "c_0_834", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_834", "role": "negated_conjecture", "formula": "horizontal_3_2|~vertical_4_5|~vertical_2_4", "inference_rule": "csr", "raw_label": "cnf(c_0_834, negated_conjecture, (horizontal_3_2|~vertical_4_5|~vertical_2_4),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_822, c_0_823]), c_0_563])).\n", "fillcolor": "lightskyblue1"}, {"id": 835, "label": "c_0_835", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_835", "role": "negated_conjecture", "formula": "vertical_4_5|horizontal_4_4|~vertical_2_4", "inference_rule": "spm", "raw_label": "cnf(c_0_835, negated_conjecture, (vertical_4_5|horizontal_4_4|~vertical_2_4),\\ninference(spm,[status(thm)],[c_0_824, c_0_825])).\n", "fillcolor": "lightskyblue1"}, {"id": 836, "label": "c_0_836", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_836", "role": "negated_conjecture", "formula": "horizontal_3_2|~horizontal_4_4|~horizontal_4_2", "inference_rule": "spm", "raw_label": "cnf(c_0_836, negated_conjecture, (horizontal_3_2|~horizontal_4_4|~horizontal_4_2),\\ninference(spm,[status(thm)],[c_0_636, c_0_826])).\n", "fillcolor": "lightskyblue1"}, {"id": 837, "label": "c_0_837", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_837", "role": "negated_conjecture", "formula": "horizontal_4_4|vertical_3_5|horizontal_3_2|~horizontal_4_2", "inference_rule": "csr", "raw_label": "cnf(c_0_837, negated_conjecture, (horizontal_4_4|vertical_3_5|horizontal_3_2|~horizontal_4_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_180, c_0_807]), c_0_827])).\n", "fillcolor": "lightskyblue1"}, {"id": 838, "label": "c_0_838", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_838", "role": "negated_conjecture", "formula": "horizontal_3_2|~horizontal_4_2|~vertical_3_5", "inference_rule": "csr", "raw_label": "cnf(c_0_838, negated_conjecture, (horizontal_3_2|~horizontal_4_2|~vertical_3_5),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_503, c_0_828]), c_0_246])).\n", "fillcolor": "lightskyblue1"}, {"id": 839, "label": "c_0_839", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_839", "role": "negated_conjecture", "formula": "vertical_3_5|horizontal_3_2|~horizontal_4_3", "inference_rule": "csr", "raw_label": "cnf(c_0_839, negated_conjecture, (vertical_3_5|horizontal_3_2|~horizontal_4_3),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_241, c_0_632]), c_0_243])).\n", "fillcolor": "lightskyblue1"}, {"id": 840, "label": "c_0_840", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_840", "role": "negated_conjecture", "formula": "horizontal_4_3|horizontal_4_2|vertical_3_5", "inference_rule": "sr", "raw_label": "cnf(c_0_840, negated_conjecture, (horizontal_4_3|horizontal_4_2|vertical_3_5),\\ninference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_829, c_0_830]), c_0_805])).\n", "fillcolor": "lightskyblue1"}, {"id": 841, "label": "c_0_841", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_841", "role": "negated_conjecture", "formula": "horizontal_3_3|vertical_2_4|horizontal_2_4", "inference_rule": "csr", "raw_label": "cnf(c_0_841, negated_conjecture, (horizontal_3_3|vertical_2_4|horizontal_2_4),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_831, c_0_832]), c_0_731])).\n", "fillcolor": "lightskyblue1"}, {"id": 842, "label": "c_0_842", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_842", "role": "negated_conjecture", "formula": "~horizontal_4_4|~vertical_3_5", "inference_rule": null, "raw_label": "cnf(c_0_842, negated_conjecture, (~horizontal_4_4|~vertical_3_5),\\nc_0_833).\n", "fillcolor": "lightskyblue1"}, {"id": 843, "label": "c_0_843", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_843", "role": "negated_conjecture", "formula": "horizontal_4_4|~vertical_2_4", "inference_rule": "sr", "raw_label": "cnf(c_0_843, negated_conjecture, (horizontal_4_4|~vertical_2_4),\\ninference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_834, c_0_835]), c_0_805])).\n", "fillcolor": "lightskyblue1"}, {"id": 844, "label": "c_0_844", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_844", "role": "negated_conjecture", "formula": "horizontal_3_2|~horizontal_4_2", "inference_rule": "csr", "raw_label": "cnf(c_0_844, negated_conjecture, (horizontal_3_2|~horizontal_4_2),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_836, c_0_837]), c_0_838])).\n", "fillcolor": "lightskyblue1"}, {"id": 845, "label": "c_0_845", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_845", "role": "negated_conjecture", "formula": "horizontal_4_2|vertical_3_5", "inference_rule": "sr", "raw_label": "cnf(c_0_845, negated_conjecture, (horizontal_4_2|vertical_3_5),\\ninference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_839, c_0_840]), c_0_805])).\n", "fillcolor": "lightskyblue1"}, {"id": 846, "label": "c_0_846", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_846", "role": "negated_conjecture", "formula": "vertical_2_4|horizontal_2_4|~vertical_2_3", "inference_rule": "spm", "raw_label": "cnf(c_0_846, negated_conjecture, (vertical_2_4|horizontal_2_4|~vertical_2_3),\\ninference(spm,[status(thm)],[c_0_531, c_0_841])).\n", "fillcolor": "lightskyblue1"}, {"id": 847, "label": "c_0_847", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_847", "role": "negated_conjecture", "formula": "~vertical_3_5|~vertical_2_4", "inference_rule": "spm", "raw_label": "cnf(c_0_847, negated_conjecture, (~vertical_3_5|~vertical_2_4),\\ninference(spm,[status(thm)],[c_0_842, c_0_843])).\n", "fillcolor": "lightskyblue1"}, {"id": 848, "label": "c_0_848", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_848", "role": "negated_conjecture", "formula": "vertical_3_5", "inference_rule": "sr", "raw_label": "cnf(c_0_848, negated_conjecture, (vertical_3_5),\\ninference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_844, c_0_845]), c_0_805])).\n", "fillcolor": "lightskyblue1"}, {"id": 849, "label": "c_0_849", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_849", "role": "negated_conjecture", "formula": "horizontal_1_2", "inference_rule": "sr", "raw_label": "cnf(c_0_849, negated_conjecture, (horizontal_1_2),\\ninference(sr,[status(thm)],[c_0_219, c_0_824])).\n", "fillcolor": "lightskyblue1"}, {"id": 850, "label": "c_0_850", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_850", "role": "negated_conjecture", "formula": "vertical_2_4|horizontal_2_4", "inference_rule": "cn", "raw_label": "cnf(c_0_850, negated_conjecture, (vertical_2_4|horizontal_2_4),\\ninference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_846, c_0_799])])).\n", "fillcolor": "lightskyblue1"}, {"id": 851, "label": "c_0_851", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_851", "role": "negated_conjecture", "formula": "~vertical_2_4", "inference_rule": "cn", "raw_label": "cnf(c_0_851, negated_conjecture, (~vertical_2_4),\\ninference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_847, c_0_848])])).\n", "fillcolor": "lightskyblue1"}, {"id": 852, "label": "c_0_852", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_852", "role": "negated_conjecture", "formula": "~horizontal_1_3", "inference_rule": "cn", "raw_label": "cnf(c_0_852, negated_conjecture, (~horizontal_1_3),\\ninference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_189, c_0_849])])).\n", "fillcolor": "lightskyblue1"}, {"id": 853, "label": "c_0_853", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_853", "role": "negated_conjecture", "formula": "horizontal_2_4", "inference_rule": "sr", "raw_label": "cnf(c_0_853, negated_conjecture, (horizontal_2_4),\\ninference(sr,[status(thm)],[c_0_850, c_0_851])).\n", "fillcolor": "lightskyblue1"}, {"id": 854, "label": "c_0_854", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_854", "role": "negated_conjecture", "formula": "vertical_1_4|horizontal_1_4", "inference_rule": "sr", "raw_label": "cnf(c_0_854, negated_conjecture, (vertical_1_4|horizontal_1_4),\\ninference(sr,[status(thm)],[c_0_176, c_0_852])).\n", "fillcolor": "lightskyblue1"}, {"id": 855, "label": "c_0_855", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_855", "role": "negated_conjecture", "formula": "~vertical_1_4", "inference_rule": "cn", "raw_label": "cnf(c_0_855, negated_conjecture, (~vertical_1_4),\\ninference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_302, c_0_853])])).\n", "fillcolor": "lightskyblue1"}, {"id": 856, "label": "c_0_856", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_856", "role": "negated_conjecture", "formula": "horizontal_1_4", "inference_rule": "sr", "raw_label": "cnf(c_0_856, negated_conjecture, (horizontal_1_4),\\ninference(sr,[status(thm)],[c_0_854, c_0_855])).\n", "fillcolor": "lightskyblue1"}, {"id": 857, "label": "c_0_857", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_857", "role": "negated_conjecture", "formula": "~horizontal_3_5", "inference_rule": "cn", "raw_label": "cnf(c_0_857, negated_conjecture, (~horizontal_3_5),\\ninference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_270, c_0_848])])).\n", "fillcolor": "lightskyblue1"}, {"id": 858, "label": "c_0_858", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_858", "role": "negated_conjecture", "formula": "~vertical_4_6", "inference_rule": "sr", "raw_label": "cnf(c_0_858, negated_conjecture, (~vertical_4_6),\\ninference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_175, c_0_856])]), c_0_857])).\n", "fillcolor": "lightskyblue1"}, {"id": 859, "label": "c_0_859", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_859", "role": "negated_conjecture", "formula": "horizontal_3_1", "inference_rule": "sr", "raw_label": "cnf(c_0_859, negated_conjecture, (horizontal_3_1),\\ninference(sr,[status(thm)],[c_0_801, c_0_805])).\n", "fillcolor": "lightskyblue1"}, {"id": 860, "label": "c_0_860", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_860", "role": "negated_conjecture", "formula": "~horizontal_3_3", "inference_rule": "cn", "raw_label": "cnf(c_0_860, negated_conjecture, (~horizontal_3_3),\\ninference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_531, c_0_799])])).\n", "fillcolor": "lightskyblue1"}, {"id": 861, "label": "c_0_861", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_861", "role": "negated_conjecture", "formula": "horizontal_5_5", "inference_rule": "sr", "raw_label": "cnf(c_0_861, negated_conjecture, (horizontal_5_5),\\ninference(sr,[status(thm)],[c_0_138, c_0_858])).\n", "fillcolor": "lightskyblue1"}, {"id": 862, "label": "c_0_862", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_862", "role": "negated_conjecture", "formula": "horizontal_5_2|~horizontal_4_2", "inference_rule": "cn", "raw_label": "cnf(c_0_862, negated_conjecture, (horizontal_5_2|~horizontal_4_2),\\ninference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_667, c_0_859])])).\n", "fillcolor": "lightskyblue1"}, {"id": 863, "label": "c_0_863", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_863", "role": "negated_conjecture", "formula": "vertical_3_4|horizontal_3_4|vertical_2_4", "inference_rule": "sr", "raw_label": "cnf(c_0_863, negated_conjecture, (vertical_3_4|horizontal_3_4|vertical_2_4),\\ninference(sr,[status(thm)],[c_0_291, c_0_860])).\n", "fillcolor": "lightskyblue1"}, {"id": 864, "label": "c_0_864", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_864", "role": "negated_conjecture", "formula": "~horizontal_3_4", "inference_rule": "cn", "raw_label": "cnf(c_0_864, negated_conjecture, (~horizontal_3_4),\\ninference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_497, c_0_848])])).\n", "fillcolor": "lightskyblue1"}, {"id": 865, "label": "c_0_865", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_865", "role": "negated_conjecture", "formula": "~vertical_5_5", "inference_rule": "cn", "raw_label": "cnf(c_0_865, negated_conjecture, (~vertical_5_5),\\ninference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_164, c_0_861])])).\n", "fillcolor": "lightskyblue1"}, {"id": 866, "label": "c_0_866", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_866", "role": "negated_conjecture", "formula": "vertical_5_5|~horizontal_4_2", "inference_rule": "csr", "raw_label": "cnf(c_0_866, negated_conjecture, (vertical_5_5|~horizontal_4_2),\\ninference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_686, c_0_859])]), c_0_862])).\n", "fillcolor": "lightskyblue1"}, {"id": 867, "label": "c_0_867", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_867", "role": "negated_conjecture", "formula": "vertical_5_5|~vertical_4_3", "inference_rule": "csr", "raw_label": "cnf(c_0_867, negated_conjecture, (vertical_5_5|~vertical_4_3),\\ninference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_640, c_0_859])]), c_0_325])).\n", "fillcolor": "lightskyblue1"}, {"id": 868, "label": "c_0_868", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_868", "role": "negated_conjecture", "formula": "vertical_3_4", "inference_rule": "sr", "raw_label": "cnf(c_0_868, negated_conjecture, (vertical_3_4),\\ninference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_863, c_0_851]), c_0_864])).\n", "fillcolor": "lightskyblue1"}, {"id": 869, "label": "c_0_869", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_869", "role": "negated_conjecture", "formula": "~horizontal_4_2", "inference_rule": "spm", "raw_label": "cnf(c_0_869, negated_conjecture, (~horizontal_4_2),\\ninference(spm,[status(thm)],[c_0_865, c_0_866])).\n", "fillcolor": "lightskyblue1"}, {"id": 870, "label": "c_0_870", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_870", "role": "negated_conjecture", "formula": "~vertical_4_3", "inference_rule": "spm", "raw_label": "cnf(c_0_870, negated_conjecture, (~vertical_4_3),\\ninference(spm,[status(thm)],[c_0_865, c_0_867])).\n", "fillcolor": "lightskyblue1"}, {"id": 871, "label": "c_0_871", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_871", "role": "negated_conjecture", "formula": "~horizontal_4_3", "inference_rule": "cn", "raw_label": "cnf(c_0_871, negated_conjecture, (~horizontal_4_3),\\ninference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_519, c_0_868])])).\n", "fillcolor": "lightskyblue1"}, {"id": 872, "label": "c_0_872", "color": {"background": "#000000", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#FFFFFF"}, "cnf_name": "c_0_872", "role": "negated_conjecture", "formula": "$false", "inference_rule": "sr", "raw_label": "cnf(c_0_872, negated_conjecture, ($false),\\ninference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_830, c_0_869]), c_0_870]), c_0_871])).\n", "fillcolor": "darkorchid1"}], "edges": [{"from": 0, "to": 95, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 1, "to": 96, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 95, "to": 97, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 2, "to": 98, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 3, "to": 99, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 96, "to": 100, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 98, "to": 101, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 97, "to": 101, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 99, "to": 102, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 4, "to": 103, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 101, "to": 104, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 100, "to": 104, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 103, "to": 105, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 102, "to": 105, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 5, "to": 106, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 6, "to": 107, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 7, "to": 108, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 105, "to": 109, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 104, "to": 109, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 8, "to": 110, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 9, "to": 111, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 10, "to": 112, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 11, "to": 113, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 106, "to": 114, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 107, "to": 115, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 12, "to": 116, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 108, "to": 117, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 110, "to": 118, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 109, "to": 118, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 13, "to": 119, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 111, "to": 120, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 14, "to": 121, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 15, "to": 122, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 16, "to": 123, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 112, "to": 124, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 113, "to": 125, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 101, "to": 126, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 114, "to": 126, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 98, "to": 127, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 115, "to": 127, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 17, "to": 128, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 18, "to": 129, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 116, "to": 130, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 19, "to": 131, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 20, "to": 132, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 118, "to": 133, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 117, "to": 133, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 120, "to": 134, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 119, "to": 134, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 21, "to": 135, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 121, "to": 136, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 122, "to": 137, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 124, "to": 138, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 123, "to": 138, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 22, "to": 139, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 105, "to": 140, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 125, "to": 140, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 105, "to": 141, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 126, "to": 141, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 127, "to": 142, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 100, "to": 142, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 128, "to": 143, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 129, "to": 144, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 131, "to": 145, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 130, "to": 145, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 132, "to": 146, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 23, "to": 147, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 134, "to": 148, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 133, "to": 148, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 136, "to": 149, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 135, "to": 149, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 24, "to": 150, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 138, "to": 151, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 137, "to": 151, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 25, "to": 152, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 26, "to": 153, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 139, "to": 154, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 134, "to": 155, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 140, "to": 155, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 134, "to": 156, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 141, "to": 156, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 27, "to": 157, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 110, "to": 158, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 142, "to": 158, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 110, "to": 159, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 143, "to": 159, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 145, "to": 160, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 144, "to": 160, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 147, "to": 161, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 146, "to": 161, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 28, "to": 162, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 149, "to": 163, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 148, "to": 163, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 150, "to": 164, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 152, "to": 165, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 151, "to": 165, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 153, "to": 166, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 155, "to": 167, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 154, "to": 167, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 29, "to": 168, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 156, "to": 169, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 154, "to": 169, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 157, "to": 170, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 30, "to": 171, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 158, "to": 172, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 117, "to": 172, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 159, "to": 173, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 117, "to": 173, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 31, "to": 174, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 161, "to": 175, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 160, "to": 175, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 32, "to": 176, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 33, "to": 177, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 162, "to": 178, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 163, "to": 179, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 154, "to": 179, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 166, "to": 180, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 165, "to": 180, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 164, "to": 180, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 170, "to": 181, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 169, "to": 181, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 168, "to": 181, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 167, "to": 181, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 171, "to": 182, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 34, "to": 183, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 35, "to": 184, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 36, "to": 185, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 149, "to": 186, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 172, "to": 186, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 149, "to": 187, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 173, "to": 187, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 37, "to": 188, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 174, "to": 189, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 176, "to": 190, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 175, "to": 190, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 38, "to": 191, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 177, "to": 192, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 39, "to": 193, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 40, "to": 194, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 41, "to": 195, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 179, "to": 196, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 178, "to": 196, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 181, "to": 197, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 180, "to": 197, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 134, "to": 198, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 182, "to": 198, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 42, "to": 199, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 183, "to": 200, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 184, "to": 201, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 185, "to": 202, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 187, "to": 203, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 168, "to": 203, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 186, "to": 203, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 188, "to": 204, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 190, "to": 205, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 189, "to": 205, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 43, "to": 206, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 191, "to": 207, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 145, "to": 208, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 192, "to": 208, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 193, "to": 209, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 194, "to": 210, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 195, "to": 211, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 197, "to": 212, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 196, "to": 212, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 180, "to": 212, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 199, "to": 213, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 198, "to": 213, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 201, "to": 214, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 165, "to": 214, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 200, "to": 214, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 44, "to": 215, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 204, "to": 216, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 203, "to": 216, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 202, "to": 216, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 45, "to": 217, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 138, "to": 218, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 205, "to": 218, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 207, "to": 219, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 206, "to": 219, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 176, "to": 220, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 208, "to": 220, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 176, "to": 221, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 209, "to": 221, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 46, "to": 222, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 131, "to": 223, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 210, "to": 223, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 47, "to": 224, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 212, "to": 225, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 211, "to": 225, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 214, "to": 226, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 213, "to": 226, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 180, "to": 226, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 48, "to": 227, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 215, "to": 228, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 216, "to": 229, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 211, "to": 229, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 49, "to": 230, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 50, "to": 231, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 51, "to": 232, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 52, "to": 233, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 217, "to": 234, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 219, "to": 235, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 218, "to": 235, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 53, "to": 236, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 220, "to": 237, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 189, "to": 237, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 221, "to": 238, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 189, "to": 238, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 222, "to": 239, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 223, "to": 240, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 144, "to": 240, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 224, "to": 241, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 226, "to": 242, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 225, "to": 242, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 227, "to": 243, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 229, "to": 244, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 228, "to": 244, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 54, "to": 245, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 230, "to": 246, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 231, "to": 247, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 226, "to": 248, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 197, "to": 248, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 232, "to": 249, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 233, "to": 250, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 235, "to": 251, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 234, "to": 251, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 236, "to": 252, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 219, "to": 253, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 237, "to": 253, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 55, "to": 254, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 219, "to": 255, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 238, "to": 255, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 56, "to": 256, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 57, "to": 257, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 161, "to": 258, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 239, "to": 258, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 161, "to": 259, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 240, "to": 259, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 58, "to": 260, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 243, "to": 261, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 242, "to": 261, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 241, "to": 261, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 246, "to": 262, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 245, "to": 262, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 244, "to": 262, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 249, "to": 263, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 248, "to": 263, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 247, "to": 263, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 251, "to": 264, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 250, "to": 264, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 163, "to": 265, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 252, "to": 265, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 255, "to": 266, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 254, "to": 266, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 253, "to": 266, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 256, "to": 267, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 59, "to": 268, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 60, "to": 269, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 257, "to": 270, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 138, "to": 271, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 258, "to": 271, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 138, "to": 272, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 259, "to": 272, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 61, "to": 273, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 260, "to": 274, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 262, "to": 275, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 261, "to": 275, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 263, "to": 276, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 204, "to": 276, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 62, "to": 277, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 264, "to": 278, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 164, "to": 278, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 265, "to": 279, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 202, "to": 279, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 267, "to": 280, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 266, "to": 280, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 250, "to": 280, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 63, "to": 281, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 268, "to": 282, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 64, "to": 283, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 269, "to": 284, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 271, "to": 285, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 270, "to": 285, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 272, "to": 286, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 270, "to": 286, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 65, "to": 287, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 273, "to": 288, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 275, "to": 289, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 274, "to": 289, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 276, "to": 290, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 274, "to": 290, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 66, "to": 291, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 277, "to": 292, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 216, "to": 293, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 280, "to": 293, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 279, "to": 293, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 278, "to": 293, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 281, "to": 294, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 149, "to": 295, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 282, "to": 295, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 67, "to": 296, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 68, "to": 297, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 283, "to": 298, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 271, "to": 299, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 284, "to": 299, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 272, "to": 300, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 284, "to": 300, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 274, "to": 301, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 286, "to": 301, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 254, "to": 301, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 285, "to": 301, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 287, "to": 302, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 235, "to": 303, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 270, "to": 303, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 280, "to": 304, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 181, "to": 304, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 278, "to": 304, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 249, "to": 305, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 242, "to": 305, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 247, "to": 305, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 289, "to": 306, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 288, "to": 306, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 292, "to": 307, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 291, "to": 307, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 290, "to": 307, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 294, "to": 308, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 293, "to": 308, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 211, "to": 308, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 296, "to": 309, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 295, "to": 309, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 297, "to": 310, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 155, "to": 311, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 298, "to": 311, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 156, "to": 312, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 298, "to": 312, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 69, "to": 313, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 70, "to": 314, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 292, "to": 315, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 300, "to": 315, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 254, "to": 315, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 299, "to": 315, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 235, "to": 316, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 284, "to": 316, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 301, "to": 317, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 164, "to": 317, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 303, "to": 318, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 302, "to": 318, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 304, "to": 319, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 280, "to": 319, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 196, "to": 319, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 278, "to": 319, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 71, "to": 320, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 305, "to": 321, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 274, "to": 321, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 308, "to": 322, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 307, "to": 322, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 306, "to": 322, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 310, "to": 323, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 309, "to": 323, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 211, "to": 323, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 163, "to": 324, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 298, "to": 324, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 228, "to": 325, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 312, "to": 325, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 168, "to": 325, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 311, "to": 325, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 313, "to": 326, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 314, "to": 327, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 72, "to": 328, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 315, "to": 329, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 164, "to": 329, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 316, "to": 330, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 302, "to": 330, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 181, "to": 331, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 317, "to": 331, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 318, "to": 332, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 164, "to": 332, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 264, "to": 333, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 200, "to": 333, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 280, "to": 334, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 213, "to": 334, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 278, "to": 334, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 294, "to": 335, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 319, "to": 335, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 211, "to": 335, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 320, "to": 336, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 292, "to": 337, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 291, "to": 337, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 321, "to": 337, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 323, "to": 338, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 322, "to": 338, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 324, "to": 339, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 178, "to": 339, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 280, "to": 340, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 325, "to": 340, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 278, "to": 340, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 265, "to": 341, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 326, "to": 341, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 327, "to": 342, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 203, "to": 342, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 326, "to": 342, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 328, "to": 343, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 73, "to": 344, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 181, "to": 345, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 329, "to": 345, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 330, "to": 346, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 164, "to": 346, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 301, "to": 347, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 200, "to": 347, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 331, "to": 348, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 196, "to": 348, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 317, "to": 348, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 331, "to": 349, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 181, "to": 349, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 332, "to": 349, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 333, "to": 350, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 294, "to": 350, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 335, "to": 351, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 334, "to": 351, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 294, "to": 351, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 338, "to": 352, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 337, "to": 352, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 336, "to": 352, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 340, "to": 353, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 280, "to": 353, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 339, "to": 353, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 278, "to": 353, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 342, "to": 354, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 280, "to": 354, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 341, "to": 354, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 278, "to": 354, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 251, "to": 355, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 343, "to": 355, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 344, "to": 356, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 315, "to": 357, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 200, "to": 357, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 345, "to": 358, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 196, "to": 358, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 329, "to": 358, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 345, "to": 359, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 181, "to": 359, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 346, "to": 359, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 347, "to": 360, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 213, "to": 360, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 317, "to": 360, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 349, "to": 361, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 348, "to": 361, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 196, "to": 361, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 332, "to": 361, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 352, "to": 362, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 351, "to": 362, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 350, "to": 362, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 294, "to": 363, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 353, "to": 363, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 211, "to": 363, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 294, "to": 364, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 354, "to": 364, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 211, "to": 364, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 74, "to": 365, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 355, "to": 366, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 164, "to": 366, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 356, "to": 367, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 266, "to": 367, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 343, "to": 367, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 357, "to": 368, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 213, "to": 368, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 329, "to": 368, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 359, "to": 369, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 358, "to": 369, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 196, "to": 369, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 346, "to": 369, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 318, "to": 370, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 200, "to": 370, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 360, "to": 371, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 213, "to": 371, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 332, "to": 371, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 294, "to": 372, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 361, "to": 372, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 211, "to": 372, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 348, "to": 373, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 211, "to": 373, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 75, "to": 374, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 362, "to": 375, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 241, "to": 375, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 364, "to": 376, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 245, "to": 376, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 363, "to": 376, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 216, "to": 377, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 279, "to": 377, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 317, "to": 377, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 365, "to": 378, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 76, "to": 379, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 77, "to": 380, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 367, "to": 381, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 181, "to": 381, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 366, "to": 381, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 325, "to": 382, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 329, "to": 382, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 330, "to": 383, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 200, "to": 383, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 368, "to": 384, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 213, "to": 384, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 346, "to": 384, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 294, "to": 385, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 369, "to": 385, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 211, "to": 385, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 358, "to": 386, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 211, "to": 386, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 370, "to": 387, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 294, "to": 387, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 372, "to": 388, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 371, "to": 388, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 294, "to": 388, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 360, "to": 389, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 373, "to": 389, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 374, "to": 390, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 352, "to": 391, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 376, "to": 391, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 375, "to": 391, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 216, "to": 392, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 377, "to": 392, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 279, "to": 392, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 332, "to": 392, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 348, "to": 393, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 378, "to": 393, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 78, "to": 394, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 379, "to": 395, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 380, "to": 396, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 381, "to": 397, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 367, "to": 397, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 196, "to": 397, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 366, "to": 397, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 382, "to": 398, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 339, "to": 398, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 329, "to": 398, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 382, "to": 399, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 325, "to": 399, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 346, "to": 399, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 342, "to": 400, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 341, "to": 400, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 329, "to": 400, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 383, "to": 401, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 294, "to": 401, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 385, "to": 402, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 384, "to": 402, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 294, "to": 402, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 368, "to": 403, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 386, "to": 403, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 216, "to": 404, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 279, "to": 404, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 329, "to": 404, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 389, "to": 405, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 388, "to": 405, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 387, "to": 405, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 338, "to": 406, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 391, "to": 406, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 390, "to": 406, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 280, "to": 407, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 294, "to": 407, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 294, "to": 408, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 392, "to": 408, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 211, "to": 408, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 212, "to": 409, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 378, "to": 409, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 274, "to": 410, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 267, "to": 410, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 280, "to": 410, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 393, "to": 410, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 394, "to": 411, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 396, "to": 412, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 369, "to": 412, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 395, "to": 412, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 358, "to": 413, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 395, "to": 413, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 396, "to": 414, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 361, "to": 414, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 395, "to": 414, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 348, "to": 415, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 395, "to": 415, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 212, "to": 416, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 395, "to": 416, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 396, "to": 417, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 319, "to": 417, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 395, "to": 417, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 355, "to": 418, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 200, "to": 418, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 367, "to": 419, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 213, "to": 419, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 366, "to": 419, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 396, "to": 420, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 397, "to": 420, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 395, "to": 420, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 399, "to": 421, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 398, "to": 421, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 339, "to": 421, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 346, "to": 421, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 342, "to": 422, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 400, "to": 422, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 341, "to": 422, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 346, "to": 422, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 403, "to": 423, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 402, "to": 423, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 401, "to": 423, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 216, "to": 424, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 404, "to": 424, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 279, "to": 424, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 346, "to": 424, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 305, "to": 425, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 405, "to": 425, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 247, "to": 425, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 407, "to": 426, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 406, "to": 426, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 275, "to": 427, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 408, "to": 427, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 280, "to": 428, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 409, "to": 428, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 267, "to": 429, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 360, "to": 429, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 410, "to": 429, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 216, "to": 430, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 378, "to": 430, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 220, "to": 431, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 411, "to": 431, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 221, "to": 432, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 411, "to": 432, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 383, "to": 433, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 396, "to": 433, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 412, "to": 434, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 384, "to": 434, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 396, "to": 434, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 368, "to": 435, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 413, "to": 435, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 370, "to": 436, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 396, "to": 436, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 414, "to": 437, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 371, "to": 437, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 396, "to": 437, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 360, "to": 438, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 415, "to": 438, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 226, "to": 439, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 416, "to": 439, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 333, "to": 440, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 396, "to": 440, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 417, "to": 441, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 334, "to": 441, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 396, "to": 441, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 280, "to": 442, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 396, "to": 442, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 367, "to": 443, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 325, "to": 443, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 366, "to": 443, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 418, "to": 444, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 396, "to": 444, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 420, "to": 445, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 419, "to": 445, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 396, "to": 445, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 367, "to": 446, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 396, "to": 446, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 79, "to": 447, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 294, "to": 448, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 421, "to": 448, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 211, "to": 448, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 294, "to": 449, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 422, "to": 449, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 211, "to": 449, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 423, "to": 450, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 241, "to": 450, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 294, "to": 451, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 424, "to": 451, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 211, "to": 451, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 426, "to": 452, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 291, "to": 452, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 425, "to": 452, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 427, "to": 453, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 288, "to": 453, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 429, "to": 454, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 226, "to": 454, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 428, "to": 454, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 228, "to": 455, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 204, "to": 455, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 340, "to": 455, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 430, "to": 455, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 432, "to": 456, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 254, "to": 456, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 431, "to": 456, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 80, "to": 457, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 435, "to": 458, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 434, "to": 458, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 433, "to": 458, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 438, "to": 459, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 437, "to": 459, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 436, "to": 459, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 249, "to": 460, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 439, "to": 460, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 247, "to": 460, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 442, "to": 461, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 441, "to": 461, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 440, "to": 461, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 81, "to": 462, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 443, "to": 463, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 367, "to": 463, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 339, "to": 463, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 366, "to": 463, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 342, "to": 464, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 367, "to": 464, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 341, "to": 464, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 366, "to": 464, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 446, "to": 465, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 445, "to": 465, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 444, "to": 465, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 325, "to": 466, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 317, "to": 466, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 447, "to": 467, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 82, "to": 468, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 450, "to": 469, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 449, "to": 469, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 245, "to": 469, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 448, "to": 469, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 453, "to": 470, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 452, "to": 470, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 451, "to": 470, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 83, "to": 471, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 243, "to": 472, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 454, "to": 472, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 241, "to": 472, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 246, "to": 473, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 245, "to": 473, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 455, "to": 473, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 267, "to": 474, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 456, "to": 474, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 250, "to": 474, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 84, "to": 475, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 457, "to": 476, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 190, "to": 477, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 411, "to": 477, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 85, "to": 478, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 396, "to": 479, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 421, "to": 479, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 395, "to": 479, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 396, "to": 480, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 422, "to": 480, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 395, "to": 480, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 458, "to": 481, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 241, "to": 481, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 460, "to": 482, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 459, "to": 482, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 247, "to": 482, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 396, "to": 483, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 353, "to": 483, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 395, "to": 483, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 396, "to": 484, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 354, "to": 484, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 395, "to": 484, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 461, "to": 485, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 241, "to": 485, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 462, "to": 486, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 396, "to": 487, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 463, "to": 487, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 395, "to": 487, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 396, "to": 488, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 464, "to": 488, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 395, "to": 488, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 465, "to": 489, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 241, "to": 489, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 274, "to": 490, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 356, "to": 490, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 367, "to": 490, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 393, "to": 490, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 360, "to": 491, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 331, "to": 491, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 466, "to": 492, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 325, "to": 492, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 332, "to": 492, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 226, "to": 493, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 467, "to": 493, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 86, "to": 494, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 170, "to": 495, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 204, "to": 495, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 304, "to": 495, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 430, "to": 495, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 170, "to": 496, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 204, "to": 496, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 359, "to": 496, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 430, "to": 496, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 468, "to": 497, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 452, "to": 498, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 469, "to": 498, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 323, "to": 499, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 470, "to": 499, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 471, "to": 500, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 473, "to": 501, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 472, "to": 501, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 476, "to": 502, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 407, "to": 502, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 475, "to": 502, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 474, "to": 502, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 396, "to": 503, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 309, "to": 503, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 378, "to": 503, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 138, "to": 504, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 477, "to": 504, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 478, "to": 505, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 481, "to": 506, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 480, "to": 506, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 245, "to": 506, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 479, "to": 506, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 291, "to": 507, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 482, "to": 507, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 485, "to": 508, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 484, "to": 508, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 245, "to": 508, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 483, "to": 508, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 87, "to": 509, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 219, "to": 510, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 486, "to": 510, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 489, "to": 511, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 488, "to": 511, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 245, "to": 511, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 487, "to": 511, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 460, "to": 512, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 274, "to": 512, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 170, "to": 513, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 204, "to": 513, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 381, "to": 513, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 430, "to": 513, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 356, "to": 514, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 360, "to": 514, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 490, "to": 514, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 263, "to": 515, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 491, "to": 515, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 247, "to": 515, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 228, "to": 516, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 204, "to": 516, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 492, "to": 516, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 430, "to": 516, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 249, "to": 517, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 493, "to": 517, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 247, "to": 517, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 263, "to": 518, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 327, "to": 518, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 494, "to": 519, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 429, "to": 520, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 226, "to": 520, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 495, "to": 520, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 497, "to": 521, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 226, "to": 521, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 496, "to": 521, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 228, "to": 522, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 204, "to": 522, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 399, "to": 522, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 430, "to": 522, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 499, "to": 523, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 498, "to": 523, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 390, "to": 523, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 88, "to": 524, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 502, "to": 525, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 501, "to": 525, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 500, "to": 525, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 280, "to": 526, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 503, "to": 526, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 504, "to": 527, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 234, "to": 527, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 89, "to": 528, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 118, "to": 529, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 505, "to": 529, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 508, "to": 530, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 507, "to": 530, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 506, "to": 530, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 509, "to": 531, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 475, "to": 532, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 510, "to": 532, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 511, "to": 533, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 288, "to": 533, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 292, "to": 534, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 291, "to": 534, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 512, "to": 534, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 514, "to": 535, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 226, "to": 535, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 513, "to": 535, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 228, "to": 536, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 204, "to": 536, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 443, "to": 536, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 430, "to": 536, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 404, "to": 537, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 378, "to": 537, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 515, "to": 538, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 204, "to": 538, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 246, "to": 539, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 245, "to": 539, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 516, "to": 539, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 519, "to": 540, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 518, "to": 540, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 245, "to": 540, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 517, "to": 540, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 243, "to": 541, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 520, "to": 541, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 241, "to": 541, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 243, "to": 542, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 521, "to": 542, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 241, "to": 542, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 246, "to": 543, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 245, "to": 543, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 522, "to": 543, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 499, "to": 544, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 523, "to": 544, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 336, "to": 544, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 524, "to": 545, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 526, "to": 546, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 525, "to": 546, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 527, "to": 547, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 250, "to": 547, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 90, "to": 548, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 528, "to": 549, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 358, "to": 550, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 378, "to": 550, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 134, "to": 551, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 529, "to": 551, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 531, "to": 552, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 530, "to": 552, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 500, "to": 552, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 310, "to": 553, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 532, "to": 553, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 396, "to": 553, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 476, "to": 554, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 446, "to": 554, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 534, "to": 554, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 533, "to": 554, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 243, "to": 555, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 535, "to": 555, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 241, "to": 555, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 246, "to": 556, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 245, "to": 556, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 536, "to": 556, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 292, "to": 557, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 356, "to": 557, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 367, "to": 557, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 537, "to": 557, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 291, "to": 558, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 538, "to": 558, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 540, "to": 559, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 539, "to": 559, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 519, "to": 559, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 473, "to": 560, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 541, "to": 560, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 543, "to": 561, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 542, "to": 561, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 544, "to": 562, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 475, "to": 562, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 502, "to": 563, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 546, "to": 563, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 545, "to": 563, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 476, "to": 564, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 502, "to": 564, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 475, "to": 564, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 547, "to": 564, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 548, "to": 565, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 147, "to": 566, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 549, "to": 566, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 292, "to": 567, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 356, "to": 567, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 367, "to": 567, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 550, "to": 567, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 551, "to": 568, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 154, "to": 568, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 554, "to": 569, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 553, "to": 569, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 552, "to": 569, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 216, "to": 570, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 279, "to": 570, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 180, "to": 570, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 556, "to": 571, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 555, "to": 571, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 476, "to": 572, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 356, "to": 572, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 558, "to": 572, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 557, "to": 572, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 561, "to": 573, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 560, "to": 573, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 291, "to": 573, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 559, "to": 573, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 476, "to": 574, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 563, "to": 574, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 562, "to": 574, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 547, "to": 574, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 564, "to": 575, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 164, "to": 575, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 103, "to": 576, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 565, "to": 576, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 566, "to": 577, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 160, "to": 577, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 367, "to": 578, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 409, "to": 578, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 398, "to": 579, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 378, "to": 579, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 400, "to": 580, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 378, "to": 580, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 356, "to": 581, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 368, "to": 581, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 567, "to": 581, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 568, "to": 582, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 178, "to": 582, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 569, "to": 583, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 296, "to": 583, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 570, "to": 584, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 378, "to": 584, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 544, "to": 585, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 532, "to": 585, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 572, "to": 586, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 571, "to": 586, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 288, "to": 586, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 531, "to": 587, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 573, "to": 587, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 500, "to": 587, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 574, "to": 588, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 164, "to": 588, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 335, "to": 589, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 181, "to": 589, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 575, "to": 589, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 576, "to": 590, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 104, "to": 590, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 91, "to": 591, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 176, "to": 592, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 577, "to": 592, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 514, "to": 593, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 226, "to": 593, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 578, "to": 593, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 514, "to": 594, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 247, "to": 594, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 292, "to": 595, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 356, "to": 595, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 367, "to": 595, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 579, "to": 595, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 292, "to": 596, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 356, "to": 596, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 367, "to": 596, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 580, "to": 596, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 581, "to": 597, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 241, "to": 597, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 576, "to": 598, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 125, "to": 598, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 576, "to": 599, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 126, "to": 599, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 181, "to": 600, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 583, "to": 600, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 582, "to": 600, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 587, "to": 601, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 586, "to": 601, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 585, "to": 601, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 584, "to": 601, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 589, "to": 602, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 196, "to": 602, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 588, "to": 602, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 574, "to": 603, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 200, "to": 603, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 110, "to": 604, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 590, "to": 604, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 591, "to": 605, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 592, "to": 606, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 189, "to": 606, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 243, "to": 607, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 593, "to": 607, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 241, "to": 607, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 540, "to": 608, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 594, "to": 608, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 597, "to": 609, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 596, "to": 609, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 245, "to": 609, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 595, "to": 609, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 92, "to": 610, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 199, "to": 611, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 598, "to": 611, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 199, "to": 612, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 599, "to": 612, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 601, "to": 613, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 600, "to": 613, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 180, "to": 613, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 363, "to": 614, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 325, "to": 614, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 575, "to": 614, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 602, "to": 615, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 378, "to": 615, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 93, "to": 616, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 603, "to": 617, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 213, "to": 617, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 588, "to": 617, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 604, "to": 618, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 505, "to": 618, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 551, "to": 619, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 298, "to": 619, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 279, "to": 620, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 605, "to": 620, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 219, "to": 621, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 606, "to": 621, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 556, "to": 622, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 607, "to": 622, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 609, "to": 623, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 476, "to": 623, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 291, "to": 623, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 608, "to": 623, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 158, "to": 624, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 505, "to": 624, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 159, "to": 625, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 505, "to": 625, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 196, "to": 626, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 166, "to": 626, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 181, "to": 627, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 166, "to": 627, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 610, "to": 628, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 467, "to": 629, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 611, "to": 629, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 298, "to": 629, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 467, "to": 630, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 612, "to": 630, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 298, "to": 630, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 325, "to": 631, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 605, "to": 631, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 226, "to": 632, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 613, "to": 632, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 614, "to": 633, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 339, "to": 633, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 588, "to": 633, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 342, "to": 634, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 341, "to": 634, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 588, "to": 634, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 476, "to": 635, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 585, "to": 635, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 615, "to": 635, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 616, "to": 636, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 587, "to": 637, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 617, "to": 637, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 495, "to": 637, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 377, "to": 638, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 378, "to": 638, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 199, "to": 639, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 618, "to": 639, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 619, "to": 640, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 178, "to": 640, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 620, "to": 641, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 378, "to": 641, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 229, "to": 642, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 532, "to": 642, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 430, "to": 642, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 152, "to": 643, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 621, "to": 643, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 572, "to": 644, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 622, "to": 644, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 288, "to": 644, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 367, "to": 645, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 503, "to": 645, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 336, "to": 646, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 623, "to": 646, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 390, "to": 646, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 625, "to": 647, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 168, "to": 647, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 624, "to": 647, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 627, "to": 648, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 626, "to": 648, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 378, "to": 648, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 355, "to": 649, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 628, "to": 649, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 631, "to": 650, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 228, "to": 650, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 630, "to": 650, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 168, "to": 650, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 629, "to": 650, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 490, "to": 651, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 585, "to": 651, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 393, "to": 651, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 585, "to": 652, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 409, "to": 652, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 249, "to": 653, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 632, "to": 653, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 247, "to": 653, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 633, "to": 654, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 378, "to": 654, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 634, "to": 655, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 378, "to": 655, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 617, "to": 656, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 635, "to": 656, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 637, "to": 657, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 636, "to": 657, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 274, "to": 658, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 267, "to": 658, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 280, "to": 658, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 638, "to": 658, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 566, "to": 659, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 239, "to": 659, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 566, "to": 660, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 240, "to": 660, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 467, "to": 661, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 639, "to": 661, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 298, "to": 661, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 325, "to": 662, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 583, "to": 662, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 640, "to": 662, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 587, "to": 663, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 586, "to": 663, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 642, "to": 663, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 585, "to": 663, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 641, "to": 663, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 274, "to": 664, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 643, "to": 664, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 234, "to": 664, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 646, "to": 665, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 645, "to": 665, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 644, "to": 665, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 551, "to": 666, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 252, "to": 666, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 327, "to": 667, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 647, "to": 667, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 326, "to": 667, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 342, "to": 668, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 395, "to": 668, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 649, "to": 669, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 648, "to": 669, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 650, "to": 670, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 636, "to": 670, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 325, "to": 671, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 180, "to": 671, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 360, "to": 672, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 651, "to": 672, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 226, "to": 673, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 652, "to": 673, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 646, "to": 674, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 653, "to": 674, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 594, "to": 674, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 476, "to": 675, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 585, "to": 675, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 654, "to": 675, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 476, "to": 676, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 585, "to": 676, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 655, "to": 676, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 656, "to": 677, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 241, "to": 677, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 658, "to": 678, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 601, "to": 678, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 657, "to": 678, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 323, "to": 679, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 532, "to": 679, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 503, "to": 679, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 567, "to": 680, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 585, "to": 680, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 550, "to": 680, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 152, "to": 681, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 659, "to": 681, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 152, "to": 682, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 660, "to": 682, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 650, "to": 683, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 661, "to": 683, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 178, "to": 683, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 663, "to": 684, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 662, "to": 684, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 605, "to": 684, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 665, "to": 685, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 664, "to": 685, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 343, "to": 685, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 666, "to": 686, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 326, "to": 686, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 246, "to": 687, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 668, "to": 687, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 296, "to": 687, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 667, "to": 687, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 669, "to": 688, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 367, "to": 688, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 648, "to": 688, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 201, "to": 689, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 213, "to": 689, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 166, "to": 689, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 671, "to": 690, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 670, "to": 690, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 243, "to": 691, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 248, "to": 691, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 241, "to": 691, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 368, "to": 692, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 496, "to": 692, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 672, "to": 693, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 247, "to": 693, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 249, "to": 694, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 674, "to": 694, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 673, "to": 694, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 247, "to": 694, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 677, "to": 695, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 676, "to": 695, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 245, "to": 695, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 675, "to": 695, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 476, "to": 696, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 679, "to": 696, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 678, "to": 696, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 368, "to": 697, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 680, "to": 697, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 497, "to": 698, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 681, "to": 698, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 284, "to": 698, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 497, "to": 699, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 682, "to": 699, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 284, "to": 699, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 315, "to": 700, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 628, "to": 700, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 684, "to": 701, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 583, "to": 701, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 683, "to": 701, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 645, "to": 702, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 649, "to": 702, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 503, "to": 702, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 586, "to": 703, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 646, "to": 703, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 685, "to": 703, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 503, "to": 703, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 246, "to": 704, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 687, "to": 704, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 583, "to": 704, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 686, "to": 704, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 689, "to": 705, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 688, "to": 705, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 691, "to": 706, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 245, "to": 706, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 690, "to": 706, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 692, "to": 707, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 636, "to": 707, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 694, "to": 708, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 693, "to": 708, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 696, "to": 709, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 695, "to": 709, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 390, "to": 709, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 697, "to": 710, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 636, "to": 710, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 700, "to": 711, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 292, "to": 711, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 699, "to": 711, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 254, "to": 711, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 698, "to": 711, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 701, "to": 712, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 636, "to": 712, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 586, "to": 713, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 646, "to": 713, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 703, "to": 713, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 702, "to": 713, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 646, "to": 714, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 665, "to": 714, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 704, "to": 714, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 366, "to": 714, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 703, "to": 715, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 705, "to": 715, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 587, "to": 716, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 586, "to": 716, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 585, "to": 716, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 638, "to": 716, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 246, "to": 717, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 706, "to": 717, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 204, "to": 717, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 497, "to": 718, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 601, "to": 718, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 707, "to": 718, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 709, "to": 719, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 291, "to": 719, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 708, "to": 719, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 711, "to": 720, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 710, "to": 720, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 586, "to": 721, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 646, "to": 721, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 713, "to": 721, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 712, "to": 721, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 246, "to": 722, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 646, "to": 722, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 714, "to": 722, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 503, "to": 722, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 243, "to": 723, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 715, "to": 723, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 241, "to": 723, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 476, "to": 724, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 291, "to": 724, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 674, "to": 724, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 717, "to": 725, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 716, "to": 725, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 587, "to": 726, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 719, "to": 726, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 718, "to": 726, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 719, "to": 727, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 720, "to": 727, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 723, "to": 728, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 722, "to": 728, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 245, "to": 728, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 721, "to": 728, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 665, "to": 729, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 724, "to": 729, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 292, "to": 729, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 726, "to": 730, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 725, "to": 730, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 288, "to": 730, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 545, "to": 731, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 727, "to": 731, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 336, "to": 731, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 504, "to": 732, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 284, "to": 732, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 729, "to": 733, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 728, "to": 733, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 288, "to": 733, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 731, "to": 734, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 679, "to": 734, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 730, "to": 734, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 592, "to": 735, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 411, "to": 735, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 604, "to": 736, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 117, "to": 736, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 732, "to": 737, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 302, "to": 737, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 733, "to": 738, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 562, "to": 738, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 733, "to": 739, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 734, "to": 739, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 152, "to": 740, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 735, "to": 740, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 149, "to": 741, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 736, "to": 741, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 504, "to": 742, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 270, "to": 742, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 315, "to": 743, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 562, "to": 743, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 737, "to": 743, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 739, "to": 744, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 738, "to": 744, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 737, "to": 744, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 497, "to": 745, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 740, "to": 745, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 284, "to": 745, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 581, "to": 746, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 636, "to": 746, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 199, "to": 747, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 741, "to": 747, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 733, "to": 748, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 585, "to": 748, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 742, "to": 749, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 302, "to": 749, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 743, "to": 750, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 164, "to": 750, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 744, "to": 751, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 164, "to": 751, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 744, "to": 752, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 200, "to": 752, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 711, "to": 753, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 745, "to": 753, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 302, "to": 753, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 356, "to": 754, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 711, "to": 754, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 746, "to": 754, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 170, "to": 755, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 747, "to": 755, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 252, "to": 755, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 587, "to": 756, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 748, "to": 756, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 430, "to": 756, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 301, "to": 757, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 562, "to": 757, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 749, "to": 757, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 567, "to": 758, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 181, "to": 758, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 750, "to": 758, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 752, "to": 759, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 213, "to": 759, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 751, "to": 759, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 743, "to": 760, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 628, "to": 760, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 754, "to": 761, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 562, "to": 761, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 753, "to": 761, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 663, "to": 762, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 756, "to": 762, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 755, "to": 762, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 202, "to": 762, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 739, "to": 763, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 738, "to": 763, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 749, "to": 763, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 757, "to": 764, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 164, "to": 764, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 759, "to": 765, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 758, "to": 765, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 754, "to": 766, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 761, "to": 766, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 760, "to": 766, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 762, "to": 767, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 378, "to": 767, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 279, "to": 768, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 166, "to": 768, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 763, "to": 769, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 164, "to": 769, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 490, "to": 770, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 181, "to": 770, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 764, "to": 770, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 766, "to": 771, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 765, "to": 771, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 636, "to": 771, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 587, "to": 772, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 748, "to": 772, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 767, "to": 772, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 756, "to": 773, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 228, "to": 773, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 768, "to": 774, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 378, "to": 774, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 770, "to": 775, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 196, "to": 775, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 769, "to": 775, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 763, "to": 776, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 200, "to": 776, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 771, "to": 777, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 196, "to": 777, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 751, "to": 777, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 772, "to": 778, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 241, "to": 778, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 246, "to": 779, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 245, "to": 779, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 773, "to": 779, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 587, "to": 780, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 586, "to": 780, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 642, "to": 780, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 585, "to": 780, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 774, "to": 780, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 748, "to": 781, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 775, "to": 781, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 378, "to": 781, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 776, "to": 782, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 213, "to": 782, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 769, "to": 782, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 748, "to": 783, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 777, "to": 783, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 378, "to": 783, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 779, "to": 784, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 778, "to": 784, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 772, "to": 785, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 636, "to": 785, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 587, "to": 786, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 761, "to": 786, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 780, "to": 786, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 782, "to": 787, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 781, "to": 787, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 759, "to": 788, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 783, "to": 788, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 772, "to": 789, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 247, "to": 789, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 784, "to": 790, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 288, "to": 790, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 786, "to": 791, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 785, "to": 791, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 694, "to": 792, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 787, "to": 792, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 247, "to": 792, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 766, "to": 793, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 788, "to": 793, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 636, "to": 793, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 733, "to": 794, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 679, "to": 794, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 791, "to": 795, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 790, "to": 795, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 678, "to": 795, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 291, "to": 795, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 789, "to": 795, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 793, "to": 796, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 709, "to": 796, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 291, "to": 796, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 792, "to": 796, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 795, "to": 797, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 794, "to": 797, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 569, "to": 798, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 309, "to": 798, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 797, "to": 799, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 796, "to": 799, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 336, "to": 799, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 342, "to": 800, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 341, "to": 800, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 317, "to": 800, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 795, "to": 801, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 583, "to": 801, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 718, "to": 802, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 600, "to": 802, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 329, "to": 802, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 716, "to": 803, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 600, "to": 803, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 317, "to": 803, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 795, "to": 804, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 798, "to": 804, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 799, "to": 805, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 545, "to": 805, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 170, "to": 806, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 639, "to": 806, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 252, "to": 806, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 246, "to": 807, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 687, "to": 807, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 296, "to": 807, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 686, "to": 807, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 800, "to": 808, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 395, "to": 808, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 684, "to": 809, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 801, "to": 809, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 683, "to": 809, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 368, "to": 810, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 802, "to": 810, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 360, "to": 811, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 803, "to": 811, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 805, "to": 812, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 804, "to": 812, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 667, "to": 813, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 806, "to": 813, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 326, "to": 813, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 704, "to": 814, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 605, "to": 814, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 342, "to": 815, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 341, "to": 815, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 180, "to": 815, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 808, "to": 816, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 807, "to": 816, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 317, "to": 816, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 809, "to": 817, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 636, "to": 817, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 601, "to": 818, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 662, "to": 818, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 180, "to": 818, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 799, "to": 819, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 500, "to": 819, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 810, "to": 820, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 636, "to": 820, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 653, "to": 821, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 811, "to": 821, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 247, "to": 821, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 804, "to": 822, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 378, "to": 822, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 264, "to": 823, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 628, "to": 823, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 812, "to": 824, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 378, "to": 824, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 563, "to": 825, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 664, "to": 825, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 250, "to": 825, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 814, "to": 826, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 246, "to": 826, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 583, "to": 826, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 813, "to": 826, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 815, "to": 827, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 395, "to": 827, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 816, "to": 828, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 704, "to": 828, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 332, "to": 828, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 818, "to": 829, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 817, "to": 829, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 819, "to": 830, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 245, "to": 830, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 711, "to": 831, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 820, "to": 831, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 291, "to": 832, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 821, "to": 832, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 94, "to": 833, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 563, "to": 834, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 823, "to": 834, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 822, "to": 834, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 825, "to": 835, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 824, "to": 835, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 826, "to": 836, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 636, "to": 836, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 827, "to": 837, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 807, "to": 837, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 180, "to": 837, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 246, "to": 838, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 828, "to": 838, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 503, "to": 838, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 243, "to": 839, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 632, "to": 839, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 241, "to": 839, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 805, "to": 840, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 830, "to": 840, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 829, "to": 840, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 731, "to": 841, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 832, "to": 841, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 831, "to": 841, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 833, "to": 842, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 805, "to": 843, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 835, "to": 843, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 834, "to": 843, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 838, "to": 844, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 837, "to": 844, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 836, "to": 844, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 805, "to": 845, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 840, "to": 845, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 839, "to": 845, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 841, "to": 846, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 531, "to": 846, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 843, "to": 847, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 842, "to": 847, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 805, "to": 848, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 845, "to": 848, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 844, "to": 848, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 824, "to": 849, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 219, "to": 849, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 799, "to": 850, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 846, "to": 850, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 848, "to": 851, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 847, "to": 851, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 849, "to": 852, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 189, "to": 852, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 851, "to": 853, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 850, "to": 853, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 852, "to": 854, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 176, "to": 854, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 853, "to": 855, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 302, "to": 855, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 855, "to": 856, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 854, "to": 856, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 848, "to": 857, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 270, "to": 857, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 857, "to": 858, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 856, "to": 858, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 175, "to": 858, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 805, "to": 859, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 801, "to": 859, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 799, "to": 860, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 531, "to": 860, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 858, "to": 861, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 138, "to": 861, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 859, "to": 862, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 667, "to": 862, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 860, "to": 863, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 291, "to": 863, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 848, "to": 864, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 497, "to": 864, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 861, "to": 865, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 164, "to": 865, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 862, "to": 866, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 859, "to": 866, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 686, "to": 866, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 325, "to": 867, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 859, "to": 867, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 640, "to": 867, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 864, "to": 868, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 851, "to": 868, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 863, "to": 868, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 866, "to": 869, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 865, "to": 869, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 867, "to": 870, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 865, "to": 870, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 868, "to": 871, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 519, "to": 871, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 871, "to": 872, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 870, "to": 872, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 869, "to": 872, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 830, "to": 872, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}]}, {"problem_file": "PUZ135_1.p", "status": "# SZS status Theorem", "nodes": [{"id": 0, "label": "0", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": null, "raw_label": "tff(laser_etching_info, axiom, (entry_of(b)!=laser_etching&![X1:knowheyan]:((entry_of(X1)=laser_etching=>(X1!=a&place_of(X1)!=second)))),\\nfile('Problems/PUZ/PUZ135_1.p', laser_etching_info)).\n", "fillcolor": "forestgreen"}, {"id": 1, "label": "1", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": null, "raw_label": "tff(collage_info, axiom, ![X1:knowheyan]:((entry_of(X1)=collage_painting=>(((X1!=a&X1!=c)&place_of(X1)!=first)&place_of(X1)!=second))),\\nfile('Problems/PUZ/PUZ135_1.p', collage_info)).\n", "fillcolor": "forestgreen"}, {"id": 2, "label": "2", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": null, "raw_label": "tff(only_entries, axiom, ![X2:entry]:((X2=collage_painting|X2=holography|X2=laser_etching|X2=reconstituted_materials_sculpture)),\\nfile('Problems/PUZ/PUZ135_1.p', only_entries)).\n", "fillcolor": "forestgreen"}, {"id": 3, "label": "3", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": null, "raw_label": "tff(a_info, axiom, (place_of(a)!=first&entry_of(a)!=holography),\\nfile('Problems/PUZ/PUZ135_1.p', a_info)).\n", "fillcolor": "forestgreen"}, {"id": 4, "label": "4", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": null, "raw_label": "tff(only_knowheyans, axiom, ![X1:knowheyan]:((X1=a|X1=b|X1=c|X1=d)),\\nfile('Problems/PUZ/PUZ135_1.p', only_knowheyans)).\n", "fillcolor": "forestgreen"}, {"id": 5, "label": "5", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": null, "raw_label": "tff(a_not_d, axiom, a!=d,\\nfile('Problems/PUZ/PUZ135_1.p', a_not_d)).\n", "fillcolor": "forestgreen"}, {"id": 6, "label": "6", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": null, "raw_label": "tff(a_not_c, axiom, a!=c,\\nfile('Problems/PUZ/PUZ135_1.p', a_not_c)).\n", "fillcolor": "forestgreen"}, {"id": 7, "label": "7", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": null, "raw_label": "tff(a_not_b, axiom, a!=b,\\nfile('Problems/PUZ/PUZ135_1.p', a_not_b)).\n", "fillcolor": "forestgreen"}, {"id": 8, "label": "8", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": null, "raw_label": "tff(unique_entrys, axiom, ![X1:knowheyan, X4:knowheyan]:((X1!=X4=>entry_of(X1)!=entry_of(X4))),\\nfile('Problems/PUZ/PUZ135_1.p', unique_entrys)).\n", "fillcolor": "forestgreen"}, {"id": 9, "label": "9", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": null, "raw_label": "tff(fourth_info2, axiom, (place_of(b)!=fourth&![X1:knowheyan]:((entry_of(X1)=laser_etching=>place_of(X1)!=fourth))),\\nfile('Problems/PUZ/PUZ135_1.p', fourth_info2)).\n", "fillcolor": "forestgreen"}, {"id": 10, "label": "10", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": null, "raw_label": "tff(third_info, axiom, (place_of(b)!=third&place_of(c)!=third),\\nfile('Problems/PUZ/PUZ135_1.p', third_info)).\n", "fillcolor": "forestgreen"}, {"id": 11, "label": "11", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": null, "raw_label": "tff(only_places, axiom, ![X3:place]:((X3=first|X3=second|X3=third|X3=fourth)),\\nfile('Problems/PUZ/PUZ135_1.p', only_places)).\n", "fillcolor": "forestgreen"}, {"id": 12, "label": "12", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": null, "raw_label": "tff(c_not_d, axiom, c!=d,\\nfile('Problems/PUZ/PUZ135_1.p', c_not_d)).\n", "fillcolor": "forestgreen"}, {"id": 13, "label": "13", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": null, "raw_label": "tff(entries, conjecture, (((entry_of(a)=reconstituted_materials_sculpture&entry_of(b)=holography)&entry_of(c)=laser_etching)&entry_of(d)=collage_painting),\\nfile('Problems/PUZ/PUZ135_1.p', entries)).\n", "fillcolor": "firebrick1"}, {"id": 14, "label": "14", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": null, "raw_label": "tff(b_not_d, axiom, b!=d,\\nfile('Problems/PUZ/PUZ135_1.p', b_not_d)).\n", "fillcolor": "forestgreen"}, {"id": 15, "label": "15", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": null, "raw_label": "tff(b_not_c, axiom, b!=c,\\nfile('Problems/PUZ/PUZ135_1.p', b_not_c)).\n", "fillcolor": "forestgreen"}, {"id": 16, "label": "16", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "fof_simplification", "raw_label": "tff(c_0_16, plain, (entry_of(b)!=laser_etching&![X1:knowheyan]:((entry_of(X1)=laser_etching=>(X1!=a&place_of(X1)!=second)))),\\ninference(fof_simplification,[status(thm)],[laser_etching_info])).\n", "fillcolor": "palegreen"}, {"id": 17, "label": "17", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "fof_simplification", "raw_label": "tff(c_0_17, plain, ![X1:knowheyan]:((entry_of(X1)=collage_painting=>(((X1!=a&X1!=c)&place_of(X1)!=first)&place_of(X1)!=second))),\\ninference(fof_simplification,[status(thm)],[collage_info])).\n", "fillcolor": "palegreen"}, {"id": 18, "label": "18", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "distribute", "raw_label": "tff(c_0_18, plain, ![X15:knowheyan]:((entry_of(b)!=laser_etching&((X15!=a|entry_of(X15)!=laser_etching)&(place_of(X15)!=second|entry_of(X15)!=laser_etching)))),\\ninference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])])])])).\n", "fillcolor": "palegreen"}, {"id": 19, "label": "19", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "variable_rename", "raw_label": "tff(c_0_19, plain, ![X6:entry]:((X6=collage_painting|X6=holography|X6=laser_etching|X6=reconstituted_materials_sculpture)),\\ninference(variable_rename,[status(thm)],[only_entries])).\n", "fillcolor": "palegreen"}, {"id": 20, "label": "20", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "fof_simplification", "raw_label": "tff(c_0_20, plain, (place_of(a)!=first&entry_of(a)!=holography),\\ninference(fof_simplification,[status(thm)],[a_info])).\n", "fillcolor": "palegreen"}, {"id": 21, "label": "21", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "distribute", "raw_label": "tff(c_0_21, plain, ![X13:knowheyan]:(((((X13!=a|entry_of(X13)!=collage_painting)&(X13!=c|entry_of(X13)!=collage_painting))&(place_of(X13)!=first|entry_of(X13)!=collage_painting))&(place_of(X13)!=second|entry_of(X13)!=collage_painting))),\\ninference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])])])).\n", "fillcolor": "palegreen"}, {"id": 22, "label": "22", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "split_conjunct", "raw_label": "tcf(c_0_22, plain, ![X1:knowheyan]:((X1!=a|entry_of(X1)!=laser_etching)),\\ninference(split_conjunct,[status(thm)],[c_0_18])).\n", "fillcolor": "palegreen"}, {"id": 23, "label": "23", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "split_conjunct", "raw_label": "tcf(c_0_23, plain, ![X2:entry]:((X2=collage_painting|X2=holography|X2=laser_etching|X2=reconstituted_materials_sculpture)),\\ninference(split_conjunct,[status(thm)],[c_0_19])).\n", "fillcolor": "palegreen"}, {"id": 24, "label": "24", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "fof_nnf", "raw_label": "tff(c_0_24, plain, (place_of(a)!=first&entry_of(a)!=holography),\\ninference(fof_nnf,[status(thm)],[c_0_20])).\n", "fillcolor": "palegreen"}, {"id": 25, "label": "25", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "split_conjunct", "raw_label": "tcf(c_0_25, plain, ![X1:knowheyan]:((X1!=a|entry_of(X1)!=collage_painting)),\\ninference(split_conjunct,[status(thm)],[c_0_21])).\n", "fillcolor": "palegreen"}, {"id": 26, "label": "26", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "variable_rename", "raw_label": "tff(c_0_26, plain, ![X5:knowheyan]:((X5=a|X5=b|X5=c|X5=d)),\\ninference(variable_rename,[status(thm)],[only_knowheyans])).\n", "fillcolor": "palegreen"}, {"id": 27, "label": "27", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "fof_simplification", "raw_label": "tff(c_0_27, plain, a!=d,\\ninference(fof_simplification,[status(thm)],[a_not_d])).\n", "fillcolor": "palegreen"}, {"id": 28, "label": "28", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "fof_simplification", "raw_label": "tff(c_0_28, plain, a!=c,\\ninference(fof_simplification,[status(thm)],[a_not_c])).\n", "fillcolor": "palegreen"}, {"id": 29, "label": "29", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "fof_simplification", "raw_label": "tff(c_0_29, plain, a!=b,\\ninference(fof_simplification,[status(thm)],[a_not_b])).\n", "fillcolor": "palegreen"}, {"id": 30, "label": "30", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "fof_simplification", "raw_label": "tff(c_0_30, plain, ![X1:knowheyan, X4:knowheyan]:((X1!=X4=>entry_of(X1)!=entry_of(X4))),\\ninference(fof_simplification,[status(thm)],[unique_entrys])).\n", "fillcolor": "palegreen"}, {"id": 31, "label": "31", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "er", "raw_label": "tcf(c_0_31, plain, entry_of(a)!=laser_etching,\\ninference(er,[status(thm)],[c_0_22])).\n", "fillcolor": "palegreen"}, {"id": 32, "label": "32", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "spm", "raw_label": "tcf(c_0_32, plain, ![X6:entry, X2:entry]:((X2=laser_etching|X2=holography|X2=collage_painting|X6=laser_etching|X6=holography|X6=collage_painting|X6=X2)),\\ninference(spm,[status(thm)],[c_0_23, c_0_23])).\n", "fillcolor": "palegreen"}, {"id": 33, "label": "33", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "split_conjunct", "raw_label": "tcf(c_0_33, plain, entry_of(a)!=holography,\\ninference(split_conjunct,[status(thm)],[c_0_24])).\n", "fillcolor": "palegreen"}, {"id": 34, "label": "34", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "er", "raw_label": "tcf(c_0_34, plain, entry_of(a)!=collage_painting,\\ninference(er,[status(thm)],[c_0_25])).\n", "fillcolor": "palegreen"}, {"id": 35, "label": "35", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "split_conjunct", "raw_label": "tcf(c_0_35, plain, ![X1:knowheyan]:((X1=a|X1=b|X1=c|X1=d)),\\ninference(split_conjunct,[status(thm)],[c_0_26])).\n", "fillcolor": "palegreen"}, {"id": 36, "label": "36", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "fof_nnf", "raw_label": "tff(c_0_36, plain, a!=d,\\ninference(fof_nnf,[status(thm)],[c_0_27])).\n", "fillcolor": "palegreen"}, {"id": 37, "label": "37", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "fof_nnf", "raw_label": "tff(c_0_37, plain, a!=c,\\ninference(fof_nnf,[status(thm)],[c_0_28])).\n", "fillcolor": "palegreen"}, {"id": 38, "label": "38", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "fof_nnf", "raw_label": "tff(c_0_38, plain, a!=b,\\ninference(fof_nnf,[status(thm)],[c_0_29])).\n", "fillcolor": "palegreen"}, {"id": 39, "label": "39", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "fof_simplification", "raw_label": "tff(c_0_39, plain, (place_of(b)!=fourth&![X1:knowheyan]:((entry_of(X1)=laser_etching=>place_of(X1)!=fourth))),\\ninference(fof_simplification,[status(thm)],[fourth_info2])).\n", "fillcolor": "palegreen"}, {"id": 40, "label": "40", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "fof_simplification", "raw_label": "tff(c_0_40, plain, (place_of(b)!=third&place_of(c)!=third),\\ninference(fof_simplification,[status(thm)],[third_info])).\n", "fillcolor": "palegreen"}, {"id": 41, "label": "41", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "fof_nnf", "raw_label": "tff(c_0_41, plain, ![X8:knowheyan, X9:knowheyan]:((X8=X9|entry_of(X8)!=entry_of(X9))),\\ninference(fof_nnf,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_30])])])).\n", "fillcolor": "palegreen"}, {"id": 42, "label": "42", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "sr", "raw_label": "tcf(c_0_42, plain, ![X2:entry]:((entry_of(a)=X2|X2=collage_painting|X2=holography|X2=laser_etching)),\\ninference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_33]), c_0_34])).\n", "fillcolor": "palegreen"}, {"id": 43, "label": "43", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "spm", "raw_label": "tcf(c_0_43, plain, ![X4:knowheyan, X1:knowheyan]:((X1=d|X1=c|X1=b|X4=d|X4=c|X4=b|X4=X1)),\\ninference(spm,[status(thm)],[c_0_35, c_0_35])).\n", "fillcolor": "palegreen"}, {"id": 44, "label": "44", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "split_conjunct", "raw_label": "tcf(c_0_44, plain, a!=d,\\ninference(split_conjunct,[status(thm)],[c_0_36])).\n", "fillcolor": "palegreen"}, {"id": 45, "label": "45", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "split_conjunct", "raw_label": "tcf(c_0_45, plain, a!=c,\\ninference(split_conjunct,[status(thm)],[c_0_37])).\n", "fillcolor": "palegreen"}, {"id": 46, "label": "46", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "split_conjunct", "raw_label": "tcf(c_0_46, plain, a!=b,\\ninference(split_conjunct,[status(thm)],[c_0_38])).\n", "fillcolor": "palegreen"}, {"id": 47, "label": "47", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "fof_nnf", "raw_label": "tff(c_0_47, plain, ![X14:knowheyan]:((place_of(b)!=fourth&(entry_of(X14)!=laser_etching|place_of(X14)!=fourth))),\\ninference(fof_nnf,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_39])])])])).\n", "fillcolor": "palegreen"}, {"id": 48, "label": "48", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "variable_rename", "raw_label": "tff(c_0_48, plain, ![X7:place]:((X7=first|X7=second|X7=third|X7=fourth)),\\ninference(variable_rename,[status(thm)],[only_places])).\n", "fillcolor": "palegreen"}, {"id": 49, "label": "49", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "fof_nnf", "raw_label": "tff(c_0_49, plain, (place_of(b)!=third&place_of(c)!=third),\\ninference(fof_nnf,[status(thm)],[c_0_40])).\n", "fillcolor": "palegreen"}, {"id": 50, "label": "50", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "split_conjunct", "raw_label": "tcf(c_0_50, plain, ![X1:knowheyan, X4:knowheyan]:((X1=X4|entry_of(X1)!=entry_of(X4))),\\ninference(split_conjunct,[status(thm)],[c_0_41])).\n", "fillcolor": "palegreen"}, {"id": 51, "label": "51", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "sr", "raw_label": "tcf(c_0_51, plain, ![X1:knowheyan, X2:entry]:((entry_of(X1)=X2|X1=b|X1=c|X1=d|X2=laser_etching|X2=holography|X2=collage_painting)),\\ninference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_44]), c_0_45]), c_0_46])).\n", "fillcolor": "palegreen"}, {"id": 52, "label": "52", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "split_conjunct", "raw_label": "tcf(c_0_52, plain, ![X1:knowheyan]:((X1!=c|entry_of(X1)!=collage_painting)),\\ninference(split_conjunct,[status(thm)],[c_0_21])).\n", "fillcolor": "palegreen"}, {"id": 53, "label": "53", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "split_conjunct", "raw_label": "tcf(c_0_53, plain, place_of(b)!=fourth,\\ninference(split_conjunct,[status(thm)],[c_0_47])).\n", "fillcolor": "palegreen"}, {"id": 54, "label": "54", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "split_conjunct", "raw_label": "tcf(c_0_54, plain, ![X3:place]:((X3=first|X3=second|X3=third|X3=fourth)),\\ninference(split_conjunct,[status(thm)],[c_0_48])).\n", "fillcolor": "palegreen"}, {"id": 55, "label": "55", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "split_conjunct", "raw_label": "tcf(c_0_55, plain, place_of(b)!=third,\\ninference(split_conjunct,[status(thm)],[c_0_49])).\n", "fillcolor": "palegreen"}, {"id": 56, "label": "56", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "er", "raw_label": "tcf(c_0_56, plain, ![X4:knowheyan, X1:knowheyan]:((entry_of(X1)=collage_painting|entry_of(X1)=holography|entry_of(X1)=laser_etching|X4=d|X4=c|X4=b|X4=X1)),\\ninference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51])])).\n", "fillcolor": "palegreen"}, {"id": 57, "label": "57", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "er", "raw_label": "tcf(c_0_57, plain, entry_of(c)!=collage_painting,\\ninference(er,[status(thm)],[c_0_52])).\n", "fillcolor": "palegreen"}, {"id": 58, "label": "58", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "fof_simplification", "raw_label": "tff(c_0_58, plain, c!=d,\\ninference(fof_simplification,[status(thm)],[c_not_d])).\n", "fillcolor": "palegreen"}, {"id": 59, "label": "59", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "assume_negation", "raw_label": "tff(c_0_59, negated_conjecture, ~((((entry_of(a)=reconstituted_materials_sculpture&entry_of(b)=holography)&entry_of(c)=laser_etching)&entry_of(d)=collage_painting)),\\ninference(assume_negation,[status(cth)],[entries])).\n", "fillcolor": "lightskyblue1"}, {"id": 60, "label": "60", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "split_conjunct", "raw_label": "tcf(c_0_60, plain, ![X1:knowheyan]:((place_of(X1)!=second|entry_of(X1)!=collage_painting)),\\ninference(split_conjunct,[status(thm)],[c_0_21])).\n", "fillcolor": "palegreen"}, {"id": 61, "label": "61", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "sr", "raw_label": "tcf(c_0_61, plain, (place_of(b)=second|place_of(b)=first),\\ninference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_55])).\n", "fillcolor": "palegreen"}, {"id": 62, "label": "62", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "split_conjunct", "raw_label": "tcf(c_0_62, plain, ![X1:knowheyan]:((place_of(X1)!=first|entry_of(X1)!=collage_painting)),\\ninference(split_conjunct,[status(thm)],[c_0_21])).\n", "fillcolor": "palegreen"}, {"id": 63, "label": "63", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "sr", "raw_label": "tcf(c_0_63, plain, ![X1:knowheyan]:((entry_of(c)=laser_etching|entry_of(c)=holography|X1=b|X1=c|X1=d)),\\ninference(sr,[status(thm)],[inference(er,[status(thm)],[inference(ef,[status(thm)],[c_0_56])]), c_0_57])).\n", "fillcolor": "palegreen"}, {"id": 64, "label": "64", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "fof_nnf", "raw_label": "tff(c_0_64, plain, c!=d,\\ninference(fof_nnf,[status(thm)],[c_0_58])).\n", "fillcolor": "palegreen"}, {"id": 65, "label": "65", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "fof_nnf", "raw_label": "tff(c_0_65, negated_conjecture, (entry_of(a)!=reconstituted_materials_sculpture|entry_of(b)!=holography|entry_of(c)!=laser_etching|entry_of(d)!=collage_painting),\\ninference(fof_nnf,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_59])])).\n", "fillcolor": "lightskyblue1"}, {"id": 66, "label": "66", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "split_conjunct", "raw_label": "tcf(c_0_66, plain, entry_of(b)!=laser_etching,\\ninference(split_conjunct,[status(thm)],[c_0_18])).\n", "fillcolor": "palegreen"}, {"id": 67, "label": "67", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "er", "raw_label": "tcf(c_0_67, plain, ![X1:knowheyan]:((entry_of(X1)=laser_etching|entry_of(X1)=holography|entry_of(X1)=collage_painting|a=X1)),\\ninference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_42])])).\n", "fillcolor": "palegreen"}, {"id": 68, "label": "68", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "csr", "raw_label": "tcf(c_0_68, plain, entry_of(b)!=collage_painting,\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_62])).\n", "fillcolor": "palegreen"}, {"id": 69, "label": "69", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "spm", "raw_label": "tcf(c_0_69, plain, ![X1:knowheyan, X4:knowheyan]:((entry_of(c)=holography|X1=d|X1=c|X1=b|c=X4|entry_of(X4)!=laser_etching)),\\ninference(spm,[status(thm)],[c_0_50, c_0_63])).\n", "fillcolor": "palegreen"}, {"id": 70, "label": "70", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "ef", "raw_label": "tcf(c_0_70, plain, ![X1:knowheyan]:((entry_of(d)=laser_etching|entry_of(d)=holography|entry_of(d)=collage_painting|X1=d|X1=b|X1=c)),\\ninference(ef,[status(thm)],[c_0_56])).\n", "fillcolor": "palegreen"}, {"id": 71, "label": "71", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "split_conjunct", "raw_label": "tcf(c_0_71, plain, c!=d,\\ninference(split_conjunct,[status(thm)],[c_0_64])).\n", "fillcolor": "palegreen"}, {"id": 72, "label": "72", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "fof_simplification", "raw_label": "tff(c_0_72, plain, b!=d,\\ninference(fof_simplification,[status(thm)],[b_not_d])).\n", "fillcolor": "palegreen"}, {"id": 73, "label": "73", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "split_conjunct", "raw_label": "tcf(c_0_73, negated_conjecture, (entry_of(a)!=reconstituted_materials_sculpture|entry_of(b)!=holography|entry_of(c)!=laser_etching|entry_of(d)!=collage_painting),\\ninference(split_conjunct,[status(thm)],[c_0_65])).\n", "fillcolor": "lightskyblue1"}, {"id": 74, "label": "74", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "sr", "raw_label": "tcf(c_0_74, plain, entry_of(b)=holography,\\ninference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_68]), c_0_46])).\n", "fillcolor": "palegreen"}, {"id": 75, "label": "75", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "sr", "raw_label": "tcf(c_0_75, plain, ![X1:knowheyan, X4:knowheyan]:((entry_of(d)=collage_painting|entry_of(d)=holography|entry_of(c)=holography|X1=c|X1=b|X1=d|X4=b|X4=c|X4=d)),\\ninference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_70]), c_0_71])).\n", "fillcolor": "palegreen"}, {"id": 76, "label": "76", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "fof_nnf", "raw_label": "tff(c_0_76, plain, b!=d,\\ninference(fof_nnf,[status(thm)],[c_0_72])).\n", "fillcolor": "palegreen"}, {"id": 77, "label": "77", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "sr", "raw_label": "tcf(c_0_77, negated_conjecture, (entry_of(b)!=holography|entry_of(c)!=laser_etching|entry_of(d)!=collage_painting),\\ninference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_23]), c_0_34]), c_0_33]), c_0_31])).\n", "fillcolor": "lightskyblue1"}, {"id": 78, "label": "78", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "spm", "raw_label": "tcf(c_0_78, plain, ![X1:knowheyan]:((b=X1|entry_of(X1)!=holography)),\\ninference(spm,[status(thm)],[c_0_50, c_0_74])).\n", "fillcolor": "palegreen"}, {"id": 79, "label": "79", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "ef", "raw_label": "tcf(c_0_79, plain, ![X1:knowheyan]:((entry_of(c)=holography|entry_of(d)=holography|entry_of(d)=collage_painting|X1=d|X1=b|X1=c)),\\ninference(ef,[status(thm)],[c_0_75])).\n", "fillcolor": "palegreen"}, {"id": 80, "label": "80", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "split_conjunct", "raw_label": "tcf(c_0_80, plain, b!=d,\\ninference(split_conjunct,[status(thm)],[c_0_76])).\n", "fillcolor": "palegreen"}, {"id": 81, "label": "81", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "fof_simplification", "raw_label": "tff(c_0_81, plain, b!=c,\\ninference(fof_simplification,[status(thm)],[b_not_c])).\n", "fillcolor": "palegreen"}, {"id": 82, "label": "82", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "cn", "raw_label": "tcf(c_0_82, negated_conjecture, (entry_of(c)!=laser_etching|entry_of(d)!=collage_painting),\\ninference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_77, c_0_74])])).\n", "fillcolor": "lightskyblue1"}, {"id": 83, "label": "83", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "sr", "raw_label": "tcf(c_0_83, plain, ![X1:knowheyan]:((entry_of(d)=collage_painting|entry_of(c)=holography|X1=c|X1=b|X1=d)),\\ninference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_79]), c_0_80])).\n", "fillcolor": "palegreen"}, {"id": 84, "label": "84", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "fof_nnf", "raw_label": "tff(c_0_84, plain, b!=c,\\ninference(fof_nnf,[status(thm)],[c_0_81])).\n", "fillcolor": "palegreen"}, {"id": 85, "label": "85", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "csr", "raw_label": "tcf(c_0_85, negated_conjecture, ![X1:knowheyan]:((entry_of(c)=holography|X1=d|X1=b|X1=c)),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_83]), c_0_63])).\n", "fillcolor": "lightskyblue1"}, {"id": 86, "label": "86", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "split_conjunct", "raw_label": "tcf(c_0_86, plain, b!=c,\\ninference(split_conjunct,[status(thm)],[c_0_84])).\n", "fillcolor": "palegreen"}, {"id": 87, "label": "87", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "sr", "raw_label": "tcf(c_0_87, negated_conjecture, ![X1:knowheyan]:((X1=c|X1=b|X1=d)),\\ninference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_85]), c_0_86])).\n", "fillcolor": "lightskyblue1"}, {"id": 88, "label": "c_0_88", "color": {"background": "#000000", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#FFFFFF"}, "cnf_name": "c_0_88", "role": "negated_conjecture", "formula": "$false", "inference_rule": "sr", "raw_label": "cnf(c_0_88, negated_conjecture, ($false),\\ninference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_87]), c_0_45]), c_0_46])).\n", "fillcolor": "darkorchid1"}], "edges": [{"from": 0, "to": 16, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 1, "to": 17, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 16, "to": 18, "label": "distribute", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 2, "to": 19, "label": "variable_rename", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 3, "to": 20, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 17, "to": 21, "label": "distribute", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 18, "to": 22, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 19, "to": 23, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 20, "to": 24, "label": "fof_nnf", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 21, "to": 25, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 4, "to": 26, "label": "variable_rename", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 5, "to": 27, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 6, "to": 28, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 7, "to": 29, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 8, "to": 30, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 22, "to": 31, "label": "er", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 23, "to": 32, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 23, "to": 32, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 24, "to": 33, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 25, "to": 34, "label": "er", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 26, "to": 35, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 27, "to": 36, "label": "fof_nnf", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 28, "to": 37, "label": "fof_nnf", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 29, "to": 38, "label": "fof_nnf", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 9, "to": 39, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 10, "to": 40, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 30, "to": 41, "label": "fof_nnf", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 34, "to": 42, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 33, "to": 42, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 32, "to": 42, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 31, "to": 42, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 35, "to": 43, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 35, "to": 43, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 36, "to": 44, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 37, "to": 45, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 38, "to": 46, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 39, "to": 47, "label": "fof_nnf", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 11, "to": 48, "label": "variable_rename", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 40, "to": 49, "label": "fof_nnf", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 41, "to": 50, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 46, "to": 51, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 45, "to": 51, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 44, "to": 51, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 43, "to": 51, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 42, "to": 51, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 21, "to": 52, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 47, "to": 53, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 48, "to": 54, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 49, "to": 55, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 51, "to": 56, "label": "er", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 50, "to": 56, "label": "er", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 52, "to": 57, "label": "er", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 12, "to": 58, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 13, "to": 59, "label": "assume_negation", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 21, "to": 60, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 55, "to": 61, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 54, "to": 61, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 53, "to": 61, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 21, "to": 62, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 57, "to": 63, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 56, "to": 63, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 58, "to": 64, "label": "fof_nnf", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 59, "to": 65, "label": "fof_nnf", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 18, "to": 66, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 42, "to": 67, "label": "er", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 50, "to": 67, "label": "er", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 62, "to": 68, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 61, "to": 68, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 60, "to": 68, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 63, "to": 69, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 50, "to": 69, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 56, "to": 70, "label": "ef", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 64, "to": 71, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 14, "to": 72, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 65, "to": 73, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 46, "to": 74, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 68, "to": 74, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 67, "to": 74, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 66, "to": 74, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 71, "to": 75, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 70, "to": 75, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 69, "to": 75, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 72, "to": 76, "label": "fof_nnf", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 31, "to": 77, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 33, "to": 77, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 34, "to": 77, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 23, "to": 77, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 73, "to": 77, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 74, "to": 78, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 50, "to": 78, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 75, "to": 79, "label": "ef", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 76, "to": 80, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 15, "to": 81, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 74, "to": 82, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 77, "to": 82, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 80, "to": 83, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 79, "to": 83, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 78, "to": 83, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 81, "to": 84, "label": "fof_nnf", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 63, "to": 85, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 83, "to": 85, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 82, "to": 85, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 84, "to": 86, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 86, "to": 87, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 85, "to": 87, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 78, "to": 87, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 46, "to": 88, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 45, "to": 88, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 87, "to": 88, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 44, "to": 88, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}]}, {"problem_file": "PUZ008-3.p", "status": "# SZS status Unsatisfiable", "nodes": [{"id": 0, "label": "two_cannibals_west", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "two_cannibals_west_to_east", "role": "axiom", "formula": "achievable(west(m(X1", "inference_rule": null, "raw_label": "cnf(two_cannibals_west_to_east, axiom, (achievable(west(m(X1),c(X2)),boatoneast,east(m(X3),c(s(s(X4)))))|~achievable(west(m(X1),c(s(s(X2)))),boatonwest,east(m(X3),c(X4)))|~safe(X1,X2)|~safe(X3,s(s(X4)))),\\nfile('Problems/PUZ/PUZ008-3.p', two_cannibals_west_to_east)).\n", "fillcolor": "forestgreen"}, {"id": 1, "label": "cannibal_east_to_w", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "cannibal_east_to_west", "role": "axiom", "formula": "achievable(west(m(X1", "inference_rule": null, "raw_label": "cnf(cannibal_east_to_west, axiom, (achievable(west(m(X1),c(s(X2))),boatonwest,east(m(X3),c(X4)))|~achievable(west(m(X1),c(X2)),boatoneast,east(m(X3),c(s(X4))))|~safe(X1,s(X2))|~safe(X3,X4)),\\nfile('Problems/PUZ/PUZ008-3.p', cannibal_east_to_west)).\n", "fillcolor": "forestgreen"}, {"id": 2, "label": "same_or_more_missi", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "same_or_more_missionaries_is_safe", "role": "axiom", "formula": "safe(X1,X2)|~greater_or_equal(X1,X2)", "inference_rule": null, "raw_label": "cnf(same_or_more_missionaries_is_safe, axiom, (safe(X1,X2)|~greater_or_equal(X1,X2)),\\nfile('Problems/PUZ/PUZ008-3.p', same_or_more_missionaries_is_safe)).\n", "fillcolor": "forestgreen"}, {"id": 3, "label": "two_missionaries_w", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "two_missionaries_west_to_east", "role": "axiom", "formula": "achievable(west(m(X1", "inference_rule": null, "raw_label": "cnf(two_missionaries_west_to_east, axiom, (achievable(west(m(X1),c(X2)),boatoneast,east(m(s(s(X3))),c(X4)))|~achievable(west(m(s(s(X1))),c(X2)),boatonwest,east(m(X3),c(X4)))|~safe(X1,X2)|~safe(s(s(X3)),X4)),\\nfile('Problems/PUZ/PUZ008-3.p', two_missionaries_west_to_east)).\n", "fillcolor": "forestgreen"}, {"id": 4, "label": "start_on_west_bank", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "start_on_west_bank", "role": "hypothesis", "formula": "achievable(west(m(s(s(s(n0)))", "inference_rule": null, "raw_label": "cnf(start_on_west_bank, hypothesis, (achievable(west(m(s(s(s(n0)))),c(s(s(s(n0))))),boatonwest,east(m(n0),c(n0)))),\\nfile('Problems/PUZ/PUZ008-3.p', start_on_west_bank)).\n", "fillcolor": "forestgreen"}, {"id": 5, "label": "no_missionaries_is", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "no_missionaries_is_safe", "role": "axiom", "formula": "safe(n0,X1)", "inference_rule": null, "raw_label": "cnf(no_missionaries_is_safe, axiom, (safe(n0,X1)),\\nfile('Problems/PUZ/PUZ008-3.p', no_missionaries_is_safe)).\n", "fillcolor": "forestgreen"}, {"id": 6, "label": "great_eq_base", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "great_eq_base", "role": "axiom", "formula": "greater_or_equal(X1,n0)", "inference_rule": null, "raw_label": "cnf(great_eq_base, axiom, (greater_or_equal(X1,n0)),\\nfile('Problems/PUZ/PUZ008-3.p', great_eq_base)).\n", "fillcolor": "forestgreen"}, {"id": 7, "label": "missionary_and_can", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "missionary_and_cannibal_east_to_west", "role": "axiom", "formula": "achievable(west(m(s(X1)", "inference_rule": null, "raw_label": "cnf(missionary_and_cannibal_east_to_west, axiom, (achievable(west(m(s(X1)),c(s(X2))),boatonwest,east(m(X3),c(X4)))|~achievable(west(m(X1),c(X2)),boatoneast,east(m(s(X3)),c(s(X4))))|~safe(s(X1),s(X2))|~safe(X3,X4)),\\nfile('Problems/PUZ/PUZ008-3.p', missionary_and_cannibal_east_to_west)).\n", "fillcolor": "forestgreen"}, {"id": 8, "label": "missionary_and_can", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "missionary_and_cannibal_west_to_east", "role": "axiom", "formula": "achievable(west(m(X1", "inference_rule": null, "raw_label": "cnf(missionary_and_cannibal_west_to_east, axiom, (achievable(west(m(X1),c(X2)),boatoneast,east(m(s(X3)),c(s(X4))))|~achievable(west(m(s(X1)),c(s(X2))),boatonwest,east(m(X3),c(X4)))|~safe(X1,X2)|~safe(s(X3),s(X4))),\\nfile('Problems/PUZ/PUZ008-3.p', missionary_and_cannibal_west_to_east)).\n", "fillcolor": "forestgreen"}, {"id": 9, "label": "missionary_east_to", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "missionary_east_to_west", "role": "axiom", "formula": "achievable(west(m(s(X1)", "inference_rule": null, "raw_label": "cnf(missionary_east_to_west, axiom, (achievable(west(m(s(X1)),c(X2)),boatonwest,east(m(X3),c(X4)))|~achievable(west(m(X1),c(X2)),boatoneast,east(m(s(X3)),c(X4)))|~safe(s(X1),X2)|~safe(X3,X4)),\\nfile('Problems/PUZ/PUZ008-3.p', missionary_east_to_west)).\n", "fillcolor": "forestgreen"}, {"id": 10, "label": "prove_can_get_to_e", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "prove_can_get_to_east_bank", "role": "negated_conjecture", "formula": "~achievable(west(m(n0", "inference_rule": null, "raw_label": "cnf(prove_can_get_to_east_bank, negated_conjecture, (~achievable(west(m(n0),c(n0)),X1,east(m(s(s(s(n0)))),c(s(s(s(n0))))))),\\nfile('Problems/PUZ/PUZ008-3.p', prove_can_get_to_east_bank)).\n", "fillcolor": "dodgerblue"}, {"id": 11, "label": "great_eq_recursive", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": "great_eq_recursive", "role": "axiom", "formula": "greater_or_equal(s(X1", "inference_rule": null, "raw_label": "cnf(great_eq_recursive, axiom, (greater_or_equal(s(X1),s(X2))|~greater_or_equal(X1,X2)),\\nfile('Problems/PUZ/PUZ008-3.p', great_eq_recursive)).\n", "fillcolor": "forestgreen"}, {"id": 12, "label": "c_0_12", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_12", "role": "plain", "formula": "achievable(west(m(X1", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_12, plain, (achievable(west(m(X1),c(X2)),boatoneast,east(m(X3),c(s(s(X4)))))|~achievable(west(m(X1),c(s(s(X2)))),boatonwest,east(m(X3),c(X4)))|~safe(X1,X2)|~safe(X3,s(s(X4)))),\\ninference(fof_simplification,[status(thm)],[two_cannibals_west_to_east])).\n", "fillcolor": "palegreen"}, {"id": 13, "label": "c_0_13", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_13", "role": "plain", "formula": "achievable(west(m(X1", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_13, plain, (achievable(west(m(X1),c(s(X2))),boatonwest,east(m(X3),c(X4)))|~achievable(west(m(X1),c(X2)),boatoneast,east(m(X3),c(s(X4))))|~safe(X1,s(X2))|~safe(X3,X4)),\\ninference(fof_simplification,[status(thm)],[cannibal_east_to_west])).\n", "fillcolor": "palegreen"}, {"id": 14, "label": "c_0_14", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_14", "role": "plain", "formula": "safe(X1,X2)|~greater_or_equal(X1,X2)", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_14, plain, (safe(X1,X2)|~greater_or_equal(X1,X2)),\\ninference(fof_simplification,[status(thm)],[same_or_more_missionaries_is_safe])).\n", "fillcolor": "palegreen"}, {"id": 15, "label": "c_0_15", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_15", "role": "plain", "formula": "achievable(west(m(X1", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_15, plain, (achievable(west(m(X1),c(X2)),boatoneast,east(m(s(s(X3))),c(X4)))|~achievable(west(m(s(s(X1))),c(X2)),boatonwest,east(m(X3),c(X4)))|~safe(X1,X2)|~safe(s(s(X3)),X4)),\\ninference(fof_simplification,[status(thm)],[two_missionaries_west_to_east])).\n", "fillcolor": "palegreen"}, {"id": 16, "label": "c_0_16", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_16", "role": "plain", "formula": "achievable(west(m(X1", "inference_rule": null, "raw_label": "cnf(c_0_16, plain, (achievable(west(m(X1),c(X2)),boatoneast,east(m(X3),c(s(s(X4)))))|~achievable(west(m(X1),c(s(s(X2)))),boatonwest,east(m(X3),c(X4)))|~safe(X1,X2)|~safe(X3,s(s(X4)))),\\nc_0_12).\n", "fillcolor": "palegreen"}, {"id": 17, "label": "c_0_17", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_17", "role": "plain", "formula": "achievable(west(m(X1", "inference_rule": null, "raw_label": "cnf(c_0_17, plain, (achievable(west(m(X1),c(s(X2))),boatonwest,east(m(X3),c(X4)))|~achievable(west(m(X1),c(X2)),boatoneast,east(m(X3),c(s(X4))))|~safe(X1,s(X2))|~safe(X3,X4)),\\nc_0_13).\n", "fillcolor": "palegreen"}, {"id": 18, "label": "c_0_18", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_18", "role": "hypothesis", "formula": "achievable(west(m(s(s(s(n0)))", "inference_rule": null, "raw_label": "cnf(c_0_18, hypothesis, (achievable(west(m(s(s(s(n0)))),c(s(s(s(n0))))),boatonwest,east(m(n0),c(n0)))),\\nstart_on_west_bank).\n", "fillcolor": "palegreen"}, {"id": 19, "label": "c_0_19", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_19", "role": "axiom", "formula": "safe(n0,X1)", "inference_rule": null, "raw_label": "cnf(c_0_19, axiom, (safe(n0,X1)),\\nno_missionaries_is_safe).\n", "fillcolor": "palegreen"}, {"id": 20, "label": "c_0_20", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_20", "role": "plain", "formula": "safe(X1,X2)|~greater_or_equal(X1,X2)", "inference_rule": null, "raw_label": "cnf(c_0_20, plain, (safe(X1,X2)|~greater_or_equal(X1,X2)),\\nc_0_14).\n", "fillcolor": "palegreen"}, {"id": 21, "label": "c_0_21", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_21", "role": "axiom", "formula": "greater_or_equal(X1,n0)", "inference_rule": null, "raw_label": "cnf(c_0_21, axiom, (greater_or_equal(X1,n0)),\\ngreat_eq_base).\n", "fillcolor": "palegreen"}, {"id": 22, "label": "c_0_22", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_22", "role": "plain", "formula": "achievable(west(m(s(X1)", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_22, plain, (achievable(west(m(s(X1)),c(s(X2))),boatonwest,east(m(X3),c(X4)))|~achievable(west(m(X1),c(X2)),boatoneast,east(m(s(X3)),c(s(X4))))|~safe(s(X1),s(X2))|~safe(X3,X4)),\\ninference(fof_simplification,[status(thm)],[missionary_and_cannibal_east_to_west])).\n", "fillcolor": "palegreen"}, {"id": 23, "label": "c_0_23", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_23", "role": "plain", "formula": "achievable(west(m(X1", "inference_rule": null, "raw_label": "cnf(c_0_23, plain, (achievable(west(m(X1),c(X2)),boatoneast,east(m(s(s(X3))),c(X4)))|~achievable(west(m(s(s(X1))),c(X2)),boatonwest,east(m(X3),c(X4)))|~safe(X1,X2)|~safe(s(s(X3)),X4)),\\nc_0_15).\n", "fillcolor": "palegreen"}, {"id": 24, "label": "c_0_24", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_24", "role": "plain", "formula": "achievable(west(m(X1", "inference_rule": "spm", "raw_label": "cnf(c_0_24, plain, (achievable(west(m(X1),c(X2)),boatoneast,east(m(X3),c(s(s(X4)))))|~achievable(west(m(X1),c(s(X2))),boatoneast,east(m(X3),c(s(X4))))|~safe(X3,s(s(X4)))|~safe(X1,s(s(X2)))|~safe(X1,X2)|~safe(X3,X4)),\\ninference(spm,[status(thm)],[c_0_16, c_0_17])).\n", "fillcolor": "palegreen"}, {"id": 25, "label": "c_0_25", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_25", "role": "hypothesis", "formula": "achievable(west(m(s(s(s(n0)))", "inference_rule": "cn", "raw_label": "cnf(c_0_25, hypothesis, (achievable(west(m(s(s(s(n0)))),c(s(n0))),boatoneast,east(m(n0),c(s(s(n0)))))|~safe(s(s(s(n0))),s(n0))),\\ninference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_18]), c_0_19])])).\n", "fillcolor": "palegreen"}, {"id": 26, "label": "c_0_26", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_26", "role": "plain", "formula": "safe(X1,n0)", "inference_rule": "spm", "raw_label": "cnf(c_0_26, plain, (safe(X1,n0)),\\ninference(spm,[status(thm)],[c_0_20, c_0_21])).\n", "fillcolor": "palegreen"}, {"id": 27, "label": "c_0_27", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_27", "role": "plain", "formula": "achievable(west(m(s(X1)", "inference_rule": null, "raw_label": "cnf(c_0_27, plain, (achievable(west(m(s(X1)),c(s(X2))),boatonwest,east(m(X3),c(X4)))|~achievable(west(m(X1),c(X2)),boatoneast,east(m(s(X3)),c(s(X4))))|~safe(s(X1),s(X2))|~safe(X3,X4)),\\nc_0_22).\n", "fillcolor": "palegreen"}, {"id": 28, "label": "c_0_28", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_28", "role": "plain", "formula": "achievable(west(m(X1", "inference_rule": "spm", "raw_label": "cnf(c_0_28, plain, (achievable(west(m(X1),c(s(X2))),boatoneast,east(m(s(s(X3))),c(X4)))|~achievable(west(m(s(s(X1))),c(X2)),boatoneast,east(m(X3),c(s(X4))))|~safe(s(s(X1)),s(X2))|~safe(s(s(X3)),X4)|~safe(X1,s(X2))|~safe(X3,X4)),\\ninference(spm,[status(thm)],[c_0_23, c_0_17])).\n", "fillcolor": "palegreen"}, {"id": 29, "label": "c_0_29", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_29", "role": "hypothesis", "formula": "achievable(west(m(s(s(s(n0)))", "inference_rule": "cn", "raw_label": "cnf(c_0_29, hypothesis, (achievable(west(m(s(s(s(n0)))),c(n0)),boatoneast,east(m(n0),c(s(s(s(n0))))))|~safe(s(s(s(n0))),s(s(n0)))|~safe(s(s(s(n0))),s(n0))),\\ninference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_19]), c_0_26]), c_0_19])])).\n", "fillcolor": "palegreen"}, {"id": 30, "label": "c_0_30", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_30", "role": "plain", "formula": "achievable(west(m(X1", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_30, plain, (achievable(west(m(X1),c(X2)),boatoneast,east(m(s(X3)),c(s(X4))))|~achievable(west(m(s(X1)),c(s(X2))),boatonwest,east(m(X3),c(X4)))|~safe(X1,X2)|~safe(s(X3),s(X4))),\\ninference(fof_simplification,[status(thm)],[missionary_and_cannibal_west_to_east])).\n", "fillcolor": "palegreen"}, {"id": 31, "label": "c_0_31", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_31", "role": "plain", "formula": "achievable(west(m(s(X1)", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_31, plain, (achievable(west(m(s(X1)),c(X2)),boatonwest,east(m(X3),c(X4)))|~achievable(west(m(X1),c(X2)),boatoneast,east(m(s(X3)),c(X4)))|~safe(s(X1),X2)|~safe(X3,X4)),\\ninference(fof_simplification,[status(thm)],[missionary_east_to_west])).\n", "fillcolor": "palegreen"}, {"id": 32, "label": "c_0_32", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_32", "role": "plain", "formula": "achievable(west(m(X1", "inference_rule": "spm", "raw_label": "cnf(c_0_32, plain, (achievable(west(m(X1),c(s(X2))),boatoneast,east(m(s(s(X3))),c(X4)))|~achievable(west(m(s(X1)),c(X2)),boatoneast,east(m(s(X3)),c(s(X4))))|~safe(s(s(X1)),s(X2))|~safe(s(s(X3)),X4)|~safe(X1,s(X2))|~safe(X3,X4)),\\ninference(spm,[status(thm)],[c_0_23, c_0_27])).\n", "fillcolor": "palegreen"}, {"id": 33, "label": "c_0_33", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_33", "role": "hypothesis", "formula": "achievable(west(m(s(n0)", "inference_rule": "cn", "raw_label": "cnf(c_0_33, hypothesis, (achievable(west(m(s(n0)),c(s(n0))),boatoneast,east(m(s(s(n0))),c(s(s(n0)))))|~safe(s(s(s(n0))),s(s(n0)))|~safe(s(s(s(n0))),s(n0))|~safe(s(s(n0)),s(s(n0)))|~safe(s(n0),s(n0))),\\ninference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_19])])).\n", "fillcolor": "palegreen"}, {"id": 34, "label": "c_0_34", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_34", "role": "plain", "formula": "achievable(west(m(X1", "inference_rule": null, "raw_label": "cnf(c_0_34, plain, (achievable(west(m(X1),c(X2)),boatoneast,east(m(s(X3)),c(s(X4))))|~achievable(west(m(s(X1)),c(s(X2))),boatonwest,east(m(X3),c(X4)))|~safe(X1,X2)|~safe(s(X3),s(X4))),\\nc_0_30).\n", "fillcolor": "palegreen"}, {"id": 35, "label": "c_0_35", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_35", "role": "plain", "formula": "achievable(west(m(s(X1)", "inference_rule": null, "raw_label": "cnf(c_0_35, plain, (achievable(west(m(s(X1)),c(X2)),boatonwest,east(m(X3),c(X4)))|~achievable(west(m(X1),c(X2)),boatoneast,east(m(s(X3)),c(X4)))|~safe(s(X1),X2)|~safe(X3,X4)),\\nc_0_31).\n", "fillcolor": "palegreen"}, {"id": 36, "label": "c_0_36", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_36", "role": "hypothesis", "formula": "achievable(west(m(n0", "inference_rule": "cn", "raw_label": "cnf(c_0_36, hypothesis, (achievable(west(m(n0),c(s(s(n0)))),boatoneast,east(m(s(s(s(n0)))),c(s(n0))))|~safe(s(s(s(n0))),s(s(n0)))|~safe(s(s(n0)),s(s(n0)))|~safe(s(s(s(n0))),s(n0))|~safe(s(n0),s(n0))),\\ninference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_19])])).\n", "fillcolor": "palegreen"}, {"id": 37, "label": "c_0_37", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_37", "role": "negated_conjecture", "formula": "~achievable(west(m(n0", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_37, negated_conjecture, (~achievable(west(m(n0),c(n0)),X1,east(m(s(s(s(n0)))),c(s(s(s(n0))))))),\\ninference(fof_simplification,[status(thm)],[prove_can_get_to_east_bank])).\n", "fillcolor": "lightskyblue1"}, {"id": 38, "label": "c_0_38", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_38", "role": "plain", "formula": "greater_or_equal(s(X1", "inference_rule": "fof_simplification", "raw_label": "cnf(c_0_38, plain, (greater_or_equal(s(X1),s(X2))|~greater_or_equal(X1,X2)),\\ninference(fof_simplification,[status(thm)],[great_eq_recursive])).\n", "fillcolor": "palegreen"}, {"id": 39, "label": "c_0_39", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_39", "role": "plain", "formula": "achievable(west(m(X1", "inference_rule": "spm", "raw_label": "cnf(c_0_39, plain, (achievable(west(m(X1),c(X2)),boatoneast,east(m(s(X3)),c(s(X4))))|~achievable(west(m(X1),c(s(X2))),boatoneast,east(m(s(X3)),c(X4)))|~safe(s(X3),s(X4))|~safe(s(X1),s(X2))|~safe(X1,X2)|~safe(X3,X4)),\\ninference(spm,[status(thm)],[c_0_34, c_0_35])).\n", "fillcolor": "palegreen"}, {"id": 40, "label": "c_0_40", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_40", "role": "hypothesis", "formula": "achievable(west(m(n0", "inference_rule": "cn", "raw_label": "cnf(c_0_40, hypothesis, (achievable(west(m(n0),c(s(n0))),boatoneast,east(m(s(s(s(n0)))),c(s(s(n0)))))|~safe(s(s(s(n0))),s(s(n0)))|~safe(s(s(n0)),s(s(n0)))|~safe(s(s(s(n0))),s(n0))|~safe(s(n0),s(n0))),\\ninference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_36]), c_0_19]), c_0_19]), c_0_26])])).\n", "fillcolor": "palegreen"}, {"id": 41, "label": "c_0_41", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_41", "role": "negated_conjecture", "formula": "~achievable(west(m(n0", "inference_rule": null, "raw_label": "cnf(c_0_41, negated_conjecture, (~achievable(west(m(n0),c(n0)),X1,east(m(s(s(s(n0)))),c(s(s(s(n0))))))),\\nc_0_37).\n", "fillcolor": "lightskyblue1"}, {"id": 42, "label": "c_0_42", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_42", "role": "plain", "formula": "greater_or_equal(s(X1", "inference_rule": null, "raw_label": "cnf(c_0_42, plain, (greater_or_equal(s(X1),s(X2))|~greater_or_equal(X1,X2)),\\nc_0_38).\n", "fillcolor": "palegreen"}, {"id": 43, "label": "c_0_43", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_43", "role": "hypothesis", "formula": "~safe(s(s(s(n0))", "inference_rule": "sr", "raw_label": "cnf(c_0_43, hypothesis, (~safe(s(s(s(n0))),s(s(s(n0))))|~safe(s(s(s(n0))),s(s(n0)))|~safe(s(s(n0)),s(s(n0)))|~safe(s(s(s(n0))),s(n0))|~safe(s(n0),s(n0))),\\ninference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_26])]), c_0_41])).\n", "fillcolor": "palegreen"}, {"id": 44, "label": "c_0_44", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_44", "role": "plain", "formula": "safe(s(X1", "inference_rule": "spm", "raw_label": "cnf(c_0_44, plain, (safe(s(X1),s(X2))|~greater_or_equal(X1,X2)),\\ninference(spm,[status(thm)],[c_0_20, c_0_42])).\n", "fillcolor": "palegreen"}, {"id": 45, "label": "c_0_45", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_45", "role": "hypothesis", "formula": "~greater_or_equal(s(s(n0)", "inference_rule": "csr", "raw_label": "cnf(c_0_45, hypothesis, (~greater_or_equal(s(s(n0)),s(s(n0)))|~safe(s(s(s(n0))),s(s(n0)))|~safe(s(s(s(n0))),s(n0))|~safe(s(n0),s(n0))),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_20])).\n", "fillcolor": "palegreen"}, {"id": 46, "label": "c_0_46", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_46", "role": "hypothesis", "formula": "~greater_or_equal(s(s(n0)", "inference_rule": "spm", "raw_label": "cnf(c_0_46, hypothesis, (~greater_or_equal(s(s(n0)),s(s(n0)))|~greater_or_equal(s(s(n0)),s(n0))|~safe(s(s(s(n0))),s(n0))|~safe(s(n0),s(n0))),\\ninference(spm,[status(thm)],[c_0_45, c_0_44])).\n", "fillcolor": "palegreen"}, {"id": 47, "label": "c_0_47", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_47", "role": "hypothesis", "formula": "~greater_or_equal(s(s(n0)", "inference_rule": "csr", "raw_label": "cnf(c_0_47, hypothesis, (~greater_or_equal(s(s(n0)),s(n0))|~greater_or_equal(s(n0),s(n0))|~safe(s(s(s(n0))),s(n0))),\\ninference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_42]), c_0_20])).\n", "fillcolor": "palegreen"}, {"id": 48, "label": "c_0_48", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_48", "role": "hypothesis", "formula": "~greater_or_equal(s(s(n0)", "inference_rule": "cn", "raw_label": "cnf(c_0_48, hypothesis, (~greater_or_equal(s(s(n0)),s(n0))|~greater_or_equal(s(n0),s(n0))),\\ninference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_44]), c_0_21])])).\n", "fillcolor": "palegreen"}, {"id": 49, "label": "c_0_49", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": "c_0_49", "role": "hypothesis", "formula": "~greater_or_equal(s(n0", "inference_rule": "cn", "raw_label": "cnf(c_0_49, hypothesis, (~greater_or_equal(s(n0),s(n0))),\\ninference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_42]), c_0_21])])).\n", "fillcolor": "palegreen"}, {"id": 50, "label": "c_0_50", "color": {"background": "#000000", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#FFFFFF"}, "cnf_name": "c_0_50", "role": "hypothesis", "formula": "$false", "inference_rule": "cn", "raw_label": "cnf(c_0_50, hypothesis, ($false),\\ninference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_42]), c_0_21])])).\n", "fillcolor": "darkorchid1"}], "edges": [{"from": 0, "to": 12, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 1, "to": 13, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 2, "to": 14, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 3, "to": 15, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 12, "to": 16, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 13, "to": 17, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 4, "to": 18, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 5, "to": 19, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 14, "to": 20, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 6, "to": 21, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 7, "to": 22, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 15, "to": 23, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 17, "to": 24, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 16, "to": 24, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 19, "to": 25, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 18, "to": 25, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 16, "to": 25, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 21, "to": 26, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 20, "to": 26, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 22, "to": 27, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 17, "to": 28, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 23, "to": 28, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 19, "to": 29, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 26, "to": 29, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 19, "to": 29, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 25, "to": 29, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 24, "to": 29, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 8, "to": 30, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 9, "to": 31, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 27, "to": 32, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 23, "to": 32, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 19, "to": 33, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 29, "to": 33, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 28, "to": 33, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 30, "to": 34, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 31, "to": 35, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 19, "to": 36, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 33, "to": 36, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 32, "to": 36, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 10, "to": 37, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 11, "to": 38, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 35, "to": 39, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 34, "to": 39, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 26, "to": 40, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 19, "to": 40, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 19, "to": 40, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 36, "to": 40, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 24, "to": 40, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 37, "to": 41, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 38, "to": 42, "label": "", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 41, "to": 43, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 26, "to": 43, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 40, "to": 43, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 39, "to": 43, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 42, "to": 44, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 20, "to": 44, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 20, "to": 45, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 44, "to": 45, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 43, "to": 45, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 44, "to": 46, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 45, "to": 46, "label": "spm", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 20, "to": 47, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 42, "to": 47, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 46, "to": 47, "label": "csr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 21, "to": 48, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 44, "to": 48, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 47, "to": 48, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 21, "to": 49, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 42, "to": 49, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 48, "to": 49, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 21, "to": 50, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 42, "to": 50, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 49, "to": 50, "label": "cn", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}]}, {"problem_file": "PUZ108^5.p", "status": "# SZS status Theorem", "nodes": [{"id": 0, "label": "0", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": null, "raw_label": "thf(cCKB_E2_def, axiom, ((cCKB_E2)=(^[X1:$i, X2:$i]:(![X3:$i > $o]:((((X3 @ X1)&![X4:$i]:(((X3 @ X4)=>(X3 @ (s @ (s @ X4))))))=>(X3 @ X2)))))),\\nfile('Problems/PUZ/PUZ108^5.p', cCKB_E2_def)).\n", "fillcolor": "forestgreen"}, {"id": 1, "label": "1", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": null, "raw_label": "thf(cCKB_L35000, conjecture, ![X1:$i, X2:$i]:(((cCKB_E2 @ X1 @ X2)=>(cCKB_E2 @ (s @ X1) @ (s @ X2)))),\\nfile('Problems/PUZ/PUZ108^5.p', cCKB_L35000)).\n", "fillcolor": "firebrick1"}, {"id": 2, "label": "2", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "fof_simplification", "raw_label": "thf(c_0_2, plain, ((cCKB_E2)=(^[Z0/* 19 */:$i, Z1:$i]:(![X3:$i > $o]:((((X3 @ Z0)&![X4:$i]:(((X3 @ X4)=>(X3 @ (s @ (s @ X4))))))=>(X3 @ Z1)))))),\\ninference(fof_simplification,[status(thm)],[cCKB_E2_def])).\n", "fillcolor": "palegreen"}, {"id": 3, "label": "3", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "apply_def", "raw_label": "thf(c_0_3, negated_conjecture, ~(![X1:$i, X2:$i]:((![X7:$i > $o]:((((X7 @ X1)&![X8:$i]:(((X7 @ X8)=>(X7 @ (s @ (s @ X8))))))=>(X7 @ X2)))=>![X9:$i > $o]:((((X9 @ (s @ X1))&![X10:$i]:(((X9 @ X10)=>(X9 @ (s @ (s @ X10))))))=>(X9 @ (s @ X2))))))),\\ninference(apply_def,[status(thm)],[inference(assume_negation,[status(cth)],[cCKB_L35000]), c_0_2])).\n", "fillcolor": "lightskyblue1"}, {"id": 4, "label": "4", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "distribute", "raw_label": "thf(c_0_4, negated_conjecture, ![X13:$i > $o, X16:$i]:(((((X13 @ (esk3_1 @ X13))|~(X13 @ esk1_0)|(X13 @ esk2_0))&(~(X13 @ (s @ (s @ (esk3_1 @ X13))))|~(X13 @ esk1_0)|(X13 @ esk2_0)))&(((epred1_0 @ (s @ esk1_0))&(~(epred1_0 @ X16)|(epred1_0 @ (s @ (s @ X16)))))&~(epred1_0 @ (s @ esk2_0))))),\\ninference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])])])).\n", "fillcolor": "lightskyblue1"}, {"id": 5, "label": "5", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "split_conjunct", "raw_label": "thf(c_0_5, negated_conjecture, ![X3:$i > $o]:(((X3 @ esk2_0)|~((X3 @ (s @ (s @ (esk3_1 @ X3)))))|~((X3 @ esk1_0)))),\\ninference(split_conjunct,[status(thm)],[c_0_4])).\n", "fillcolor": "lightskyblue1"}, {"id": 6, "label": "6", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "split_conjunct", "raw_label": "thf(c_0_6, negated_conjecture, ![X1:$i]:(((epred1_0 @ (s @ (s @ X1)))|~((epred1_0 @ X1)))),\\ninference(split_conjunct,[status(thm)],[c_0_4])).\n", "fillcolor": "lightskyblue1"}, {"id": 7, "label": "7", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "split_conjunct", "raw_label": "thf(c_0_7, negated_conjecture, (epred1_0 @ (s @ esk1_0)),\\ninference(split_conjunct,[status(thm)],[c_0_4])).\n", "fillcolor": "lightskyblue1"}, {"id": 8, "label": "8", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "split_conjunct", "raw_label": "thf(c_0_8, negated_conjecture, ~((epred1_0 @ (s @ esk2_0))),\\ninference(split_conjunct,[status(thm)],[c_0_4])).\n", "fillcolor": "lightskyblue1"}, {"id": 9, "label": "9", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "split_conjunct", "raw_label": "thf(c_0_9, negated_conjecture, ![X3:$i > $o]:(((X3 @ (esk3_1 @ X3))|(X3 @ esk2_0)|~((X3 @ esk1_0)))),\\ninference(split_conjunct,[status(thm)],[c_0_4])).\n", "fillcolor": "lightskyblue1"}, {"id": 10, "label": "10", "color": {"background": "#FFFFFF", "border": "#000000"}, "borderWidth": 2, "font": {"color": "#000000"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "sr", "raw_label": "thf(c_0_10, negated_conjecture, ~((epred1_0 @ (s @ (esk3_1 @ (^[Z0/* 3 */:$i]:((epred1_0 @ (s @ Z0)))))))),\\ninference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_5, c_0_6]), c_0_7])]), c_0_8])).\n", "fillcolor": "lightskyblue1"}, {"id": 11, "label": "11", "color": {"background": "#000000", "border": "#000000"}, "borderWidth": 4, "font": {"color": "#FFFFFF"}, "cnf_name": null, "role": null, "formula": null, "inference_rule": "sr", "raw_label": "thf(c_0_11, negated_conjecture, ($false),\\ninference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_7]), c_0_10]), c_0_8])).\n", "fillcolor": "darkorchid1"}], "edges": [{"from": 0, "to": 2, "label": "fof_simplification", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 2, "to": 3, "label": "apply_def", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 1, "to": 3, "label": "apply_def", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 3, "to": 4, "label": "distribute", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 4, "to": 5, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 4, "to": 6, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 4, "to": 7, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 4, "to": 8, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 4, "to": 9, "label": "split_conjunct", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 8, "to": 10, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 7, "to": 10, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 6, "to": 10, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 5, "to": 10, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 8, "to": 11, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 10, "to": 11, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 7, "to": 11, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}, {"from": 9, "to": 11, "label": "sr", "arrows": "to", "color": {"color": "#000000"}, "width": 1.5, "smooth": {"type": "straight"}}]}];
        const graphsPerPage = 25;
        let currentPage = 0;
        const totalPages = 1;
        const networks = Array(25).fill(null);

        function renderPage(page) {
            const startIdx = page * graphsPerPage;
            const endIdx = Math.min(startIdx + graphsPerPage, allGraphsData.length);

            for (let i = 0; i < 25; i++) {
                const labelEl = document.getElementById('label-' + i);
                const networkEl = document.getElementById('network-' + i);
                if (networkEl) networkEl.innerHTML = '';
                if (labelEl) labelEl.textContent = '';
                if (networks[i]) {
                    networks[i].destroy();
                    networks[i] = null;
                }
            }

            for (let i = 0; i < endIdx - startIdx; i++) {
                const graphIdx = startIdx + i;
                const graphData = allGraphsData[graphIdx];
                if (!graphData) continue;

                const labelEl = document.getElementById('label-' + i);
                const networkEl = document.getElementById('network-' + i);
                if (!labelEl || !networkEl) continue;

                const name = graphData.problem_file.split('/').pop() || graphData.problem_file;
                const stats = 'N' + graphData.nodes.length + ' E' + graphData.edges.length;
                labelEl.textContent = name + ' | ' + stats;

                const nodes = new vis.DataSet(graphData.nodes);
                const edges = new vis.DataSet(graphData.edges);

                const data = { nodes: nodes, edges: edges };
                const options = {
                    nodes: {
                        shape: 'box',
                        font: { size: 8, face: 'monospace' },
                        margin: 4,
                        widthConstraint: { maximum: 90 },
                        heightConstraint: { maximum: 30 }
                    },
                    edges: {
                        font: { size: 6, align: 'middle' },
                        smooth: { type: 'straight' },
                        arrows: { to: { enabled: true, scaleFactor: 0.5 } },
                        width: 1.5
                    },
                    layout: {
                        hierarchical: {
                            enabled: true,
                            direction: 'LR',
                            sortMethod: 'directed',
                            levelSeparation: 50,
                            nodeSpacing: 80,
                            treeSpacing: 100
                        }
                    },
                    physics: { enabled: false },
                    interaction: {
                        dragNodes: true,
                        dragView: true,
                        zoomView: true,
                        hover: false,
                        tooltipDelay: 0
                    }
                };

                networks[i] = new vis.Network(networkEl, data, options);

                const originalNodeFontSize = options.nodes.font.size;
                const originalEdgeFontSize = options.edges.font.size;

                networks[i].on('click', function(params) {
                    const detailBox = document.getElementById('node-detail-box');

                    if (params.nodes.length > 0) {
                        const nodeId = params.nodes[0];
                        const nodeData = nodes.get(nodeId);

                        const newNodeFontSize = Math.max(4, Math.floor(originalNodeFontSize * 0.7));
                        const newEdgeFontSize = Math.max(4, Math.floor(originalEdgeFontSize * 0.7));

                        nodes.forEach(function(node) {
                            nodes.update({
                                id: node.id,
                                font: { size: newNodeFontSize, face: 'monospace' }
                            });
                        });

                        edges.forEach(function(edge) {
                            edges.update({
                                id: edge.id,
                                font: { size: newEdgeFontSize, align: 'middle' }
                            });
                        });

                        const lines = [];
                        lines.push('Problem: ' + graphData.problem_file);
                        lines.push('Node: ' + nodeData.label + ' (id=' + nodeId + ')');
                        if (nodeData.cnf_name) lines.push('cnf_name: ' + nodeData.cnf_name);
                        if (nodeData.role) lines.push('role: ' + nodeData.role);
                        if (nodeData.inference_rule) lines.push('rule: ' + nodeData.inference_rule);
                        if (nodeData.formula) lines.push('formula: ' + nodeData.formula);

                        detailBox.textContent = lines.join('\n');
                        detailBox.classList.add('active');

                        setTimeout(function() {
                            detailBox.classList.remove('active');
                            nodes.forEach(function(node) {
                                nodes.update({
                                    id: node.id,
                                    font: { size: originalNodeFontSize, face: 'monospace' }
                                });
                            });
                            edges.forEach(function(edge) {
                                edges.update({
                                    id: edge.id,
                                    font: { size: originalEdgeFontSize, align: 'middle' }
                                });
                            });
                        }, 3500);
                    } else {
                        detailBox.classList.remove('active');
                    }
                });
            }

            document.getElementById('page-info').textContent = 'Page ' + (page + 1) + ' / ' + totalPages;
            document.getElementById('prev-btn').disabled = (page === 0);
            document.getElementById('next-btn').disabled = (page >= totalPages - 1);
        }

        function changePage(delta) {
            const newPage = currentPage + delta;
            if (newPage >= 0 && newPage < totalPages) {
                currentPage = newPage;
                renderPage(currentPage);
            }
        }

        renderPage(0);
    </script>
</body>
</html>